<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25735 -->
  <question type="stack">
    <name>
				<text>2.4.6.1 Classify odd, even, or neither</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Classify the following functions as odd, even or neither. If necessary sketch a graph to help you decide.</p><p><br></p>
<ol class="HELM_parts">
<li><p>\(f(x)=6\) is [[1]]</p></li>
<li><p>\(f(x)=x^2\) is [[1]]</p></li>
<li><p>\(f(x)=2x+1\) is [[3]]</p></li>
<li><p>\(f(x)=x\) is [[2]]</p></li>
<li><p>\(f(x)=2x\) is [[2]]</p></li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200618185538+n5fD63</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200812172645+pkcUXY</text></tag>
    </tags>
    
</question>
</quiz>