<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27289 -->
  <question type="stack">
    <name>
				<text>36.1.5.2 Measures of location and Mean</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4><span>Measures of location</span></h4>
<p>There are three widely used measures of location, these are:</p>
<ul><li>The Mean, the arithmetic average of the data;</li>
<li>The Median, the central value of the data;</li>
<li>The Mode, the most frequently occurring value in the data set.</li>
</ul>
<h4><span>Mean</span></h4>
<p>This section of the booklet will show you how to calculate the mean. Strictly, we should say that this is the <i>arithmetic mean</i>. There are other kinds of means as well. This is by far the most common, though, and so we often just say "mean''.</p>
<div class="HELM_keypoint">
<h4>Key Point</h4>
<p>If we take a set of numbers \(x_1,x_2,\dots , x_n\), its arithmetic mean is defined as:
\[\bar x =\dfrac{ x_1+x_2+x_3+\cdots+x_n}{n} = \dfrac{1}{n}\sum_{i=1}^nx_i.\]</p>
<p>This is often shortened to: \[\bar x=\frac{1}{n}\sum x.\]</p>
</div>
<p> </p>
<p>In words, this formula says:&nbsp;<i>sum the values of \(x\) and divide by the number of numbers you have summed.</i></p>
<p>In the case of discrete data, we have a frequency for each possible value of the variable and we can express the calculation in terms of these frequencies. Suppose that the possible values are \(x_{(1)}, x_{(2)}, \ldots , x_{(J)}\) (where we could have \(J \rightarrow \infty\)) and that these values have observed frequencies \( f_{1}, f_{2}, \ldots , f_{J}\). Then it is easily seen that
\[ \sum_{i=1}^{n} x_{i} = \sum_{j=1}^{J} f_{j} x_{(j)}. \]</p>
<p>A common case is when the possible values are the nonnegative integers, \( 0, 1, 2, \ldots\). Then
\[ \bar{x} = \dfrac{1}{n} \sum_{k=0}^{\infty} k f_{k} = \dfrac{\sum k f_{k}}{\sum f_{k}}. \]</p>
<hr>
<h4 class="HELM_example">Example</h4>
<p>The numbers of telephone calls made on a particular day by 500 subscribers to a mobile telephone network gave the following frequency table, from which we can calculate the sample mean.</p>
<p></p><table class="HELM_table" style="text-align: center; width: 75%;">
<thead>
<tr><th>Number of <br>Calls \(x\)</th> <th>Frequency<br>\(f\)</th> <th>\(f x\)</th></tr> 
</thead>
<tbody>
<tr><td> 0</td> <td> 13</td><td> 0</td></tr>
<tr><td>1</td> <td> 30</td><td> 30</td></tr>  
<tr><td> 2</td> <td> 85</td> <td> 170 </td></tr>
<tr><td>3</td> <td>  98</td> <td> 294</td></tr>
<tr><td>4</td> <td> 96</td> <td> 384</td></tr>
<tr><td>5</td> <td> 81</td> <td>405 </td></tr>
<tr><td>6</td> <td>42</td> <td>252</td></tr>
<tr><td>7</td> <td>32</td> <td>224</td></tr>
<tr><td>8</td> <td>17</td> <td>136</td></tr>
<tr><td>9</td><td> 4</td> <td> 36</td></tr>
<tr><td>10</td> <td> 1 </td> <td> 10</td> </tr>
<tr style="border-bottom: 1px solid #000;"><td>11</td> <td> 1 </td> <td> 11</td></tr> 
<tr><td></td><td>\(\sum f = 500\)</td><td>\( \sum f x = 1952 \)</td></tr>
</tbody>
</table>
<p> </p>	 	 
<p>The sample mean is given by
\(\bar x=\dfrac{\sum fx}{\sum f}=\dfrac{1952}{500}=3.904\).</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200914145829+l4LWYA</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200914145829+Qngqgq</text></tag>
    </tags>
    
</question>
</quiz>