<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62168 -->
  <question type="stack">
    <name>
				<text>43.2.01 Introduction</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p></p><h3>Introduction</h3>

While medical researchers might be interested in knowing the answers to questions such as `Is age
related to blood pressure?' engineers might be interested in knowing the answers to questions such
as `Is the shear strength of a weld related to its diameter?' or `Is the rate of wear of a petrol engine
related to its operating temperature?' Hopefully you remember from the Regression section, statisticians measure the strength of a relationship between two variables by using a quantity called the correlation coefficient. As you might expect, tests exist which allow us to interpret the meaning of a calculated correlation coefficient.]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200705183206+bzGEFA</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200705183236+ywOOSi</text></tag>
    </tags>
    
</question>
</quiz>