<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25990 -->
  <question type="stack">
    <name>
				<text>18.2 Introduction</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span>Partial Derivatives</span><br></h3>

<p>When a function of more than one independent input variable changes because of changes in one or more of the input variables, it is important to calculate the change in the function itself. This can be investigated by holding all but one of the variables constant and finding the rate of change of the function with respect to the one remaining variable. This process is called partial differentiation. In this Section we show how to carry out the process.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200818045524+QDm5VL</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200818045905+DZiAhs</text></tag>
    </tags>
    
</question>
</quiz>