<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24620 -->
  <question type="stack">
    <name>
				<text>14.3.2.2 Find volume of parabola around y-axis</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Find the volume generated when the graph of \(y={@f@}\) between \(x={@low@}\) and \(x={@upp@}\) is rotated around the \(y\)-axis.</p>
<p>Using the latest Key Point write down the required integral:</p>
<table>
    <tbody>
        <tr>
            <td>\(\displaystyle \int\)</td>
            <td>
                <table>
                    <tbody>
                        <tr>
                            <td>[[input:ans2up]] [[validation:ans2up]]</td>
                        </tr>
                        <tr>
                            <td>[[input:ans2dn]] [[validation:ans2dn]]</td>
                        </tr>
                    </tbody>
                </table>
            </td>
            <td> [[input:ans2int]] [[validation:ans2int]] \(\mathrm{d}y\)</td>
        </tr>
    </tbody>
</table>
<p>[[feedback:prt2]]</p><br>

<p>This integral can be written entirely in terms of \(y\), using the fact that \(y={@f@}\) to eliminate \(x\). Do this now, and then evaluate the integral:</p>
<p>\(\quad V =\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_solution">Solution</h4>
<p>Using the latest Key Point, the required integral is \[\int_{@low@}^{@upp@} {@ta2int@} \,\mathrm{d}y,\] since \(c={@low@}\) and \(d={@upp@}\).</p>
<p>We have \(y={@f@}\), so the volume becomes \[\begin{align} \int_{@low@}^{@upp@} {@pi*x_sq@} \,\mathrm{d}y &amp;= \pi \left[ {@integrate(x_sq,y)@} \right]_{@low@}^{@upp@} \\ &amp;=\pi \left[ {@subst(y=upp,integrate(x_sq,y))@}-{@subst(y=low,integrate(x_sq,y))@} \right]
    \\ &amp;= {@ta3@} \end{align}\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>const: rand([2,3,4]);
x_sq: y/const;
f: const*x^2;
low: 0;
upp: rand(2)+1;

ta2up: const*upp^2;
ta2dn: const*low^2;
ta2int: pi*x^2;

ta3: simplify(integrate(pi*x_sq,y,ta2dn,ta2up));</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(y=\) {@f@}; \({@low@}\le x\le {@upp@}, V= {@ta3@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.</span>]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans2dn</name>
      <type>algebraic</type>
      <tans>ta2dn</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2int</name>
      <type>algebraic</type>
      <tans>ta2int</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2up</name>
      <type>algebraic</type>
      <tans>ta2up</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2up</sans>
        <tans>ta2up</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3300000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2dn</sans>
        <tans>ta2dn</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3300000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>EqualComAss</answertest>
        <sans>ans2int</sans>
        <tans>ta2int</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3400000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201201170236+Wl3UwC</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210607115033+D8qNF0</text></tag>
    </tags>
    <deployedseed>10883121</deployedseed>
    <deployedseed>1579797532</deployedseed>
    <deployedseed>718353704</deployedseed>
    <deployedseed>1670537042</deployedseed>
    <deployedseed>1137587650</deployedseed>
    <deployedseed>150421186</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans2dn</name>
        <value>ta2dn</value>
      </testinput>
      <testinput>
        <name>ans2int</name>
        <value>ta2int</value>
      </testinput>
      <testinput>
        <name>ans2up</name>
        <value>ta2up</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans2dn</name>
        <value>ta2dn+1</value>
      </testinput>
      <testinput>
        <name>ans2int</name>
        <value>ta2int+1</value>
      </testinput>
      <testinput>
        <name>ans2up</name>
        <value>ta2up+1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3+1</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-3-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans2dn</name>
        <value>ta2up</value>
      </testinput>
      <testinput>
        <name>ans2int</name>
        <value>ta2int</value>
      </testinput>
      <testinput>
        <name>ans2up</name>
        <value>ta2dn</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value></value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.3400000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore></expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <testinput>
        <name>ans2dn</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2int</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2up</name>
        <value>ta2up</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value></value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.3300000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-3-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore></expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <testinput>
        <name>ans2dn</name>
        <value>ta2dn</value>
      </testinput>
      <testinput>
        <name>ans2int</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2up</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value></value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.3300000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-3-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore></expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>