<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24367 -->
  <question type="stack">
    <name>
				<text>16.1.5.1 Finding convergence</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Which of the following sequences are convergent?</p><br>
<ol class="HELM_parts">
    <li>\(\displaystyle {@seqd[1]@},{@seqd[2]@},{@seqd[3]@},{@seqd[4]@},\ldots\)</li>
    <li>\(\displaystyle {@seqc[1]@},{@seqc[2]@},{@seqc[3]@},{@seqc[4]@},\ldots\)</li>
</ol><br>
<p>Select one:</p>
<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<ol class="HELM_parts">
    <li>This sequence is \[ \displaystyle {@seqd[1]@},{@seqd[2]@},{@seqd[3]@},{@seqd[4]@},\ldots =\displaystyle {@ev(seqd[1],simp:true)@},{@ev(seqd[2],simp:true)@},{@ev(seqd[3],simp:true)@},{@ev(seqd[4],simp:true)@},{@ev(seqd[5],simp:true)@},\ldots\] which does not converge.</li>
    <li>This sequence is \[\displaystyle {@seqc[1]@},{@seqc[2]@},{@seqc[3]@},{@seqc[4]@},\ldots = \displaystyle {@ev(seqc[1],simp:true)@},{@ev(seqc[2],simp:true)@},{@ev(seqc[3],simp:true)@},{@ev(seqc[4],simp:true)@},{@ev(seqc[5],simp:true)@},\ldots\] which converges to zero.</li>
</ol>
<br>
<p>Remember, a sequence is convergent if the limit of \(a_n\) as \(n\) increases can be found.</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>const: rand([pi/2, pi/3, pi/4]);
seq: [const, 2*const, 3*const, 4*const, 5*const];
simp:false;
seqd: map(lambda([x], sin(x)),seq);
seqc: map(lambda([x], sin(x)/x),seq);
choice:[[a,false,"(a)"],[b,true,"(b)"]];</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>\[\displaystyle {@seqd[1]@},{@seqd[2]@},{@seqd[3]@},{@seqd[4]@},\ldots \quad \text{divergent} \]
\[\displaystyle {@seqc[1]@},{@seqc[2]@},{@seqc[3]@},{@seqc[4]@},\ldots \quad \text{convergent}\]</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.</span>]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>choice</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(choice)[1]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201022101954+0UIjUy</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210609153210+YyoGmH</text></tag>
    </tags>
    <deployedseed>764221736</deployedseed>
    <deployedseed>2060465884</deployedseed>
    <deployedseed>130005795</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>a</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>b</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>