<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 30697 -->
  <question type="stack">
    <name>
				<text>13.4.1.3 By parts multiple times</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>Sometimes it is necessary to apply the formula more than once, as the next Example shows.</p>
<hr>
<h4 class="HELM_example">Example</h4>
<p></p>
<p>Find \({\displaystyle\int 2x^2 e^{-x} \, \mathrm{d}x}\).</p>
<h4 class="HELM_solution">Solution</h4>
<p>We let \(f=2x^2\) and \(g=e^{-x}\). Then \(\dfrac{\mathrm{d} f}{\mathrm{d} x}=4x\) and \({\displaystyle\int g \, \mathrm{d}x} = -e^{-x}\)</p>
<p>Using the formula for integration by parts we find</p>
\({\displaystyle\int 2x^2 e^{-x} \, \mathrm{d}x = 2x^2(-e^{-x})-\displaystyle\int 4x(-e^{-x}) \, \mathrm{d}x = -2x^2e^{-x}+\displaystyle\int4xe^{-x} \, \mathrm{d}x}\)

<p>We now need to find \({\displaystyle\int4xe^{-x} \, \mathrm{d}x}\) using integration by parts again. We get</p>

\begin{eqnarray*} \int 4xe^{-x}\ \mathrm{d} x&amp;=&amp;4x(-e^{-x})-\int4(-e^{-x})\ \mathrm{d} x\\ &amp;=&amp;-4xe^{-x}+\int 4e^{-x}\ \mathrm{d} x\\ &amp;=&amp;-4xe^{-x}-4e^{-x}\end{eqnarray*}
<p>Altogether we have</p>

\({\displaystyle\int 2x^2e^{-x} \, \mathrm{d}x} = -2x^2e^{-x}-4x e^{-x}-4 e^{-x}+c = -2e^{-x}(x^2+2x+2)+c\)]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210604110209+VSVw0D</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210609133633+gFB78r</text></tag>
    </tags>
    
</question>
</quiz>