<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24629 -->
  <question type="stack">
    <name>
				<text>14.4.2  The area of a surface of revolution</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>&nbsp;The area of a surface of revolution</h4>
<p>In previous workbooks we found an expression for the volume of a solid of revolution. Here we consider the more complicated problem of formulating an expression for the surface area of a solid of revolution.</p>
<p style="text-align: center;"><img src="@@PLUGINFILE@@/14-4-fig-12-Revolution.svg" alt="See text below." width="411" height="300" class="img-responsive atto_image_button_text-bottom"><br></p>
<p>The Figure above shows the portion of the curve \(y(x)\) between \(x = a\) and \(x = b\) which is rotated around the \(x\) axis through \(360^{\circ}\). A small disc, of thickness \(\delta x\), of the solid of revolution has been selected. Its radius
    is \(y\) and so its circumference has length \(2 \pi y\). (As usual we assume \(\delta x\) is ‘small’ so that the <b>curved</b> part of \(y(x)\) representing the hypotenuse of the highlighted ‘triangle’ can be regarded as <b>straight</b>). This surface
    ‘ribbon’, shown shaded, has a length \(2 \pi y\) and a width \(\sqrt{(\delta x)^2 + (\delta y)^2}\) and so its area is, to a good approximation, \(2 \pi y \sqrt{(\delta x)^2 + (\delta y)^2}\). We now let \(\delta x \to 0\) to obtain the result in
    the Key Point below:</p>
<div class="HELM_keypoint">
    <h4>Key Point</h4>
    <p>Given a curve with equation \(y = f(x)\), then the surface area of the solid generated by rotating that part of the curve between the points where \(x = a\) and \(x = b\) around the \(x\) axis is given by the formula: \[\text{area of surface} = \int_a^b
        2 \pi y \sqrt{1+\left(\frac{\mathrm{d}y}{\mathrm{d}x}\right)^2} \ \mathrm{d}x\]</p>
</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201202064419+qgUgTd</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210608092821+dC78Hj</text></tag>
    </tags>
    
</question>
</quiz>