<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25158 -->
  <question type="stack">
    <name>
				<text>2.1.3.1 Composition of Functions</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span style="color: inherit; font-family: inherit;">Composition of Functions</span><br></h3>
<p>Consider the two functions \(g(x)=x^2\), and \(h(x)=3x+5\). The Block diagrams below show the rules for these functions.</p>
<p style="text-align:center;"><img width="288" height="152" class="img-responsive atto_image_button_text-bottom" alt="composite functions" src="@@PLUGINFILE@@/2-1-fig-4.svg"><br></p>
<p style="text-align:center;"><span><strong>Block diagrams of two functions \(g\) and \(h\)</strong></span> </p>
<p>Suppose we place these Block diagrams together in series as shown below, so that the output from function \(g\) is used as the input to function \(h\).</p>
<p style="text-align:center;"><img width="468" height="70" class="img-responsive atto_image_button_text-bottom" alt="composite functions 2" src="@@PLUGINFILE@@/2-1-fig-5.svg"><br></p>
<p style="text-align:center;"><span><strong>The composition of the two functions to give \(h(g(x))\)</strong></span> </p>
<p>Study the figure above carefully and deduce that when the input to \(g\) is \(x\) the output from the two functions in series is \(3x^2+5\). Since the output from \(g\) is used as input to \(h\) we write</p>
<p style="text-align:center;"> \(h(g(x)) = h(x^2) = 3x^2+5\). </p>
<p>The form \(h(g(x))\) is known as the <span><strong>composition</strong></span> of the functions \(g\) and \(h\).</p>

<p>Suppose we interchange the two functions so that \(h\) is applied first as shown below.</p>
<p style="text-align:center;"><img width="487" height="67" class="img-responsive atto_image_button_text-bottom" alt="composite functions 3" src="@@PLUGINFILE@@/2-1-fig-6.svg"><br></p>
<p style="text-align:center;"><span><strong>The composition of the two functions to give \(g(h(x))\)</strong></span></p>
<p>Study the figure above and note that when the input to \(h\) is \(x\) the final output is \((3x+5)^2\). We write </p>
<p style="text-align:center;"> \(g(h(x)) = (3x+5)^2\)</p>
<p>Note that the function \(h(g(x))\) is different from \(g(h(x))\).</p>

<hr>
<h4 class="HELM_example">Example</h4>
<p>Given two functions \(g(t)=3t+2\) and \(h(t)=t+3\) obtain an expression for the composition \(g(h(t))\).</p>
<h4 class="HELM_solution">Solution</h4>
<p>We have \(g(h(t))= g(t+3)\). Now the rule for \(g\) is ‘triple the input and add \(2\)’, and so we can write \(g(t+3) = 3(t+3) +2= 3t+11\) so, \(g(h(t))= 3t+11\).</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>1</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200619130521+dGB6dr</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210831134128+r928ZV</text></tag>
    </tags>
    
</question>
</quiz>