<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24294 -->
  <question type="stack">
    <name>
				<text>5.1.1.9</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p></p>
<h5>Other contexts for linear models</h5>
<p></p>
<p>Linear functions may arise in other contexts. In each of these situations, the slope and intercept values will have some modelling significance. Indeed the behaviour and hence the suitability of a linear function, of the form \(y = ax + b\), when modelling
    any given situation will be determined by the values of <em>a</em> and <em>b</em>.</p>

<hr>
<h4 class="HELM_example">Example</h4>
<p>During \(20\) minutes of rain, a cylindrical rain barrel that is initially empty is filled to a depth of \(1.5 \, \textrm{cm}.\)</p>

<ol class="HELM_parts">
    <li>Choose variables to represent the level of water in the barrel and time. Sketch a graph representing the level of water in the barrel if the intensity of rainfall remains constant over the \(20\) minute period.</li>
    <li>Write down a linear function that represents the level of water in the vessel together with its range of validity</li>
    <li> State any assumptions that you have made:</li>
    <li> Write down the amended form of your answer to (b), if the vessel contains \(2 \, \textrm{cm}\) of water initially.
    </li>
</ol>

<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
    <li>In this answer \(h \, \textrm{cm}\) is used for the level of water measured from the bottom of the barrel and t minutes for time.&nbsp;<img src="@@PLUGINFILE@@/Block7_Fig7.svg" alt="A cylinder barrel filled with water up height h and graph of the water height as a function of time." width="505" height="169" class="img-responsive atto_image_button_text-bottom"></li>
    <li>
        <p>The intensity of rainfall is stated to be constant, so the rate at which the barrel fills may be taken as constant.</p>
      <p>The gradient of an appropriate linear function relating level of water (\(h \, \textrm{cm}\)) measured from the bottom of the vessel and time (\(t\), minutes) measured would be \( \frac{1.5}{20} = 0.075\) and an appropriate linear function would be \(h = 0.075t + c.\)</p>
        <p>Since the barrel is empty to start with, \(h = 0\) when \(t = 0,\) implying that \(c = 0.\) So the appropriate linear function and its range of validity are expressed by \( h = 0.075t, \quad (0 \leq t \leq 20).\)</p>
    </li>
    <li>It is assumed that the barrel has a uniformly cylindrical cross section, that no water is removed during the rainfall and there are no holes or leaks up to \(1.5 \, \textrm{cm}\) depth.
</li>
    <li>\( h = 0.075t+2, \quad (0 \leq t \leq 20).\)</li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200724081021+HbegUG</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220324190345+1QWDGq</text></tag>
    </tags>
    
</question>
</quiz>