<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26769 -->
  <question type="stack">
    <name>
				<text>10.3.3.4 Osborne example</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Feedback applied to an amplifier</h4>

<p>Feedback is applied to an amplifier such that <span class="math inline">\[\qquad A'=\dfrac{A}{1-\beta A}\]</span>&nbsp;where

    <span class="math inline">\(A', A\)</span> and <span class="math inline">\(\beta\)</span> are complex quantities.

    <span class="math inline">\(A\)</span> is the amplifier gain, <span class="math inline">\(A'\)</span> is the gain with feedback and

    <span class="math inline">\(\beta\)</span> is the proportion of the output which has been fed back.</p>

<figure class="HELM_figure"><img src="@@PLUGINFILE@@/look%20svg.svg" alt="Diagram of amplifier with feedback" width="369" height="176" role="presentation" class="img-responsive atto_image_button_text-bottom">
    <figcaption>An amplifier with feedback</figcaption>
</figure>

<ol class="HELM_parts">

    <li>If at 30 Hz, <span class="math inline">\(A=-500\)</span> and <span class="math inline">\(\beta=0.005e^{8\pi i /9}\)</span>, calculate <span class="math inline">\(A'\)</span> in exponential form.</li>

    <li>At a particular frequency it is desired to have <span class="math inline">\(A'=300e^{5\pi i /9}\)</span> where it is known that <span class="math inline">\(A=400e^{11\pi i /18}\)</span>. Find the value of <span class="math inline">\(\beta\)</span>        necessary to achieve this gain modification.</li>

</ol>

<p><strong>Mathematical statement of the problem</strong></p>

<ol class="HELM_parts">

    <li> Substitute

        <span class="math inline">\(A=-500\)</span> and <span class="math inline">\(\beta=0.005e^{8\pi i /9}\)</span> into

        <span class="math inline">\(A'=\dfrac{A}{1-\beta A}\)</span> in order to find <span class="math inline">\(A'\)</span>.</li>

    <li> We need to solve for <span class="math inline">\(\beta\)</span> when <span class="math inline">\(A'=300e^{5\pi i /9}\)</span> and

        <span class="math inline">\(A=400e^{11\pi i /18}\)</span>.</li>

</ol>

<p><strong>Mathematical analysis</strong></p>

<ol class="HELM_parts">

    <li><span class="math inline">We have \[A'=\dfrac{A}{1-\beta A}=\dfrac{-500}{1-0.005e^{8\pi i /9}\times (-500)}=\dfrac{-500}{1+2.5e^{8\pi i /9}}\]</span>&nbsp;Expressing the bottom line of this expression in Cartesian form this becomes:

        <span class="math inline">\[\qquad A'=\dfrac{-500}{1+2.5\cos\left(\dfrac{8\pi}{9}\right)+2.5i\sin\left(\dfrac{8\pi}{9}\right)}\approx \dfrac{-500}{-1.349+0.855i}\]</span>&nbsp;Expressing both the top and bottom lines in exponential form we get:

        <span class="math inline">\[\qquad A'\approx\dfrac{500e^{i\pi}}{1.597e^{i 2.576}}\approx 313e^{0.566i}\]</span></li>

    <li><span class="math inline">We have \[A'=\dfrac{A}{1-\beta A}\rightarrow A'(1-\beta A)=A\rightarrow

-\beta AA'=A-A'\] and </span><span class="math inline">\[\text{i.e.}\qquad \beta=\dfrac{A'-A}{AA'}\rightarrow\beta=\dfrac{1}{A}-\dfrac{1}{A'}\]</span>&nbsp;So

        <span class="math inline">\[\beta=\dfrac{1}{A}-\dfrac{1}{A'}=\dfrac{1}{400e^{11\pi i /18}}-\dfrac{1}{300e^{5\pi i /9}}\approx 0.0025e^{-11\pi i/18}-0.00333e^{-5\pi i /9}\]</span>&nbsp;Expressing both complex numbers in Cartesian form gives

        <span class="math inline">\[\beta = 0.0025\cos\left(-\dfrac{11\pi}{18}\right)+0.0025i\sin \left(-\dfrac{11\pi}{18}\right)\] \[\qquad-0.00333\cos \left(-\dfrac{5\pi}{9}\right)-0.00333i\sin \left(-\dfrac{5\pi}{9}\right)\] \[=-2.768\times 10^{-4}+9.3017\times 10^{-4}i=9.7048\times 10^{-4}e^{1.86i}\]</span>&nbsp;So
        to 3 significant figures <span class="math inline">\(\beta=9.70\times 10^{-4}e^{1.86i}\)</span></li>

</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200911145750+jfNT4r</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210913222308+kq7ZYV</text></tag>
    </tags>
    
</question>
</quiz>