<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 64396 -->
  <question type="stack">
    <name>
				<text>3.4.1.10 Use a graph to solve simultaneous</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>

<p>Use graphs to solve the following pairs of simultaneous equations.<br></p>

<p>\[{@sysl1[1,1]@}={@sysr1[1,1]@}, \quad (eq1)\]</p>
<p>\[{@sysl1[2,1]@}={@sysr1[2,1]@}, \quad (eq2)\]</p>
<p>By rearranging into the standrd form for tte equation of a straight line we get:</p>
<p>\( eq1: \quad y= \)&nbsp;[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>\( eq2: \quad y = \)&nbsp;[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]<br></p>
<p>By drawing the graphs we find that the two lines intersect at the point:</p>
<div>\({@B1@}=\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>We are given the equations:</p>
<p>\[{@sysl1[1,1]@}={@sysr1[1,1]@}, \quad (eq1)\]</p>
<p>\[{@sysl1[2,1]@}={@sysr1[2,1]@}, \quad (eq2)\]</p>

<p>By rearranging into the standrd form for tte equation of a straight line we get:</p>
<p>\( eq1: \quad y= {@eq1@}\)</p>
<p>\( eq2: \quad y = {@eq2@}\)</p>

<p>Now we draw the graphs of the two lines:</p>
<p>{@plot([eq1, eq2], [x,s1-3,s1+3], grid2d, [yx_ratio, 1],[legend, true] )@}</p>
<p>We find that the two lines intersect at the point:</p>
<div>\({@B1@}= {@ta@}\)</div>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021120900</text>
		</stackversion>
    <questionvariables>
				<text>ordergreat(x,y);

a1: rand([2,4]);
b1: rand([3,5])
A1: matrix([a1,-a1],[b1,1]);
B1:matrix([x],[y]);
sysl1:A1.B1;
s1: -(rand(4)+2);
s2: rand_with_prohib(-4,-1,[s1]);
ta:matrix([s1],[s2]);
sysr1:A1.ta;

eq1: (a1*x)/a1-sysr1[1,1]/a1;
eq2: -b1*x+sysr1[2,1];
</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\[{@sysl1=sysr1@}, {@B1=ta@}\]</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>eq1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>eq2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>matrix</type>
      <tans>ta</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>eq1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>eq2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+211217095550+JDQtO2</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220420155820+Di3DZh</text></tag>
    </tags>
    <deployedseed>541751447</deployedseed>
    <deployedseed>899340595</deployedseed>
    <deployedseed>1342916126</deployedseed>
    <deployedseed>1374571056</deployedseed>
    <deployedseed>22659542</deployedseed>
    <deployedseed>805503052</deployedseed>
    <deployedseed>1271752133</deployedseed>
    <deployedseed>1764685283</deployedseed>
    <deployedseed>2003108158</deployedseed>
    <deployedseed>2144178646</deployedseed>
    <deployedseed>1707030201</deployedseed>
    <deployedseed>67881432</deployedseed>
    <deployedseed>1236495857</deployedseed>
    <deployedseed>769780679</deployedseed>
    <deployedseed>802393112</deployedseed>
    <deployedseed>1919170488</deployedseed>
    <deployedseed>1818817386</deployedseed>
    <deployedseed>394849086</deployedseed>
    <deployedseed>105735676</deployedseed>
    <deployedseed>679319087</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>eq1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>eq2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(eq1+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(eq2+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ev(ta+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>