<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24384 -->
  <question type="stack">
    <name>
				<text>16.2.3.3 Note</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>Note that in all of these examples and exercises we have decided upon the convergence or divergence of various series; we have not been able to use the tests to discover what actual number the convergent series converges to.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+201013133659+uUD4Oc</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210609161319+nepmbK</text></tag>
    </tags>
    
</question>
</quiz>