<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27368 -->
  <question type="stack">
    <name>
				<text>39.1.9.1 The Standard Normal Probability Table</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>The Standard Normal Probability Table</h3>

<p> The Table below shows the Standard Normal Probability Table of values.</p>

<p> Use this Table to find \(Z\)-values when needed. It is recommended that you open this Table in another tab to refer to throughout the quiz.</p><p><img src="@@PLUGINFILE@@/39-1-standardnormaltable.svg" alt="The Standard Normal Probability Table" width="621" height="788" class="img-responsive atto_image_button_text-bottom"><br></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200731095241+O7f7Bi</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200826171705+h9vkCn</text></tag>
    </tags>
    
</question>
</quiz>