<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24688 -->
  <question type="stack">
    <name>
				<text>L10.1 Q06 Linear or not</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>Decide whether or not the following differential equations are linear:</p><p>\[ \sin ({@a@}t) \frac{d{@v@}}{dt}+{@v@} = {@b@} t \]
</p><p>[[input:ans1]] [[validation:ans1]]
</p><p>\[ \sin ({@a+3@}{@v@}) \frac{d{@v@}}{dt}+{@v@} = {@b-5@} t \]</p>
<p>[[input:ans2]] [[validation:ans2]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p></p><p>Linear equations are those in which the dependent variable, {@v@}, and it's derivative does not occur as products, reaised to powers on in nonlinear functions.&nbsp;</p><p>The equation&nbsp; \( \displaystyle \sin ({@a@}t) \frac{d{@v@}}{dt}+{@v@} = {@b@} t \)&nbsp; is {@ta_1@}.</p><p>The equation&nbsp; \(&nbsp;\displaystyle \sin ({@a+3@}{@v@}) \frac{d{@v@}}{dt}+{@v@} = {@b-5@} t\)&nbsp; is {@ta_2@}</p><br><p></p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2018080600</text>
		</stackversion>
    <questionvariables>
				<text>v: rand([y,z,x,f,g,h])
a: rand(10)+2
b: rand(8)+10
ta1: [["Non-linear",false], ["Linear",true]];
ta_1: "Linear"
ta2: [["Non-linear",true], ["Linear",false]];
ta_2: "Non-linear"</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]] [[feedback:prt2]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>\(  \sin ({@a@}t) \frac{d{@v@}}{dt}+{@v@} = {@b@} t \) is {@ta_1@}, \( \sin ({@a+3@}{@v@}) \frac{d{@v@}}{dt}+{@v@} = {@b-5@} t\) is {@ta_2@}</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>radio</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:moodle2.maths.ed.ac.uk+210319111609+2E3rc6</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+210319111609+XXA0AR</text></tag>
    </tags>
    <deployedseed>615490445</deployedseed>
    <deployedseed>1303371701</deployedseed>
    <deployedseed>1364970966</deployedseed>
    <deployedseed>1878929988</deployedseed>
    <deployedseed>835064316</deployedseed>
    <deployedseed>1215682561</deployedseed>
    <deployedseed>1592432774</deployedseed>
    <deployedseed>876266222</deployedseed>
    <deployedseed>721525947</deployedseed>
    <deployedseed>1437601476</deployedseed>
    <deployedseed>1465655894</deployedseed>
    <deployedseed>571267673</deployedseed>
    <deployedseed>1009068646</deployedseed>
    <deployedseed>1785727489</deployedseed>
    <deployedseed>1607281106</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta_2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>"Non-linear"</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>"Linear"</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>