<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27181 -->
  <question type="stack">
    <name>
				<text>35.1.3.4 Union of sets</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Union</h4>
<p>The <b>union</b> of two sets \(A\) and \(B\) is a set which contains all the elements of \(A\) together with all the elements of \(B\).  We write \(A \cup B\) to denote the union of \(A\) and \(B\).  We can describe the set \(A \cup B\) formally by:</p>
<div class="HELM_keypoint">
<h4>Key Point</h4>
<p>\(A \cup B = \{x: x \in A \hbox{ or }x \in B \hbox{ or both}\}\).</p>
</div><br>
<p>Thus the elements of the set \(A \cup B\) are those quantities \(x\) such that \(x\) is a member of 
\(A\) or  a member of \(B\) or  a member of both \(A\) and \(B\). The deeply shaded areas of the figures below represents \(A \cup B\).</p><p style="text-align: center;"><img src="@@PLUGINFILE@@/35-1-figure-5-union1.svg" alt="" width="235" height="185" class="img-responsive atto_image_button_text-bottom"><br></p><p>
</p><p style="text-align: center;">Figure (a)</p>
<p style="text-align: center;"><img src="@@PLUGINFILE@@/35-1-figure-5-union2.svg" alt="" width="235" height="185" class="img-responsive atto_image_button_text-bottom"></p><p></p><p></p>
<p style="text-align: center;">Figure (b)</p>
<p>In Figure (a) the sets intersect, whereas in Figure (b) the sets have no region in common. We say they are <b>disjoint</b>.</p>
<hr>
<h4 class="HELM_example">Example</h4>
<p>Given \(A = \{0,1\}\), \(B = \{1,2,3\}\) and \(C = \{2,3,4,5\}\) write
down </p>
<ol class="HELM_parts_inline">
<li>\(A \cup B\)</li>
<li>\(A \cup C\)</li>
<li>\(B \cup C\)</li>
</ol>
<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
<li>\(A \cup B = \{0,1,2,3\}\).</li>
<li>\(A \cup C = \{0,1,2,3,4,5\}\).</li>
<li>\(B \cup C =\{1,2,3,4,5\}\).</li>
</ol>
<p>Recall that there is no need to repeat elements in a set. Clearly the order of the union is
unimportant so \(A \cup B = B \cup A\).</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200701010939+6pak7q</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200727223858+SFHCCR</text></tag>
    </tags>
    
</question>
</quiz>