<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26515 -->
  <question type="stack">
    <name>
				<text>6.2.3.2 tanh^2</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise&nbsp;</p>
<p>Which of the following identities are correct?</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>Note the following</p>
<p> \( 1 - \tanh^2 x = 1 - \dfrac{\sinh^2 x}{\cosh^2 x} = \dfrac{\cosh^2 x - \sinh^2 x}{\cosh^2 x} = \dfrac{1}{\cosh^2 x} = \) \( \rm{sech}^2 \) \(x\). </p>
<p>To see that each of the other identities is incorrect, specify a particular value for \(x\) for which the identity is wrong. Also note that they are inconsistent with the above correct identities. </p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200619172311+jJ2scK</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200715142026+CqNQn4</text></tag>
    </tags>
    
</question>
</quiz>