<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24618 -->
  <question type="stack">
    <name>
				<text>14.3.1.7 Finding volume of a solid of revolution</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>The parabola \(y^2={@y^2@}\) for \({@x1@}\le x\le {@x2@}\), is rotated around the \(x\)-axis. Find the volume of the solid formed.</p>
<p>\(\quad V=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_solution">Solution</h4>
<p>Using the Key Point, we have \(y^2={@y^2@}\), \(a={@x1@}\) and \(b={@x2@}\). Then, \[\begin{aligned} V &amp;= \int_a^b \pi y^2 \,\mathrm{d}x = \int_{@x1@}^{@x2@} \pi({@y^2@}) \,\mathrm{d}x \\ &amp;= \pi \left[{@integrate(y^2,x)@}\right]_{@x1@}^{@x2@} \\ &amp;= \pi \left[{@subst(x=x2,integrate(y^2,x))@}-{@subst(x=x1,integrate(y^2,x))@}\right]
    \\ &amp;= {@vol@} \end{aligned}\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>y: sqrt((rand([4,9]))*x);
x1: rand(2);
x2: x1+1;
vol: integrate(pi*y^2,x,x1,x2);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(y=\) {@y@}; {@x1@}\(\le x\le\) {@x2@}</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.</span>]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>vol</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>vol</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201201163617+qhdvQK</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210607114815+ftFB08</text></tag>
    </tags>
    <deployedseed>518225142</deployedseed>
    <deployedseed>1351974419</deployedseed>
    <deployedseed>595403344</deployedseed>
    <deployedseed>339351056</deployedseed>
    <deployedseed>1603754338</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>vol</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>vol+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>