<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27176 -->
  <question type="stack">
    <name>
				<text>35.1.1.2 Subsets</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Subsets</h4><p>Sometimes one set is contained completely within another set.  For example, if \(X = \{2,3,4,5,6\} \hbox{ and } Y = \{2,3,6\}\), then all the elements of \(Y\) are also elements of \(X\).  We say that \(Y\) is a <b>subset</b> of \(X\) and write \(Y \subseteq X\).<br></p>
<hr>
<h4 class="HELM_example">Example</h4>
<p>Given \(A = \{0,1,2,3\}\), \(B = \{0,1,2,3,4,5,6\}\) and \(C = \{0,1\}\), state which sets are subsets.</p>
<h4 class="HELM_solution">Solution</h4>
<p>\(A\) is a subset of \(B\), that is \(A \subseteq B\).</p>
<p>\(C\) is a subset of \(B\), that is \(C \subseteq B\).</p>
<p>\(C\) is a subset of \(A\), that is \(C \subseteq A\).</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200630232604+Wy5xSF</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200630232653+9n5LSC</text></tag>
    </tags>
    
</question>
</quiz>