<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24288 -->
  <question type="stack">
    <name>
				<text>5.1.1.2 Constant functions</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Constant Functions Example</h4>
<p>There are two physical interpretations of constancy that are of interest here. A very common form is <strong>constancy with time</strong>. Motion under gravity may be modelled as motion with constant acceleration. By definition, Fixed Rate Mortgages (increasingly
    popular in the late 1990s) offer a constant rate of interest over a specified period. In these examples, the constancy will be limited to a certain time interval. Motion under gravity will only involve the constant acceleration due to the Earth’s
    gravitational pull as long as the motion is close to the Earth’s surface. In any case the acceleration will only be from the time the object is released to the time it stops. Unfortunately, increases in base interest rates eventually feed into mortgage
    rates. So mortgage lenders are only able to offer fixed rates for a certain time. A mathematical statement of these limits is a statement of the <em>range of validity</em> of the constant function model.</p>
<p> Another type of constancy is <strong>constancy in space</strong>. Long stretches of Roman roads were built in a fixed direction. For at least part of their lengths, roads have constant width. In modelling the formation and movement of seismic waves in
    the Earth’s crust it is convenient to assume that the layers from which the Earth’s crust is formed have constant thickness with respect to the Earth’s surface. In these cases the assumption of constancy will only be valid within certain limits in space.</p>
<p></p>
<p><br></p>
<hr>
<h4 class="HELM_example">Example</h4>
<p></p>
<p>The rate of flow of water from a tap is denoted as \(r\) (litres per minute). The time for which it is turned on is denoted by \(t\) (minutes). Suppose that a tap is turned on and that the rate of water running out of a tap is assumed to be constant at
    3 litres per minute and that it is turned off after 10 minutes.</p>
<p></p>
<ol class="HELM_parts">
    <p></p>
    <li>Write down a mathematical statement of the model for the flow from the tap, including its range of validity.</li>
    <li>Sketch a graph of the variation of <em>r</em> with <em>t</em>.</li>
    <li>Find an equation for the number of litres of water that have run out of the tap after <em>t</em> minutes.</li>
    <li>Calculate the volume of water that has run out of the tap three minutes before it is turned off.</li>
</ol>
<p></p>
<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
    <li> r = 3 (0 ≤ <em>t</em> ≤ 10)</li>
    <li><img width="429" height="139" class="img-responsive atto_image_button_text-bottom" alt="graph showing r vs t at a constant value of 3 litres per minute" src="@@PLUGINFILE@@/5-1-example%20solution%20part%20b.svg"></li>
    <li> <em>V = rt</em> (0 ≤<em> t </em>≤ 10)</li>
    <li> The tap will have run for seven minuters; 3 litres per minute x 7 minutes = 21 litres</li>
</ol>
<p> Note that a more sophisticated model would allow for the variation in flow rate as the tap is turned on and off.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200722145217+bCSCCK</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220215165502+7aNDqt</text></tag>
    </tags>
    
</question>
</quiz>