<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62123 -->
  <question type="stack">
    <name>
				<text>23.3.1c key point</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[These results are summarized in the following Key Point.
<p></p>
<div class="HELM_keypoint">
<h4>Key Point</h4>
<p>Products of functions</p>
\[
\begin{align}
\text{(even) } \times \text{(even)} &amp;= \text{(even)} \\
\text{(even) } \times \text{(odd)} &amp;= \text{(odd)} \\
\text{(odd) } \times \text{(odd)} &amp;= \text{(even)} 
\end{align}
\]
<p>Sums of functions</p>
\[
\begin{align}
\text{(even) } + \text{(even)} &amp;= \text{(even)} \\
\text{(even) } + \text{(odd)} &amp;= \text{(neither)} \\
\text{(odd) } + \text{(odd)} &amp;= \text{(odd)} 
\end{align}
\]
</div>
Useful properties of even and of odd functions in connection with integrals can be readily deduced if
we recall that a definite integral has the significance of giving us the value of an area as shown in the Figure below.<p><br></p><p style="text-align: center;"><img src="@@PLUGINFILE@@/23-3-1-integral-area.svg" alt="area under a curve" width="201" height="124" class="img-responsive atto_image_button_text-bottom"><br></p><p style="text-align: center;"><br></p>
\( \displaystyle{ \int_a^b } f(t) dt\) gives us the <b>net</b> value of the shaded area, that above the \(t\)-axis being positive, that below being negative.
<p></p>
For the case of a symmetrical interval \( (-a, a)\) for an odd function \(g(t)\) we have 
\[ \displaystyle{\int_{-a} ^a } g(t) dt = 0\]
and for an even function \(h(t)\) we have
\[ \displaystyle{\int_{-a} ^a } h(t) dt = 2 \displaystyle{\int_{0} ^a } h(t) dt \]
This can be seen in the Figure below.<p><br></p><p style="text-align: center;"><img src="@@PLUGINFILE@@/23-3-1-even-odd-integrals.svg" alt="even-odd-integrals" width="440" height="153" class="img-responsive atto_image_button_text-bottom"><br></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200901163055+38E9KY</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200901164719+TXwOhh</text></tag>
    </tags>
    
</question>
</quiz>