<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27206 -->
  <question type="stack">
    <name>
				<text>35.2.2.3 Defining probability</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h5> Notation </h5>
<p> If \(A\) is an event associated with a sample space \(S\) then the probability of \(A\) occurring is denoted by \(\mathbb{P}(A)\). Referring to the examples above we may immediately deduce that</p>
<ol type="a">
<li> \(\displaystyle \text{P}(H) = \text{P}(T) =\frac12\) </li>
<li> \(\displaystyle \text{P}(1) = \text{P}(2) = \text{P}(3) = \text{P}(4) = \text{P}(5) = \text{P}(6) = \frac16\) </li>
<li> \(\displaystyle \text{P}(HHH) = \text{P}(HHT) = \text{P}(HTH) = \text{P}(THH)= \text{P}(HTT)\)\[ =\text{P}(THT) = \text{P}(TTH) = \text{P}(TTT) = \frac18\] </li>
</ol>
<h5> Definition </h5>
<p> We can now define probability using the Principle of Equally Likely Outcomes as follows:
If a sample space \(S\) consists of \(n\) simple outcomes which are equally likely and an event \(A\) consists of \(m\) of those simple outcomes, then\[\text{P}(A) = \frac{m}{n} = \frac{\text{the number of simple outcomes in }A}{\text{the number of simple outcomes in }S}.\]It follows from definition that \(0\leq\text{P}(A)\leq1\).</p>
<ul>
<li> If \(\text{P}(A) = 1\) we say that the event \(A\) is <strong>certain</strong> because \(A\) is identical to \(S\). </li>
<li> If \(\text{P}(A) = 0\) we say that the event \(A\) is <strong>impossible</strong> because \(A\) is empty. </li>
</ul>
<p> The set with no outcomes in it is called the <strong>empty set</strong> and is written \(\varnothing\); therefore \(\text{P}(\varnothing)=0\).</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200701105802+icma4A</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+211020103135+Wqy38b</text></tag>
    </tags>
    
</question>
</quiz>