<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25991 -->
  <question type="stack">
    <name>
				<text>18.2.1.1 First partial derivatives</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span>First partial derivatives</span><br></h3>

 <p>For a function of a single variable, \(y=f(x)\), changing the independent variable \(x\) leads to a corresponding change in the dependent variable \(y\). The of \(y\) with respect to \(x\) is given by the <strong>derivative</strong>, written \(\dfrac{df}{dx}\). A similar situation occurs with functions of more than one variable. For clarity we shall concentrate on functions of just two variables. In the relation \(z=f(x,y)\) the <strong>independent</strong> variables are \(x\) and \(y\) and the <strong>dependent variable</strong> \(z\). We have seen in Section 18.1 that as \(x\) and \(y\) vary the \(z\)-value traces out a surface. Now both of the variables \(x\) and \(y\) may change <em>simultaneously</em> inducing a change in \(z\). However, rather than consider this general situation, to begin with we shall hold one of the independent variables <strong>fixed</strong>. This is equivalent to moving along a curve obtained by intersecting the surface by one of the coordinate planes.</p>
<p>Consider \(f(x,y)=x^3+2x^2y+y^2+2x+1.\)</p>
<p>Suppose we keep \(y\) constant and vary \(x\); then what is the rate of change of the function \(f\)?</p>
<p> Suppose we hold \(y\) at the value \(3\) then</p>
<p></p>
<p>\(f(x,3)=x^3+6x^2+9+2x+1 = x^3+6x^2+2x+10\)</p>

<p>In effect, we now have a function of \(x\) only. If we differentiate it with respect to \(x\) we obtain the expression: \(3x^2+12x+2\).</p>
<p>We say that \(f\) has been <strong>partially differentiated</strong> with respect to \(x\). We denote the partial derivative of \(f\) with respect to \(x\) by \(\dfrac{\partial f}{\partial x}\) (to be read as ‘partial dee \(f\) by dee \(x\)’ ). In this example, when \(y=3\): \(\dfrac{\partial f}{\partial x} = 3x^2+12x+2.\)</p>
<p>In the same way if \(y\) is held at the value 4 then \(f(x,4)=x^3+8x^2+16+2x+1=x^3+8x^2+2x+17\) and so, for this value of \(y\)</p>
<p>\(\dfrac{\partial f}{\partial x} = 3x^2+16x+2.\)</p>

<p>Now if we return to the original formulation \(f(x,y)=x^3+2x^2y+y^2+2x+1\) and treat \(y\) as a constant then the process of partial differentiation with respect to \(x\) gives</p>
\(\dfrac{\partial f}{\partial x} = 3x^2+4xy+0+2+0\\ \phantom{\dfrac{\partial f}{\partial x}} = 3x^2+4xy+2.\)

<div class="HELM_keypoint">
<h4>Key Point</h4>
<p><strong>The Partial Derivative of  \(\pmb f\) with respect to \(\pmb x\)</strong></p>
<p> For a function of two variables \(z=f(x,y)\) the partial derivative of \(f\) with respect to \(x\) is denoted by \(\dfrac{\partial f}{\partial x}\) and is obtained by differentiating \(f(x,y)\) with respect to \(x\) in the usual way but treating the \(y\)-variable as if it were a constant.</p>
<p>Alternative notations for \(\dfrac{\partial f}{\partial x}\) are \(f_x(x,y)\) or \(f_x\) or \(\dfrac{\partial z}{\partial x}\)
</p></div>

<hr>
<h4 class="HELM_example">Example</h4>
<ol class="HELM_parts_inline">
<p>Find \(\dfrac{\partial f}{\partial x}\) for</p> 
<li> \(f(x,y)=x^3+x+y^2+y\), </li>
<li> \(f(x,y)=x^2y+xy^3\).</li>
</ol>

<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
<li>\(\dfrac{\partial f}{\partial x}=3x^2+1+0+0=3x^2+1\)</li>
<li>\(\dfrac{\partial f}{\partial x}=2x\times y+1\times y^3=2xy+y^3\)</li>
</ol><p>&nbsp;</p>
<hr>

<p><strong>The \(y\) partial derivative</strong></p>
<p>For functions of two variables \(f(x,y)\) the \(x\) and \(y\) variables are on the same footing, so what we have done for the \(x\)-variable we can do for the \(y\)-variable. We can thus imagine keeping the \(x\)-variable fixed and determining the rate of change of \(f\) as \(y\) changes. This rate of change is denoted by \(\dfrac{\partial f}{\partial y}\).</p>

<div class="HELM_keypoint">
<h4>Key Point</h4>
<p><strong>The Partial Derivative of  \(\pmb f\) with respect to \(\pmb y\)</strong></p>
<p> For a function of two variables \(z=f(x,y)\) the partial derivative of \(f\) with respect to \(y\) is denoted by \(\dfrac{\partial f}{\partial y}\) and is obtained by differentiating \(f(x,y)\) with respect to \(y\) in the usual way but treating the \(x\)-variable as if it were a constant.</p>
<p>Alternative notations for \(\dfrac{\partial f}{\partial y}\) are \(f_y(x,y)\) or \(f_y\) or \(\dfrac{\partial z}{\partial y}\)
</p></div>
<p><br></p>
<p>Returning to \(f(x,y)=x^3+2x^2y+y^2+2x+1\) once again, we therefore obtain:</p>
<p>\(\dfrac{\partial f}{\partial y} = 0+2x^2\times 1+2y+0+0 = 2x^2+2y\).</p>

<hr>
<h4 class="HELM_example">Example</h4>
<ol class="HELM_parts_inline">
<p>Find \(\dfrac{\partial f}{\partial y}\) for</p> 
<li> \(f(x,y)=x^3+x+y^2+y\), </li>
<li> \(f(x,y)=x^2y+xy^3\).</li>
</ol>

<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
<li>\(\dfrac{\partial f}{\partial y}=0+0+2y+1=2y+1\)</li>
<li>\(\dfrac{\partial f}{\partial y}=x^2\times 1+x\times 3y^2=x^2+3xy^2\)</li>
</ol>
<hr>

<p>We can calculate the partial derivative of \(f\) with respect to \(x\) and the value of \(\dfrac{\partial f}{\partial x}\) at a specific point e.g. \(x=1, y=-2\).</p>

<hr>
<h4 class="HELM_example">Example</h4>
<p>[Remember \(f_x\) means \(\dfrac{\partial f}{\partial x}\) and \(f_y\) means \(\dfrac{\partial f}{\partial y}\).]
</p><p>Find \(f_x(1,-2)\) and \(f_y(-3,2)\) for \(f(x,y)=x^2+y^3+2xy\)</p> 

<h4 class="HELM_solution">Solution</h4>
<p>\(f_x(x,y)=2x+2y\), so  \(f_x(1,-2)=2-4=-2\);</p>
<p>\(f_y(x,y)=3y^2+2x,\) so \(f_y(-3,2)=12-6=6\)</p>

<hr>
<h4 class="HELM_example">Example</h4>
<p>Given \(f(x,y)=3x^2+2y^2+xy^3\) find \(f_x(1,-2)\) and \(f_y(-1,-1)\).</p>
 
<h4 class="HELM_solution">Solution</h4>
<p>First find expressions for \(\dfrac{\partial f}{\partial x}\) and \(\dfrac{\partial
f}{\partial y}\):</p>
<p>\(\dfrac{\partial f}{\partial x}=6x+y^3\),</p>
<p>\(\dfrac{\partial f}{\partial y}=4y+3xy^2\)</p>
<p>Now calculate \(f_x(1,-2)\) and \(f_y(-1,-1)\):</p>
<p>\(f_x(1,-2)=6\times 1+(-2)^3=-2\);</p>
<p>\(f_y(-1,-1)=4\times (-1)+3(-1)\times 1=-7\)</p>
<hr>

<p><strong>Functions of several variables</strong></p>
<p>As we have seen, a function of two variables \(f(x,y)\) has two partial derivatives, \(\dfrac{\partial f}{\partial x}\) and \(\dfrac{\partial f}{\partial y}\). In an exactly analogous way a function of three variables \(f(x,y,u)\) has three partial derivatives \(\dfrac{\partial f}{\partial x}\), \(\dfrac{\partial f}{\partial y}\) and \(\dfrac{\partial f}{\partial u}\), and so on for functions of more than three variables. Each partial derivative is obtained in the same way as stated in the following Key Point:</p>

<div class="HELM_keypoint">
<h4>Key Point</h4>
<p><strong>The Partial Derivatives of \(\pmb f(\) \(\pmb x\), \(\pmb y\), \(\pmb u\), \(\pmb v\), \(\pmb w\), \(\dots\)\()\)</strong></p>
<p>For a function of several variables \(z=f(x, y, u, v, w, \dots)\) the partial derivative of
\(f\) with respect to \(v\) (say) is denoted by \(\dfrac{\partial f}{\partial v}\) and is obtained by differentiating \(f(x, y, u, v, w, \dots)\) with respect to \(v\) in the usual way but treating all the other variables as if they were constants.</p>
<p>Alternative notations for \(\dfrac{\partial f}{\partial v}\) when \(z = f(x, y, u, v, w, \dots )\) are
\(f_v(x, y, u, v, w \dots )\) and \(f_v\) and \(\dfrac{\partial z}{\partial v}\).</p>
</div>

<hr>
<h4 class="HELM_example">Example</h4>
<p>Find \(\dfrac{\partial f}{\partial x}\) and \(\dfrac{\partial f}{\partial u}\) for \(f(x, y, u, v) = x^2+xy^2+y^2u^3-7uv^4\)</p>
 
<h4 class="HELM_solution">Solution</h4>
<p>\(\dfrac{\partial f}{\partial x}=2x+y^2+0+0=2x+y^2\),</p>
<p>\(\dfrac{\partial f}{\partial u}=0+0+y^2\times 3u^2-7v^4=3y^2u^2-7v^4\)</p>
<hr>

<h4 class="HELM_example">Example</h4>
<p>The pressure, \(P\), for one mole of an ideal gas is related to its absolute temperature, \(T\), and specific volume, \(v\), by the equation \(Pv=RT\) where \(R\) is the gas constant. Obtain simple expressions for</p> 
<ol>
<li> the coefficient of thermal expansion, \(\alpha\), defined by: \(\alpha = \dfrac{1}{v}\left(\dfrac{\partial v}{\partial T}\right)_P\)</li><br>
<li> the isothermal compressibility, \(\kappa_T\), defined by: \(\kappa_T=-\dfrac{1}{v}\left(\dfrac{\partial v}{\partial P}\right)_T\)</li>

<h4 class="HELM_solution">Solution</h4>
<p>
\(v=\dfrac{RT}{P}\quad\Rightarrow\quad \left(\dfrac{\partial v}{\partial
T}\right)_P=\dfrac{R}{P}\)</p>
<p>so \(\alpha=\dfrac{1}{v}\left(\dfrac{\partial v}{\partial
T}\right)_P=\dfrac{R}{Pv}=\dfrac{1}{T}\)
</p><br>
<p>
\(v=\dfrac{RT}{P}\quad\Rightarrow\quad \left(\dfrac{\partial v}{\partial
P}\right)_T=-\dfrac{RT}{P^2}\)</p>
<p>so \(\kappa_T=-\dfrac{1}{v}\left(\dfrac{\partial v}{\partial
P}\right)_T=\dfrac{RT}{vP^2}=\dfrac{1}{P}\)</p>
<hr></ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200818045608+iczQLI</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200818092352+Gt3oxD</text></tag>
    </tags>
    
</question>
</quiz>