<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24989 -->
  <question type="stack">
    <name>
				<text>W4 Q04 Lagrange - point to plane</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>Use Lagrange multipliers to find the distance from the point \( P=({@a@},{@b@},{@c@}) \) to the plane {@pl@} as follows.</p>
<p>Let \(f(x,y,z)\) be the square of the distance from \( (x,y,z) \) to \( P=({@a@},{@b@},{@c@}) \).</p>
<p>\( f(x,y,x) = \) [[input:ans1]] [[validation:ans1]]</p>
<p>Let \( g(x,y,z)=0 \) be the constraint that \((x,y,z)\) lies on {@pl@}.  We define the Lagrangian \[ L= f(x,y,z) + lg(x,y,z)\]</p>
<p>Calculate \( \frac{\partial L}{\partial x} = \) [[input:ans2]] [[validation:ans2]]</p>
<p>Calculate \( \frac{\partial L}{\partial l} = \) [[input:ans3]] [[validation:ans3]]</p>
<p>Solve the equations \( \frac{\partial L}{\partial x} = 0\),  \( \frac{\partial L}{\partial y} = 0\), \( \frac{\partial L}{\partial z} = 0\), \( \frac{\partial L}{\partial l} = 0\), and find the value of \(l\).</p><p>
</p><p>\( l= \) [[input:ans4]] [[validation:ans4]]</p>
<p>Hence, find the shortest distance from the point to the plane.</p>
<p>Shortest distance is  [[input:ans5]] [[validation:ans5]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[We define
\[ L= f(x,y,z) + lg(x,y,z)\]
In our case this gives us
\[ L = {@taf@} - l({@tag@}).\]
Now we calculate the derivatives as
\[ \frac{\partial L}{\partial x} = {@Lx@}\]
\[ \frac{\partial L}{\partial y} = {@Ly@}\]
\[ \frac{\partial L}{\partial z} = {@Lz@}\]
\[ \frac{\partial L}{\partial l} = {@Ll@}\]
Now we need to equate to zero and solve this system of algebraic equations.
\[ {@Lx@}=0,\quad {@Ly@}=0,\quad {@Lz@}=0,\quad {@Ll@}=0.\]
In this case, we can solve for \(x\), \(y\) and \(z\) which give us solutions in terms of \(l\).
\[ {@taxl@} \]
\[ {@tayl@} \]
\[ {@tazl@} \]
When we use these values in \({@Ll@}=0\) we get 
\[ {@L1@}=0.\]
Although this looks complicated it is a linear equation in \(l\) which we can readily solve to give \({@tal@}\).
So, this gives the actual values of \(x\), \(y\) and \(z\) as
\[ {@tax@}\]
\[ {@tay@}\]
\[ {@taz@}\]
Now, we need to evaluate \(f\) at this point to give us the square of the distance.
\[ f(x,y,z) = {@fv@}\]
So, the closest distance of \( P=({@a@},{@b@},{@c@}) \) to the plane {@pl@} is {@sqrt(fv)@}.]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>1</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>a:1;
b:1+rand(2);
c:-1;
A:2;
B:-1;
C:2;
r:3+rand(4);
/* Distance to the point (a,b,c): */
pt:matrix([a,b,c]);
taf:(x-a)^2+(y-b)^2+(z-c)^2;
/* To a plane */
pl:A*x+B*y+C*z=r;
tag:A*x+B*y+C*z-r;
L:taf+l*tag;
/* Derivatives */
Lx:diff(L,x);
Ly:diff(L,y);
Lz:diff(L,z);
Ll:diff(L,l);
/* Start to solve these. Lx, Ly, Lz all have unique solutions! */
taxl:first(solve(Lx=0,x));
tayl:first(solve(Ly=0,y));
tazl:first(solve(Lz=0,z));
/* Evaluate the constraint: this is the hideous part! */
L1:ev(Ll,[taxl,tayl,tazl]);
tal:first(solve(L1,l));
/* Sub back in l, to get x,y,z. */
tax:ev(taxl,tal);
tay:ev(tayl,tal);
taz:ev(tazl,tal);
fv:ev(taf,[tax,tay,taz])</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>Distance from \( P=({@a@},{@b@},{@c@}) \) to the plane {@pl@} is {@sqrt(fv)@}.</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>taf</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>Lx</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>Ll</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>rhs(tal)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans5</name>
      <type>numerical</type>
      <tans>sqrt(fv)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>taf</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.2000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text><p>Your function \(f(x)\) is correct.</p></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text><p>Your function \(f(x)\) is incorrect.<br></p></text>
				</falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>Lx</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
					<text><p>Your value of \(\frac{\partial L}{\partial x}\) is correct.</p></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
					<text><p>Your value of \(\frac{\partial L}{\partial x}\) is incorrect.<br></p></text>
				</falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>Ll</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
					<text><p>Your value of \(\frac{\partial L}{\partial l}\) is correct.<br></p></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
					<text><p>Your value of \(\frac{\partial L}{\partial l}\) is incorrect.<br></p></text>
				</falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>rhs(tal)</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
					<text><p>Your value of \(l\) is correct!</p></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
					<text><p>Your value of \(l\) is incorrect!<br></p></text>
				</falsefeedback>
      </node>
      <node>
        <name>4</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>sqrt(fv)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
					<text><p>Correct final distance, well done.</p></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>5</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>fv</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
					<text><p>You have the square of the distance!</p></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:moodle2.maths.ed.ac.uk+190308114721+tXktmq</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+210314100738+xaSzRr</text></tag>
    </tags>
    <deployedseed>1057140986</deployedseed>
    <deployedseed>1939874711</deployedseed>
    <deployedseed>346418144</deployedseed>
    <deployedseed>1260033592</deployedseed>
    <deployedseed>745204119</deployedseed>
    <deployedseed>1398748464</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>Lx</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>Ll</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ev(rhs(tal),simp)</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>sqrt(fv)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-5-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>taf</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>Lx</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>Ll</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ev(rhs(tal),simp)</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>fv</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.8000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-6-T</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>