<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24305 -->
  <question type="stack">
    <name>
				<text>5.3.2.3 Oscillating functions: modelling tides</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p></p>
<p>The peak levels of the graph correspond to times when the sine function has the value 1. The lowest points correspond to times when the sine function is \(-1\). At these times the arguments of the sine function (i.e. \(2.7t+8.5\)) are an odd number of
    \(\pi/2\) starting at \(3\pi /2\) for the first low tide.</p>

<p>So far all of this may be deduced from the general form \(y = a\)sin\((bx)\) and from the modelling context. However there is an additional term in the function being considered here. This is a constant \(8.5\) within the sine function. When \(t = 0\)
    the presence of this constant means that the intercept on the height axis is 3.2sin\((8.5)\) = 2.56, implying that the water level is 2.56 m above the mean value at the start of timing. The constant 8.5 has displaced the sine curve sideways. This
    constant is known as the <b>phase</b>of the function. Phase is measured in radians as it is an angle.</p>

<p>As remarked earlier, at \(t = 0\), this function has the value \(3.2\sin (8.5)\). Since \(\sin (8.5) = \sin(8.5 - 2\pi) \approx \sin (2.2168)\), we can replace the constant 8.5 by 2.2168 without altering the values on the graph. This means that the function</p>
<p>\(\qquad h = 3.2\sin (2.74t + 2.2168)\)</p>
<p>does just as well as the original function in representing the tidal variation in the harbour. We now rewrite this latest form of the function, representing the variation of water level in the harbour, so that time is measured in hours rather than in
    six-hourly invervals. The effect of changing the units of time to hours from 6 hours is to decrease the coefficient of \(t\) in the sine function by a factor of 6, so that the new function is</p>
<p>\(\qquad h = 3.2 \sin (0.45t + 2.2168).\)<br></p>
<p>See the figure below</p>
<p><img src="@@PLUGINFILE@@/5-3-fig-3-height_vs_time_example.svg" alt="" width="407" height="151" role="presentation" class="img-responsive atto_image_button_text-bottom"><br></p>

<p>We can use the latest form of the function to calculate the time of the first low tide assuming that \(t = 0\) corresponds to midnight.</p>
<p>At the first low tide, \(h = -3.2\)\(\quad\)and\(\quad \sin(0.45t + 2.2168) = -1,\)</p>
<p> Using the fact that \(\sin(\dfrac{3\pi}{2})=-1\), we have</p>
<p>\(\qquad 0.45t + 2.2168 = 3\pi /2\),\(\quad\)giving\(\quad\)\(t = 5.5458=5.55\)\(\quad\)to \(2\) d.p.</p>
<p>so the first low tide is at 5:30 a.m.<br></p>

<p><br></p>
<p>There are various rules connected with sine and cosine functions that can be summarised at this point.</p>
<p><b>(1)</b>\(\quad\)Placing a multiplier before \(\sin x\) or \(\cos x\) (e.g. \(2\sin x\)) changes the amplitude without changing the period.</p>
<p><b>(2)</b>\(\quad\)Placing a multiplier before \(x\) in \(\sin x\) or \(\cos x\), (e.g. \(\sin 3x\)), changes the period or frequency without changing the amplitude.</p>
<p><b>(3)</b>\(\quad\)As with any function, the addition of a constant (e.g. \(4+\sin x\)) raises or lowers the whole graph of the sine or cosine function. It alters the mean value without changing the amplitude.</p>
<p><b>(4)</b>\(\quad\)Changing the sign within a cosine function has no effect, (e.g. \(\cos (-x)\) = \(\cos x\)).</p>
<p><b>(5)</b>\(\quad\)Changing the sign within a sine function changes the sign of the function, (e.g. \(\sin (-x)\) = \(-\sin x\)).</p>
<p><b>(6)</b>\(\quad\)Placing a constant or altering the constant \(b\) in \(\sin (ax + b)\) or \(\cos (ax + b)\) changes the phase and shifts the sine or cosine function along the \(x\)-axis.</p><br>
<p></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200821125154+BLEtxY</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220214123429+motGxh</text></tag>
    </tags>
    
</question>
</quiz>