<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24990 -->
  <question type="stack">
    <name>
				<text>W4 Q03 Lagrange multiplier</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>Find the maximum and minimum values of \[f({@v1@},{@v2@})={@f@}\] subject to \[ {@C@}.\] Write the Lagrange function in the form \(f({@v1@},{@v2@}) - \lambda g({@v1@},{@v2@})\).</p>
<p><b>(a)</b> For which values of \({@v1@}\) do critical points of \( f- \lambda g\) occur? </p><p>\( x = \) [[input:ansx]] [[validation:ansx]]</p>
<p><b>(b)</b> The maximum/minimum of \(f({@v1@},{@v2@})={@f@}\) subject to \( {@C@}\) is<br></p><p>\(\mbox{max} = \) [[input:ansmax]] [[validation:ansmax]]</p><p>\(\mbox{min} = \) [[input:ansmin]] [[validation:ansmin]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>To find the maximum and minimum values of 
\[f({@v1@},{@v2@})={@f@}\] 
subject to \( {@C@}\) we use the method of Lagrange multipliers.</p>
<p>In general we have the function \(f({@v1@},{@v2@})\) and the constraint \(g({@v1@},{@v2@}) = {@lhs(C)-rhs(C)@}=0\).</p>
<p>We write \[ \phi({@v1@},{@v2@},\lambda) = f({@v1@},{@v2@})+\lambda\, g({@v1@},{@v2@}) = {@f@}+ \lambda ({@(lhs(C)-rhs(C))@})^{**}\]
and then solve
\[ \frac{\partial \phi}{\partial{@v1@}} = {@DP1@}=0\]
\[ \frac{\partial \phi}{\partial{@v2@}} = {@DP2@}=0\]
\[ \frac{\partial \phi}{\partial {@l@}} = {@DPl@}=0\]
This gives us three equations in three unknowns, \( {@v1@},{@v2@},{@l@}\), which can be solved to get:
\[ {@ssol@}.\]
Evaluating \(f\) at each of these points gives respective values as \( \left[{@tamax@}, {@tamin@} \right] \).&nbsp;</p><p>Hence the minimum value is {@tamin@} and the maximum value is {@tamax@}.</p><p>\( ** \) Because we are using the variable \(l\) as part of the Maxima code we replace \( \lambda \) with \(l.\)</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>1</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2019121800</text>
		</stackversion>
    <questionvariables>
				<text>v1:x;
v2:y;
a: rand(5)+3;
b: rand_with_prohib(2,5,[a]);
orderless(l,v1,v2);
f: v1*v2;
C: x^2/a+y^2/b=1;
/*  Problem setup above this line.  */
P1:f-l*(lhs(C)-rhs(C));
DP1:diff(P1,v1);
DP2:diff(P1,v2);
DPl:diff(P1,l);
H1:diff(f,v1,2)*diff(f,v2,2)-(diff(f,v1,1,v2,1))^2;
ssol:solve([DP1,DP2,DPl],[v1,v2,l]);
/* Extract the real solutions only */
ssol:sublist(ssol,lambda([ex],all_listp(real_numberp,maplist(rhs,ex))));
/* returns the list of critical points*/
extvr:maplist(rhs,flatten(maplist(lambda([ex],sublist(ex,lambda([ex2], lhs(ex2)=v1))),ssol)));
extv:setify(extvr);  /* x coordinates of critical points*/

/* returns the list of y values*/
extyr:maplist(rhs,flatten(maplist(lambda([ex],sublist(ex,lambda([ex2], lhs(ex2)=v2))),ssol)));
exty:setify(extyr);  /* y coordinates of critical points*/

f(v1,v2):=v1*v2;
x1: lmax(extvr);
x2: lmin(extvr);
y1: lmax(extyr);
y2: lmin(extyr);
tamax: f(x1, y1);
tamin: f(x1,y2);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prtx]]
[[feedback:prtmax]]
[[feedback:prtmin]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>\(f({@v1@},{@v2@})={@f@}\) subject to \( {@C@}\) so \( x= {@extv@}\), max val = {@tamax@}, min val = {@tamin@}</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ansmax</name>
      <type>algebraic</type>
      <tans>tamax</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansmin</name>
      <type>algebraic</type>
      <tans>tamin</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansx</name>
      <type>algebraic</type>
      <tans>extv</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>{?,?,....}</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prtmax</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ansmax</sans>
        <tans>tamax</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtmax-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtmax-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtmin</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ansmin</sans>
        <tans>tamin</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtmin-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtmin-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtx</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text>ansxL: listify(ansx)
ansxMin: lmin(ansxL)
ansxMax: lmax(ansxL)
extvL: listify(extv)
extvMin: lmin(extvL)
extvMax: lmax(extvL)</text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>length(ansxL)</sans>
        <tans>2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtx-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtx-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ansxMin</sans>
        <tans>extvMin</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prtx-2-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prtx-2-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ansxMax</sans>
        <tans>extvMax</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtx-3-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtx-3-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:moodle2.maths.ed.ac.uk+190308114732+frGoqm</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+200109191224+3GMVPA</text></tag>
    </tags>
    <deployedseed>1269450543</deployedseed>
    <deployedseed>1728583457</deployedseed>
    <deployedseed>88787539</deployedseed>
    <deployedseed>2062507803</deployedseed>
    <deployedseed>945085450</deployedseed>
    <deployedseed>164482871</deployedseed>
    <deployedseed>2114267236</deployedseed>
    <deployedseed>1460114613</deployedseed>
    <deployedseed>1383914030</deployedseed>
    <deployedseed>1563535883</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ansmax</name>
        <value>tamax</value>
      </testinput>
      <testinput>
        <name>ansmin</name>
        <value>tamin</value>
      </testinput>
      <testinput>
        <name>ansx</name>
        <value>extv</value>
      </testinput>
      <expected>
        <name>prtmax</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtmax-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtmin</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtmin-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtx</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtx-3-T</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>