<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26735 -->
  <question type="stack">
    <name>
				<text>10.2.2.3 Example-Polar form</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_example">Example</h4>
<p>Find the polar coordinate form of</p>
<ol class="HELM_parts_inline">
    <li><span class="math inline">\(z=3+4i\)</span> </li>
    <li> <span class="math inline">\(z=-3-i\)</span></li>
</ol>
<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
    <li>
        <p> Here <span class="math inline">\(r=|z|=\sqrt{3^2+4^2}=\sqrt{25}=5\qquad \theta=\text{arg}(z)=\tan^{-1}({4\over 3})=0.93 \mathrm{rad}\)</span> so that<span class="math inline">\(z=5(\cos 0.93 \mathrm{rad}+i\sin 0.93 \mathrm{rad}).\)</span></p>
    </li>
    <li>
        <p>Here <span class="math inline">\(r=|z|=\sqrt{(-3)^2+(-1)^2}=\sqrt{10}\approx 3.16\qquad
\theta=\text{arg}(z)=\tan^{-1}{(-1)\over {(-3)}}\)</span> It is natural to assume that <span class="math inline">\(\tan^{-1}\dfrac{(-1)}{(-3)}=\tan^{-1}\left(\dfrac{1}{3}\right)\)</span>. </p>

        <p>Using this value on your calculator (unless it is very sophisticated) you will obtain a value of about <span class="math inline">\(18.43^{\circ}\), which is approximately \(0.321\ \mathrm{radians}\),</span>&nbsp;for <span class="math inline">\({\tan^{-1}({1\over 3})}\)</span>.
            This is incorrect since if we use the Argand diagram to plot <span class="math inline">\(z=-3-i\)</span> we get:&nbsp;</p>
    </li>
    <p></p>

    <br>
    <figure class="HELM_figure"><img src="@@PLUGINFILE@@/10-2-fig4-example-polarform.svg" alt="" role="presentation" class="img-responsive atto_image_button_middle" width="383" height="150">&nbsp;</figure>
    <br></ol>
<ol class="HELM_parts">
    <p>The angle <span class="math inline">\(\theta\)</span> is clearly <span class="math inline">\(-180^{\circ}+18.43^{\circ}=-161.57^{\circ}=-2.82 \mathrm{rad}\)</span>. This example warns us to take care when determining arg<span class="math inline">\((z)\)</span>        purely using algebra. You will always find it helpful to construct the Argand diagram to locate the particular quadrant into which your complex number is pointing. <span><strong>Your calculator cannot do this for you</strong></span>. Finally,
        in this example, <span class="math inline">\(z=3.16\ (\cos -2.82 \mathrm{rad}+i \sin -2.82 \mathrm{rad})\)</span>.</p>
</ol>
<p></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200629123208+vuOk45</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210701151237+qzPzhH</text></tag>
    </tags>
    
</question>
</quiz>