<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26527 -->
  <question type="stack">
    <name>
				<text>6.2.2.6 (not in quiz) Compute sinh(x+y) knowing sinhx, sinhy</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise&nbsp;</p>
<p> Let \( \sinh x = {@xx@} \), and \( \sinh y = {@yy@} \), calculate \(\sinh(x+y)\).</p>
<p>Note: Write down a formula for \( \sinh(x+y) \) in terms of \( \sinh x, \sinh y \), and only sub in numbers and use a calculator at the final step. Round up your answer and enter 4 decimal places. If the answer has less than 4 decimal places, then enter zero(s) at the end.</p>
<p> \( \sinh(x+y) = \) [[input:ans1]] [[validation:ans1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>Using the fundamental identity \( \cosh^2x - \sinh^2x = 1\), we get \(\cosh^2 x = 1 + \sinh^2 x\). Similarly, \(\cosh^2 y = 1 + \sinh^2 y\).</p>
<p>Next, take the (positive) square root to obtain \( \cosh x\), noting that \( \cosh x &gt; 0 \) for all value of \(x\). Similar for \( \cosh y \).</p>
<p>Then use the formula \( \sinh(x+y) = \sinh x \cosh y + \cosh x \sinh y \), we obtain
\[\begin{align*} \sinh(x+y) &amp; = \sinh x \cosh y + \cosh x \sinh y \\
&amp; = \sinh x \cdot \sqrt{1 + \sinh^2 y} + \sqrt{1 + \sinh^2 x} \cdot \sinh y \\
&amp; = ({@xx@}) \cdot \sqrt{1 + ({@yy@})^2} + \sqrt{1 +({@xx@})^2} \cdot ({@yy@}) \\
&amp; = {@ta@} \text{ after rounding up.} \end{align*}\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020070100</text>
		</stackversion>
    <questionvariables>
				<text>aa: -(rand(9)+1)/10;
bb: (rand(9)+11)/10;
xx: decimalplaces(float(sinh(aa)),4);
yy: decimalplaces(float(sinh(bb)),4);
uu: sqrt(1+xx*xx);
vv: sqrt(1+yy*yy);
sinhsum: xx*vv+uu*yy;
ta: decimalplaces(float(sinhsum),4);
</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>\( \sinh x = {@xx@} \), and \( \sinh y = {@yy@} \)
\[\cosh(x+y) = {@ta@} \]</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>sinh,cosh,tanh,coth,sech,cosech,[[BASIC-ALGEBRA]]</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200625214557+DAyGvK</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200716111730+fgHuCL</text></tag>
    </tags>
    <deployedseed>91144347</deployedseed>
    <deployedseed>1061835517</deployedseed>
    <deployedseed>1910134799</deployedseed>
    <deployedseed>2068373406</deployedseed>
    <deployedseed>1143371224</deployedseed>
    <deployedseed>2053469723</deployedseed>
    <deployedseed>1434403637</deployedseed>
    <deployedseed>651001782</deployedseed>
    <deployedseed>798957904</deployedseed>
    <deployedseed>772423529</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>0.75</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(ta+0.005,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>