<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 38397 -->
  <question type="stack">
    <name>
				<text>1.4.1.8 Cancelling Common Factors</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_example">Example</h4>
<p>Simplify \(\dfrac{6(4-8x)(x-2)}{1-2x}\)</p>

<h4 class="HELM_solution">Solution</h4>
<p>The factor \(4-8x\) can be factorised to \(4(1-2x)\). Thus</p>
<p>\(\qquad \dfrac{6(4-8x)(x-2)}{1-2x} = \dfrac{(6)(4)(1-2x)(x-2)}{1-2x} = 24(x-2)\)</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210914134011+CyGaEi</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220117172845+UXvY4Y</text></tag>
    </tags>
    
</question>
</quiz>