<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24574 -->
  <question type="stack">
    <name>
				<text>15.2.1.3 Intuition behind moments</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>Intuition tells us that a large moment corresponds to a large turning effect. A mass placed 8 metres from the origin has a smaller turning effect than the same mass placed 10 metres from the origin. This is, of course, is why it is easier to rock a see-saw by pushing it at a point further from the pivot. Our intuition also tells us the side of the pivot on which the masses are placed is important. Those placed to the left of the pivot have a turning effect opposite to those placed to the right.</p>

<p>For a collection of masses, the moment of the total mass located at the centre of mass is equal to the sum of the moments of the individual masses. This definition enables us to calculate the position of the centre of mass. It is conventional to label the \(x\) coordinate of the centre of mass as \(\bar{x}\), pronounced ‘x bar’.</p>
<div class="HELM_keypoint">
<h4>Key Point</h4>
<p>The <strong>moment</strong> of the total mass located at the centre of mass is equal to the sum of the moments of the individual masses.</p>
</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200818055919+0QQ5Cr</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200909092437+tdvsCZ</text></tag>
    </tags>
    
</question>
</quiz>