<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24302 -->
  <question type="stack">
    <name>
				<text>5.1.1.17 DELETE</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>When the price of a luxury consumer item is £\({@a@}\), a market survey reveals that the demand is {@b@} items per year. However another survey has shown that at a price of £\({@c@}\), the demand for the item is {@d@} items per year. Assuming that both surveys are valid, find a linear function that relates demand \(Q\) to price \(P\). What demand would be predicted by the linear function at a price of £\({@e@}\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]and a price of £500? [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<p>Comment on the validity of both predictions.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>The linear function will be of the form:</p>
<p><em>Q = aP + b </em>    ({@c@}≤<em>P</em>≤{@a@})</p>
<p>The limits on P represent the given range of data on price. Substituting the  values of <em>Q</em> and <em>P</em> to calculate the gradient:</p>
<p>\(\dfrac{{@d@} - {@b@}}{{@c@} - {@a@}}\)<p> = {@m@}</p>
This can be substituted into the equation <em>Q = aP + b </em>, using one of the pairs of data, eg.</p>
<p> {@b@} = {@m@} * {@a@} + <em>b</em></p>
<p>This can be rearranged to give:</p>
<p> <em>b<em> = {@y@}</p>
<p>Therefore the equation of the function is:</p>
<p><em>Q</em> = {@m@} * <em>P</em> + {@y@}</p>
<p>A precautionary check is to make sure that this result is consistent with the other pairs of values. </p>
<p><em>Q</em> = {@m@} * {@c@} + {@y@} = {@d@}</p>
<p>This is the same as quoted in the question therefore all calculations are correct up to this point.</p>
To find the demand with a price of £{@e@} we can substitute the values in.</p>
<p>Q = {@m@} * {@e@} + {@y@} = {@ta1@}</p>
<p>This value is within the two values the market survey found, therefore interpolation is being used, this would suggest that this is a valid prediction. This comes with a caveat however such as the assumption it is a linear function. </p>
<p>If we were to use this model further to predict the demand with a price of £500, we would have to use extreme caution as it is outside our study parameters and therefore we would be extrapolating. 
<p>Q = {@m@} * 500 + {@y@} = {@ta2@}</p>
<p>If the given data points are close to each other then interpolation between these points is more reliable than extrapolation to points further away.]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021120900</text>
		</stackversion>
    <questionvariables>
				<text>a: rand([1000,1050,1100,1150])
b: rand([90000,95000,100000,105000])
c: rand([550,600,650,700])
d: rand([190000,195000,200000,205000,210000])
e: rand([750,800,850,900])
m: (d-b)/(c-a)
y: b-m.a
f: 500
ta1:m.e+y
ta2: m.f+y</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>a</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200801174416+DIABxl</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220224125600+hBnRij</text></tag>
    </tags>
    
</question>
</quiz>