<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25970 -->
  <question type="stack">
    <name>
				<text>19.2.5.2 Solve an exact equation</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_example">Example</h4>
<p> Solve the equation \[ x^3 \dfrac{\mathrm{d}y}{\mathrm{d}x}+3x^2y=x\]</p>
<h4 class="HELM_solution">Solution</h4>
<p>Comparing this equation with the form in the Key Point before we see that \(f(x)=x^3\) and \(g(x)=x\). Hence the equation can be written \[ \dfrac{\mathrm{d}}{\mathrm{d}x}(yx^3)=x,\] which has solution \[yx^3=\int x\, \mathrm{d}x=\dfrac{1}{2}x^2+C.\]Therefore \[ y=\dfrac{1}{2x}+\dfrac{C}{x^3}.\]
    </p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:moodle2.maths.ed.ac.uk+210308154806+T5i6vd</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210614120408+3lkm64</text></tag>
    </tags>
    
</question>
</quiz>