<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27129 -->
  <question type="stack">
    <name>
				<text>36.2.1.4 (not used for Edinburgh) steam and leaf section: task1</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Consider the following rounded data:</p><p></p><table><thead><tr><th scope="col"><span><span><span style="font-weight: normal;">{@L[1]@}</span>&nbsp; &nbsp; &nbsp; </span></span><br></th><th scope="col"><span style="font-weight: normal;">{@L[3]@}</span>&nbsp; &nbsp; &nbsp; <br></th><th scope="col"><span style="font-weight: normal;"><span></span>{@L[6]@} &nbsp; &nbsp;&nbsp; </span><br></th><th scope="col"><span style="font-weight: normal;">{@L[8]@}</span> &nbsp;&nbsp;&nbsp;&nbsp; <br></th><th scope="col"><span style="font-weight: normal;">{@L[11]@}</span> &nbsp; &nbsp;&nbsp; <br></th><th scope="col"><span style="font-weight: normal;">{@L[10]@}</span> &nbsp;&nbsp; <br></th><th scope="col"><span style="font-weight: normal;">{@L[5]@}</span> &nbsp;&nbsp;&nbsp; <br></th><th scope="col"><span style="font-weight: normal;">{@L[14]@}</span> &nbsp;&nbsp;&nbsp; <br></th><th scope="col"><span style="font-weight: normal;">{@L[15]@}</span> &nbsp;&nbsp;&nbsp; <br></th><th scope="col"><span style="font-weight: normal;">{@L[16]@}</span> &nbsp;&nbsp; <br></th></tr><tr><td>{@L[12]@} &nbsp; &nbsp; <br></td><td>{@L[2]@}&nbsp; &nbsp;&nbsp;&nbsp;&nbsp; <br></td><td>{@L[13]@} &nbsp; <br></td><td>{@L[17]@} &nbsp;&nbsp;&nbsp; <br></td><td>{@L[18]@}&nbsp; &nbsp;&nbsp; <br></td><td>{@L[7]@} &nbsp;&nbsp; &nbsp; <br></td><td>{@L[4]@} &nbsp;&nbsp;&nbsp;&nbsp; <br></td><td>{@L[19]@} &nbsp;&nbsp;&nbsp; <br></td><td>{@L[9]@} &nbsp;&nbsp;&nbsp; <br></td><td>{@L[50]@} &nbsp;&nbsp; <br></td></tr><tr><td>{@L[20]@}&nbsp; &nbsp; &nbsp; <br></td><td>{@L[27]@} &nbsp; &nbsp;&nbsp;&nbsp; <br></td><td>{@L[33]@} &nbsp;&nbsp;&nbsp; <br></td><td>{@L[41]@} &nbsp;&nbsp;&nbsp; <br></td><td>{@L[26]@} &nbsp; &nbsp;&nbsp; <br></td><td>{@L[42]@} &nbsp;&nbsp;&nbsp;&nbsp; <br></td><td>{@L[25]@} &nbsp;&nbsp;&nbsp;&nbsp; <br></td><td>{@L[39]@} &nbsp;&nbsp;&nbsp; <br></td><td>{@L[46]@} &nbsp;&nbsp;&nbsp; <br></td><td>{@L[40]@} &nbsp;&nbsp; <br></td></tr><tr><td>{@L[45]@} &nbsp;&nbsp;&nbsp;&nbsp; <br></td><td>{@L[34]@} &nbsp; &nbsp; <br></td><td>{@L[44]@} &nbsp;&nbsp;&nbsp; <br></td><td>{@L[21]@} &nbsp;&nbsp;&nbsp;&nbsp; <br></td><td>{@L[32]@}&nbsp; &nbsp;&nbsp; <br></td><td>{@L[31]@} &nbsp;&nbsp;&nbsp;&nbsp; <br></td><td>{@L[47]@} &nbsp;&nbsp;&nbsp;&nbsp; <br></td><td>{@L[49]@} &nbsp;&nbsp;&nbsp; <br></td><td>{@L[30]@} &nbsp;&nbsp; <br></td><td>{@L[38]@} &nbsp;&nbsp; <br></td></tr><tr><td>{@L[35]@}&nbsp; &nbsp; &nbsp; <br></td><td>{@L[28]@} &nbsp; &nbsp; &nbsp; <br></td><td>{@L[36]@} &nbsp;&nbsp;&nbsp; <br></td><td>{@L[48]@} &nbsp;&nbsp;&nbsp;&nbsp; <br></td><td>{@L[22]@} &nbsp;&nbsp; <br></td><td>{@L[23]@} &nbsp;&nbsp;&nbsp;&nbsp; <br></td><td>{@L[37]@}&nbsp; <br></td><td>{@L[24]@} &nbsp;&nbsp; <br></td><td>{@L[43]@} &nbsp; <br></td><td>{@L[29]@} &nbsp;&nbsp;&nbsp; <br></td></tr></thead></table><p></p><p><br></p>

<ol class="HELM_parts">
<li><p>Use the summary to check the data for normality.</p></li>

<br>
<center>
<table class="HELM_table">
<tbody>
<tr style="border-bottom: 1px solid #000;"><td style="text-align: center;"><b>Lowest Value</b></td><td>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</td></tr>
<tr style="border-bottom: 1px solid #000;"><td style="text-align: center;"><b>Lower Hinge</b></td><td>[[input:ans2]] [[validation:ans2]][[feedback:prt2]]</td></tr>
<tr style="border-bottom: 1px solid #000;"><td style="text-align: center;"><b>Low-Spread</b></td><td>[[input:ans3]] [[validation:ans3]][[feedback:prt3]]</td></tr>
<tr style="border-bottom: 1px solid #000;"><td style="text-align: center;"><b>Hinge to Extreme</b></td><td>[[input:ans4]] [[validation:ans4]][[feedback:prt4]]</td></tr>
<tr style="border-bottom: 1px solid #000;"><td style="text-align: center;"><b>Median</b></td><td>[[input:ans5]] [[validation:ans5]][[feedback:prt5]]</td></tr>
<tr style="border-bottom: 1px solid #000;"><td style="text-align: center;"><b>Median to Lower Hinge</b></td><td>[[input:ans6]] [[validation:ans6]][[feedback:prt6]]</td></tr>
<tr style="border-bottom: 1px solid #000;"><td style="text-align: center;"><b>Median to Upper Hinge</b></td><td>[[input:ans7]] [[validation:ans7]][[feedback:prt7]]</td></tr>
<tr style="border-bottom: 1px solid #000;"><td style="text-align: center;"><b>Upper Hinge</b></td><td>[[input:ans8]] [[validation:ans8]][[feedback:prt8]]</td></tr>
<tr style="border-bottom: 1px solid #000;"><td style="text-align: center;"><b>High-Spread</b></td><td>[[input:ans9]] [[validation:ans9]][[feedback:prt9]]</td></tr>
<tr style="border-bottom: 1px solid #000;"><td style="text-align: center;"><b>Hinge to Extreme</b></td><td>[[input:ans10]] [[validation:ans10]][[feedback:prt10]]</td></tr>
<tr><td style="text-align: center;"><b>Highest Value</b></td><td>[[input:ans11]] [[validation:ans11]][[feedback:prt11]]</td></tr>
</tbody><tbody>
</tbody></table>
</center>
<br><br>

<li><p>Is the distribution of the data symmetric? [[input:ans12]] [[validation:ans12]][[feedback:prt12]]</p></li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[a) Data for normality<br><ul><li>The lowest value is \({@L[1]@}\)</li></ul>
<ul><li>The Lower Hinge is the 25th percentile: \(({@L[12]@}+{@L[13]@})/2\)<br></li></ul>
<ul><li>The Low-Spread is the difference between the lowest value and the median: \(({@L[25]@}+{@L[26]@})/2-{@L[1]@}\)<br></li></ul>
<ul><li>The Hinge to Extreme is the difference Lower Hinge-Lowest Value:\(({@L[12]@}+{@L[13]@})/2-{@L[1]@}\)<br></li></ul>
<ul><li>The Median is the middle number of the distribution: \(({@L[25]@}+{@L[26]@})/2\)<br></li></ul>
<ul><li>The Median to Lower Hinge is the difference  Median-Lower Hinge: \(({@L[25]@}+{@L[26]@})/2-({@L[12]@}+{@L[13]@})/2\)<br></li></ul>
<ul><li>The Median to Upper Hinge is the difference  Upper Hinge-Median: \(({@L[37]@}+{@L[38]@})/2-({@L[25]@}+{@L[26]@})/2\)<br></li></ul>
<ul><li>The Upper Hinge is the 75th percentile: \(({@L[37]@}+{@L[38]@})/2\)<br></li></ul>
<ul><li>The High-Spread is the difference between the highest value and the median: \({@L[50]@}-({@L[25]@}+{@L[26]@})/2\)<br></li></ul>
<ul><li>The Hinge to Extreme is the difference Highest value-Upper Hinge: \({@L[50]@}]-({@L[37]@}+{@L[38]@})/2\)<br></li></ul>
<ul><li>The highest value is \({@L[50]@}\)<br></li></ul>

<br>
b) <ul><li>While there are no hard-and-fast rules for comparing figures such as those obtained here, many authors suggest that the figures should be within \(10%\) of each other before normality can be assumed. This is clearly not the case here. We conclude that the distribution of data being investigated is not symmetrical. In fact, the figures above suggest that the distribution is skewed to the left, a fact supported by the steam-and-leaf diagram of the same data to be found above. Remember that the term 'skewness' refers to the location of the 'tail' of a distribution.</li></ul>

<img src="@@PLUGINFILE@@/36.2.3%20Skewness%201.svg" alt="Right skew on the left and left skew on the right are represented" class="img-responsive atto_image_button_text-bottom" width="345" height="69">]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020070100</text>
		</stackversion>
    <questionvariables>
				<text>L:[rand([9,10,11]),12,rand([12,13,14]),13,rand([14,15,16]),17,rand([17,18,19]), 20, rand([21,22,23]), rand([24,25,26]),26,rand([27,28]), 29, rand([31,32,33]), 34, rand([34,35,36]), rand([38,39]), 40, 41, rand([41,42]), rand([41,42]), 41, 42, rand([41,42]), rand([42,43]),43,44,44,rand([44,45]), rand([44,45]), 45,45,rand([45,46]), rand([45,46]), 46, rand([46,47]), 47, 47, rand([47,48]), rand([47,48]), rand([47,48]), rand([48,49]), rand([48,49]), rand([48,49]),49,49, 50, 50, rand([50,51]),51];

ta1:L[1];
ta2:(L[12]+L[13])/2;
ta3:(L[25]+L[26])/2-L[1];
ta4:(L[12]+L[13])/2-L[1];
ta5:(L[25]+L[26])/2;
ta6:(L[25]+L[26])/2-(L[12]+L[13])/2;
ta7:(L[37]+L[38])/2-(L[25]+L[26])/2;
ta8:(L[37]+L[38])/2;
ta9:L[50]-(L[25]+L[26])/2;
ta10:L[50]-(L[37]+L[38])/2;
ta11:L[50];
ta12: [["Yes", false], ["No", true]];</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[<p><br></p><p><br></p>]]></text>
		</specificfeedback>
    <questionnote>
				<text>\(lowest value={@ta1@},  lower hinge={@ta2@},  low-spread={@ta3@},  hinge to extreme={@ta4@},  median={@ta5@},  median to lower hinge={@ta6@},  median to upper hinge={@ta7@}, upper hinge={@ta8@},  high-spread={@ta9@},  hinge to extreme={@ta10@},  highest value={@ta11@} \),  Is the distribution of the data symmetric? \({@ta12@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans10</name>
      <type>numerical</type>
      <tans>ta10</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans11</name>
      <type>numerical</type>
      <tans>ta11</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans12</name>
      <type>dropdown</type>
      <tans>ta12</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>ta4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans5</name>
      <type>numerical</type>
      <tans>ta5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans6</name>
      <type>numerical</type>
      <tans>ta6</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans7</name>
      <type>numerical</type>
      <tans>ta7</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans8</name>
      <type>numerical</type>
      <tans>ta8</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans9</name>
      <type>numerical</type>
      <tans>ta9</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt10</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans10</sans>
        <tans>ta10</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt11</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans11</sans>
        <tans>ta11</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt11-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt11-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt12</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>[ans12]</sans>
        <tans>mcq_correct(ta12)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt12-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt12-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans8</sans>
        <tans>ta8</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans9</sans>
        <tans>ta9</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200915102434+jfjXf6</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200915102435+c3PAMJ</text></tag>
    </tags>
    <deployedseed>163178386</deployedseed>
    <deployedseed>849336731</deployedseed>
    <deployedseed>1955953421</deployedseed>
    <deployedseed>1283784405</deployedseed>
    <deployedseed>171250866</deployedseed>
    <deployedseed>904552302</deployedseed>
    <deployedseed>1739178614</deployedseed>
    <deployedseed>173111493</deployedseed>
    <deployedseed>729280991</deployedseed>
    <deployedseed>1995177506</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>ta10</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>ta11</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>"No"</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ta8</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>ta9</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt10-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt11</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt11-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt12</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt12-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt9-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+5,simp)</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>ev(ta10+5,simp)</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>ev(ta11+5,simp)</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>"Yes"</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta2+5,simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ev(ta3+5,simp)</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ev(ta4+5,simp)</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ev(ta5+5,simp)</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ev(ta6+5,simp)</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ev(ta7+5,simp)</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ev(ta8+5,simp)</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>ev(ta9+5,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt10-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt11</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt11-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt12</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt12-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt5-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt6-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt7-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt8-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt9-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>