<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 96169 -->
  <question type="stack">
    <name>
				<text>23.7.2.1 Applying Fourier series to solve a differential equation</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>Applying Fourier series to solve a differential equation<br></h3>

<p>The following Task which is quite long will provide useful practice in applying Fourier series to a practical problem. Essentially you should follow Steps 1 to 3 above carefully.</p>

<hr>
<h4 class="HELM_example">Example</h4>
<p>The problem is to find the steady state response \(y(t)\) of a spring/mass/damper system modelled by</p>
<p>\(\qquad m\dfrac{d^2y}{dt^2} + c\dfrac{dy}{dt}+k y = F(t) \qquad \qquad \qquad\) (4)</p>
<p>where \(F(t)\) is the <strong>periodic square wave</strong> function shown in the diagram.<br></p>

<figure class="HELM_figure">
  <img src="@@PLUGINFILE@@/Figure9.svg" alt="Image of a periodic square wave of the function F(t). The maximum amplitude is F_0." width="536" height="149" class="img-responsive atto_image_button_text-bottom">
</figure>

<h4 class="HELM_solution">Solution</h4>
<p><strong>Step 1:</strong> Obtain the Fourier series of \(F(t)\) noting that it is an odd function.</p>
<p>The calculation is similar to those you have performed earlier in this Workbook.</p>
<p>Since \(F(t)\) is an odd function and has period \(2t_0\) so that \(\, \omega = \dfrac{2\pi}{2t_0} = \dfrac{\pi}{t_0}\), it has Fourier coefficients:</p>
<p>\(\begin{align} b_n &amp;= \frac{2}{t_0}\int_0^{t_o}F_0\sin\left(\frac{n\pi t}{t_0}\right)\,dt \qquad n=1,2,3,\dots \\ &amp;= \left( \frac{2F_0}{t_0} \right) \left( \frac{t_0}{n\pi} \right) \left[ -\cos \frac{n\pi t}{t_0} \right]_0^{t_0}\\ &amp;= \frac{2F_0}{n\pi}(1-\cos
    n\pi) = \left\{\begin{matrix}\dfrac{4F_0}{n\pi}&amp;n\ \text{ odd}\\ 0&amp;n\ \text{ even} \end{matrix}\right. \end{align}\)</p>

<p>so \( \quad F(t) = \dfrac{4F_0}{\pi}\sum_{n=1}^{\infty}\dfrac{\sin n\omega t}{n} \qquad\)(where the sum is over odd \(n\) only).</p>

<p><br></p>
<p><strong>Step 2(a):</strong> Since each term in the Fourier series is a sine term you must now solve (4) to find the steady state response \(y_n\) to the \(n^{\text{ th}}\) harmonic input: \(\quad F_n(t) = b_n\sin n\omega t \qquad n=1,3,5,\dots\)</p>
<p>From the basic theory of linear differential equations this response has the form</p>
<p>\( \qquad y_n = A_n\cos n\omega t + B_n\sin n\omega t \qquad \qquad\) (5)</p>
<p>where \(A_n\) and \(B_n\) are coefficients to be determined by substituting (5) into (4) with \(F(t)=F_n(t).\) Do this to obtain simultaneous equations for \(A_n\) and \(B_n :\)

</p>
<p>We have, differentiating (5),</p>
<p>\( \qquad y_n' = n\omega (-A_n\sin n\omega t+B_n\cos n\omega t)\)</p>
<p>\( \qquad y_n''=(n\omega )^2(-A_n\cos n\omega t-B_n\sin n\omega t) \)</p>
<p>from which, substituting into (4) and collecting terms in \(\cos n\omega t\) and \(\sin n\omega t,\) \[(-m(n\omega)^2A_n + cn\omega B_n+kA_n) \cos n\omega t + (-m(n\omega)^2B_n-cn\omega A_n + kB_n)\sin n\omega t = b_n\sin n\omega t \]</p>
<p>Then, by comparing coefficients of \( \cos n\omega t\) and \(\sin n\omega t\), we obtain the simultaneous equations:</p>
<p>\(\qquad (k-m(n\omega)^2)A_n + c(n\omega)B_n =0 \qquad \qquad \qquad \) (6)</p>
<p>\(\qquad -c(n\omega)A_n + (k-m(n\omega)^2)B_n =b_n \qquad \qquad \qquad \) (7)</p>

<p><br></p>
<p><strong>Step 2(b):</strong> Now solve (6) and (7) to obtain \(A_n\) and \(B_n:\)</p>
<p>\[A_n = -\frac{c\omega_nb_n}{(k-m\omega_n^2)^2 + \omega_n^2c^2} \qquad \qquad (8)\]</p>
<p>\[B_n = \frac{(k-m\omega_n^2)b_n}{(k-m\omega_n^2)^2 + \omega_n^2c^2}\qquad \qquad (9)\]</p>
<p>where we have written \(\omega_n\) for \(n\omega\) as the frequency of the \(n^{\text{ th}}\) harmonic.</p>

<p>It follows that the steady state response \(y_n\) to the \(n^{\text{ th}}\) harmonic of the Fourier series of the forcing function is given by (5). The amplitudes \(A_n\) and \(B_n\) are given by (8) and (9) respectively in terms of the systems parameters
    \(k,c,m\), the frequency \(\omega_n\) of the harmonic and its amplitude \(b_n\). In practice it is more convenient to represent \(y_n\) in the so-called <strong>amplitude</strong> form:</p>
<p>\(\qquad y_n = C_n\sin (\omega_nt +\phi_n) \qquad \qquad \qquad \) (10)</p>
<p>where, from (5) and (10),</p>
<p>\( \qquad A_n\cos\omega_nt + B_n\sin\omega_nt = C_n(\cos\phi_n\sin\omega_nt + \sin\phi_n\cos\omega_nt).\)</p>
<p>Hence</p>
<p>\(\qquad C_n\sin \phi_n = A_n \qquad C_n\cos \phi_n=B_n\)</p>
<p>so</p>
<p>\(\qquad \tan \phi_n = \dfrac{A_n}{B_n} = \dfrac{c\omega_n}{(m\omega_n^2-k)^2} \qquad \qquad \qquad \qquad\) (11)</p>
<p>\(\qquad C_n = \sqrt{A_n^2+B_n^2} = \dfrac{b_n}{\sqrt{(m\omega_n^2-k)^2 + \omega_n^2c^2}} \qquad \qquad \) (12)</p>

<p><br></p>
<p><strong>Step 3:</strong> Finally, use the superposition principle, to state the complete steady state response of the system to the periodic square wave forcing function:</p>
<p>\( \qquad y(t) = \sum_{n=1}^{\infty}y_n(t) = \sum_{{n=1}\atop (n \text{ odd})} C_n(\sin\omega_nt + \phi_n) \quad\) where \(C_n\) and \(\phi_n\) are given by (11) and (12).</p>
<p>In practice, since \(\, b_n=\dfrac{4F_0}{n\pi}\) it follows that the amplitude \(C_n\) also decreases as \(\, \frac{1}{n}\). However, if one of the harmonic frequencies say \(\omega_n'\) is close to the natural frequency \(\, \sqrt{\frac{k}{m}}\) of the
    undamped oscillator then that particular frequency harmonic will dominate in the steady state response. The particular value \(\omega_n'\) will, of course, depend on the values of the system parameters \(k\) and \(m.\)</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+220601180906+rFenQF</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220602124813+zvgde1</text></tag>
    </tags>
    
</question>
</quiz>