<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24541 -->
  <question type="stack">
    <name>
				<text>13.5.1.3 Substitution</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>

<p>Use suitable substitutions to find</p>

<ol class="HELM_parts">
<li>\({\displaystyle\int {@exp1@}\ \mathrm{d}x}=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]] </li>

<li>\({\displaystyle\int {@exp2@}\ \mathrm{d}t}=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]] </li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<ol class="HELM_parts">

    <li>
        <p>We need to calculate the integral \( I =\int {@exp1@}\ \mathrm{d}x .\)</p>
        <p>Using \( u = {@u1@}\) we get \(\dfrac{\mathrm{d}u}{\mathrm{d}x} = {@diff(u1,x)@} \Rightarrow \ \mathrm{d}x = \frac{1}{{@diff(u1,x)@}}\ \mathrm{d}u.\)</p>
        <p>\[\begin{aligned} I = \int {@exp1@}\ \mathrm{d}x = \int u^{{@bb@}}\ \mathrm{d}x = \int \dfrac{u^{{@bb@}}}{{@aa@}}\ \mathrm{d}u = {@ta1@}. \end{aligned}\]</p>
    </li>

    <li>
        <p>We need to calculate the integral \( I =\int {@exp2@}\ \mathrm{d}t .\)</p>
        <p>Using \( u = {@v1@}\) we get \(\dfrac{\mathrm{d}u}{\mathrm{d}x} = {@diff(v1,t)@} \Rightarrow \ \mathrm{d}t = \frac{1}{{@diff(v1,t)@}}\ \mathrm{d}u.\)</p>
        <p>\[\begin{aligned} I = \int {@exp2@}\ \mathrm{d}t = \int {@p1@} \sin u \, \mathrm{d}t &amp;=&nbsp; \int {@p1@} \sin u \, \frac{1}{{@diff(v1,t)@}}\ \mathrm{d}u, \\ &amp;= \int \dfrac{\sin (u)}{3}, \\ &amp;= {@ta2@}.\end{aligned}\]</p>
    </li>
</ol>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>aa: rand([2, 4, 6, 8]); 
bb: rand([6, 7, 8, 9]);

u1: aa*x + 1;
v1: t^3 + 1;
exp1: (aa*x + 1)^bb; 
p1: t^2;
exp2: (t^2)*sin(t^3 + 1);

ta1: int(exp1, x) + c; 
ta2: int(exp2, t) + c;</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\( \int {@exp1@}\mathrm{d}x= {@ta1@}, \quad  \int {@exp2@}\mathrm{d}t = {@ta2@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>int</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>int</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>0</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Int</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>t</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200804174229+b68pip</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210609133026+6uYnPD</text></tag>
    </tags>
    <deployedseed>1651997333</deployedseed>
    <deployedseed>26321727</deployedseed>
    <deployedseed>1024763884</deployedseed>
    <deployedseed>1657208556</deployedseed>
    <deployedseed>1390433546</deployedseed>
    <deployedseed>1394793454</deployedseed>
    <deployedseed>1026384303</deployedseed>
    <deployedseed>1014867547</deployedseed>
    <deployedseed>874207373</deployedseed>
    <deployedseed>2115982990</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>2*ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>2*ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>