<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24835 -->
  <question type="stack">
    <name>
				<text>L7.2 Q4 Improper integral 3</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>Compute the improper integral. If it convergent evaluate it, otherwise enter the value <code>none</code> as your answer. \[ I = \int_0^{\infty} {@p@} dx. \]
<p>\( I = \) [[input:SAns]][[validation:SAns]]</p>
<p>Give your answer in exact form.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[We have the integral: \[ I = \int_0^{\infty} {@p@} dx = \lim_{t \rightarrow \infty}&nbsp; &nbsp;\int_0^{t} {@p@} dx &nbsp;=&nbsp; \lim_{t \rightarrow \infty}&nbsp; \left[ {@int@}\right]_0^{t} =&nbsp; \lim_{t \rightarrow \infty}&nbsp; (e^{t^2}/2-1/2)&nbsp;\]<br><br>Note that to preform the integral \(\int_0^t xe^{x^2} dx\) we use the substitution \(u(x)=x^2\) so that \(du=2xdx\) and our integral becomes \[\int_0^t xe^{x^2} dx=\int_{0}^{t^2} \frac{e^{u}}{2} du = \left[\frac{e^{u}}{2}\right]_0^{t^2}=\left[\frac{e^{x^2}}{2}\right]_0^{t}.\]<br><p></p><p>The integral is divergent as \(\lim_{t \rightarrow \infty} e^{t^2}/2\) diverges to infinity.</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>1</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>p1: x
p2: exp(x^2)
p: p1*p2
ta : none
int: integrate(p,x)</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>\( I = \int_0^{\infty} {@p@} dx = {@ta@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>SAns</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>[[basic-calculus]],[[basic-algebra]]</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>SAns</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:maths-moodle.maths.ed.ac.uk+180222200854+ATe1TI</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+210302104251+sRQWsj</text></tag>
    </tags>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>SAns</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>SAns</name>
        <value>5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>