<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27123 -->
  <question type="stack">
    <name>
				<text>36.2.1.5 (for Edinburgh) The box-and-whisker diagram</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>The box-and-whisker diagram</h4><br>
<p>In order to visually summarise a data set, we can use a <b>box and whisker</b> plot. To draw a box-and-whisker diagram we first need to evaluate the of the five-number summaries. Consider the Inter-Party Competition data:</p><br>
<table style="width:75%">
<tbody>
<tr><td>9.5</td><td>11.9</td><td>20.0</td><td>33.4</td><td>40.1</td><td>50.0</td><td>12.7</td><td>21.0</td><td>33.6</td><td>40.6</td></tr>
<tr><td>50.0</td><td>15.5</td><td>26.4</td><td>35.4</td><td>41.1</td><td>50.0</td><td>17.7</td><td>37.9</td><td>41.3</td><td>50.0</td></tr>
<tr><td>41.9</td><td>50.4</td><td>43.0</td><td>43.3</td><td>43.6</td><td>43.7</td><td>43.8</td><td>44.7</td><td>44.9</td><td>45.0</td></tr>
<tr><td>45.1</td><td>45.2</td><td>45.3</td><td>45.5</td><td>46.1</td><td>46.5</td><td>46.6</td><td>47.1</td><td>48.0</td><td>48.2</td></tr>
<tr><td>48.5</td><td>48.4</td><td>48.6</td><td>48.7</td><td>48.8</td><td>48.9</td><td>49.4</td><td>49.5</td><td>49.6</td><td>49.8</td></tr>
</tbody>
</table><br>
<p>The five-number summaries are:
</p><ul>
<li>Minimum \(= 9.5\)</li>
<li>Lower Quartile \(= 40.1\)</li>
<li>Median \(= 45.05\)</li>
<li>Upper Quartile \(= 48.6\)</li>
<li>Maximum \(= 50.4\)</li>
</ul><p></p>
<p>The box-and-whisker diagram for this data is given below:</p><p style="text-align: center;"><img src="@@PLUGINFILE@@/36-2-fig2-boxplot.svg" alt="" width="647" height="136" class="img-responsive atto_image_button_text-bottom"><br></p>
<p>It is constructed as follows:</p>
<ol>
<li><b>The Box</b>
<ul>
	<li> The left-hand vertical is placed at the lower quartile (\(40.1\));</li>
	<li> The right-hand vertical is placed at the upper quartile (\(48.6\));</li>
	<li> The vertical in the box is placed at the median (\(45.05\)).</li></ul></li>
<li><b>The Whiskers</b>
<p>These dashed lines span out from the lower/upper quartiles towards the minimum/maximum value, but their lengths should not exceed \(1.5\) times the inter-quartile range. Notice that the inter-quartile range of the data is \(8.5\).</p>
<ul>
	<li>The furthest extend permitted for the upper whisker is \(48.6 + 1.5 \times 8.5 = 61.35\).</li> 
	<li>The upper whisker therefore extends to the largest value in the data that does not exceed this limit, the maximum value \(50.4\).</li>
	<li>The furtherest extent permitted for the lower whisker is \(40.1 - 1.5 \times 8.5 = 27.35\).</li> 
	<li>The lower whisker therefore extends to the smallest value in the data that does not exceed this limit, which is \(33.4\).</li>
</ul></li>
<li><b>The Outlying Values</b>
<p>Mark as a dot any value that lie beyond furthest reaches of the whiskers. In this example, there are no values greater than \(61.35\), but there are \(8\) values which are less than \(37.35\). </p></li></ol>
<p>Notice that about half of the data values lie in the box and that the tails show up well in the diagram. The diagram shows the left-skew (skewness refers to the tail) present in the data.</p>
<p>Note: There are several similar methods recommended by different authors for drawing box-and-whisker plots and so the methods recommended in statistical texts may vary a little from those given here.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200915102140+ZUFOF4</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200915102140+52u89q</text></tag>
    </tags>
    
</question>
</quiz>