<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24695 -->
  <question type="stack">
    <name>
				<text>18.4.2.1 Measuring the height of a building</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p></p>
<h3><span>Engineering example</span></h3>The height \(h\) of a building is estimated from
<p></p>
<ol type="i">
    <li>the known horizontal distance \(x\) between the point of observation \(M\) and the foot of the building and</li>
    <li>the elevation angle \(\theta\) between the horizontal and the line joining the point of observation to the top of the building (see Figure below).
    </li>
</ol>
<p>If the measured horizontal distance is \(x = 150 \text{m}\) and the elevation angle is \(\theta = 40^{\circ}\), estimate the error in measured building height due to an error of \(0.1^{\circ}\) degree in the measurement of the angle of elevation.</p>
<figure class="HELM_figure"><img src="@@PLUGINFILE@@/18-4-fig-10-EngEx1.svg" alt="18-4-fig-10-geometry.svg" width="350" height="176" class="img-responsive atto_image_button_text-bottom"><br></figure>
<p>The variables \(x, \theta\), and \(h\) are related by \[\tan \theta = \dfrac{h}{x}\] or \[x \tan \theta = h.\]</p>
<p>The error in \(h\) resulting from a measurement error in \(\theta\) can be deduced by differentiating \(x \tan \theta = h\) as follows. \[\dfrac{\mathrm{d}(x \tan \theta)}{\mathrm{d} \theta} = \dfrac{\mathrm{d} h}{\mathrm{d} \theta} \quad \implies \quad
    \tan \theta \dfrac{\mathrm{d} x}{\mathrm{d} \theta} + x \dfrac{\mathrm{d}(\tan \theta)}{\mathrm{d} \theta} = \dfrac{\mathrm{d} h}{\mathrm{d} \theta}.\] This can be written \[\tan \theta \dfrac{\mathrm{d} x}{\mathrm{d} \theta} + x \sec^2 \theta = \dfrac{\mathrm{d}
    h}{\mathrm{d} \theta}.\]</p>
<p>The equation above gives the relationship among the small variations in variables \(x, h\) and \(\theta\). Since \(x\) is assumed to be without error and independent of \(\theta\), \(\dfrac{\mathrm{d}x}{\mathrm{d} \theta} = 0\) and the equation becomes
    \[x \sec^2 \theta = \dfrac{\mathrm{d}h}{\mathrm{d} \theta}.\]</p>
<p>The above equation can be considered to relate the error in building height \(\delta h\) to the error in angle \(\delta \theta\): \[\dfrac{\delta h}{\delta \theta} \simeq x \sec^2 \theta.\]</p>
<p>It is given that \(x = 150 \text{m}\).</p>
<p>The incidence angle \(\theta = 40^{\circ}\) can be converted to radians i.e \(\theta = \dfrac{40 \pi}{180} \text{ rad} = \dfrac{2 \pi}{9} \text{ rad}\).</p>
<p>Then the error in angle \(\delta \theta = 0.1^{\circ}\) needs to be expressed in radians for consistency of the units.</p>
<p>So \(\delta \theta = \dfrac{0.1 \pi}{180} \text{ rad} = \dfrac{\pi}{1800} \text{ rad}\). Hence, \[\delta h = 150 \dfrac{\pi}{1800 \times \cos^2\left(\dfrac{2\pi}{9}\right)} \approx 0.45 \text{ m}. \] </p>
<p>So the error in building height resulting from an error in elevation angle of \(0.1^{\circ}\) is about \(0.45 \text{ m}\).</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201215022713+o2beBa</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210614102651+jlaIiH</text></tag>
    </tags>
    
</question>
</quiz>