<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26608 -->
  <question type="stack">
    <name>
				<text>8.3.2.2 Row echelon form</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>We say a matrix is in <b>row echelon form</b> if it has the shape resulting from Gaussian elimination, where Gaussian elimination has operated on the rows. Furthermore, we say a matrix is in <b>reduced row echelon form</b> if it satisfies the following conditions:</p><ol><li>It is in row echelon form.</li><li>The leading entry in each nonzero row is a \(1 \) (called a leading \( 1 \)).</li><li>Each column containing a leading \( 1 \) has zeros in all its other entries.</li></ol>
<hr>
<h4 class="HELM_example">Example</h4>
<p> The matrix
<span class="math display">\[ \left [ \begin{matrix}
    2 &amp; -1 &amp; 2 &amp; 0 \\
    0 &amp; 1 &amp; 4 &amp; -2 \\
    0 &amp; 0 &amp; -2 &amp; 3 \\
\end{matrix} \right ] \]</span></p>
is in row echelon form, but not reduced row echelon form. The matrices
<p><span class="math display">\[\left [ \begin{matrix}
    1 &amp; 0 &amp; 0 &amp; -1 \\
    0 &amp; 1 &amp; 0 &amp; 2 \\
    0 &amp; 0 &amp; 1 &amp; -5\\
\end{matrix} \right], \hspace{3mm}
\left [ \begin{matrix}
    1 &amp; 3 &amp; 0 &amp; 0 \\
    0 &amp; 0 &amp; 1 &amp; 0 \\
    0 &amp; 0 &amp; 0 &amp; 1 \\
\end{matrix} \right]\]</span></p>
are in reduced row echelon form.]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200813084441+paUT7c</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200813100119+q2kQIS</text></tag>
    </tags>
    
</question>
</quiz>