<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 84679 -->
  <question type="stack">
    <name>
				<text>18.3.1.1 The stationary points of a function of two variables (copy)</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span>The stationary points of a function of two variables</span><br></h3>

<p>The Figure below shows a computer generated picture of the surface defined by the function \(z=x^3+y^3-3x-3y,\) where both \(x\) and \(y\) take values in the interval \([-1.8,1.8]\).

</p>
<p style="text-align: center;"><img src="@@PLUGINFILE@@/18.3%20Surface.svg" alt="The plot of the surface z = x^3 + y^3 - 3x - 3y in (x, y, z) coordinates" width="513" height="393" class="img-responsive atto_image_button_text-bottom"><br></p>

<p>There are four features of particular interest on the surface. At point \(A\) there is a <strong>local maximum</strong>, at \(B\) there is a <strong>local minimum</strong>, and at \(C\) and \(D\) there are what are known as <strong>saddle points</strong>.
    At \(A\) the surface is at its greatest height in the immediate neighbourhood. If we move on the surface from \(A\) we immediately loose height no matter in which direction we travel. At \(B\) the surface is at its least height in the neighbourhood.
    If we move on the surface from \(B\) we immediately gain height, no matter in which direction we travel. The features at \(C\) and \(D\) are quite different. In some directions as we move away from these points along the surface we lose height whilst
    in others we gain height. The similarity in shape to a horse’s saddle is evident. At each point \(P\) of a <em>smooth</em> surface one can draw a unique plane which touches the surface there. This plane is called the <strong>tangent plane</strong>    at \(P\). (The tangent plane is a natural generalisation of the tangent line which can be drawn at each point of a smooth curve.) In the Figure above, at each of the points \(A,B,C,D\) the tangent plane to the surface is horizontal at the point of
    interest. Such points are thus known as <strong>stationary points</strong> of the function. In the next subsections we show how to locate stationary points and how to determine their nature using partial differentiation of the function \(f(x,y)\).</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+220307124647+IjpQ8H</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220307124647+Wv0AwO</text></tag>
    </tags>
    
</question>
</quiz>