<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27359 -->
  <question type="stack">
    <name>
				<text>39.1.3.1 Probabilities and the Standard Normal Distribution</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>Probabilities and the Standard Normal Distribution</h3>

<p>Since the standard normal distribution is used so frequently, a table of values has been produced to help us calculate probabilities. This is located at the end of this quiz and at the end of this page. You should keep a copy open in another tab for later questions but for now you can refer to the copy below. It is based upon the following diagram: </p>
<center><img src="@@PLUGINFILE@@/39-1-fig-3-z1diagram.svg" alt="The normal distribution curve is shaded between 0 and the positive-valued z1." width="273" height="127" class="img-responsive atto_image_button_text-bottom"><br></center><p>Since the total area under the curve is equal to 1, it follows from the symmetry in the curve that the area under the curve in the region \(Z&gt;0\) is equal to \(0.5\). In the figure above, the shaded area is the probability that \(Z\) takes values between \(0\) and \(z_1\). When we ‘look-up’ a value in the table we obtain the value of the shaded area.</p> 

<p> Note from the Table that as \(Z\) increases from \(0\), the entries increase, rapidly at first and then more slowly, toward \(5000\) i.e. a probability of \(0.5\). This is consistent with the shape of the curve.</p>
<p>After \(Z=3\) the increase is quite slow so that we tabulate entries for values of \(Z\) rising by increments of \(0.1\) instead of \(0.01\), as in the rest of the Table.</p><p>Note: There are different versions of the table of standard normal probabilities depending on
the interval of interest. The table used in this section give the values of \(\text{P}(0&lt;Z&lt;z)\) for \(z \geq 0\), this is the area shaded in the figure above. Another common table of standard probabilities are of \(\text{P}(Z&lt;z)\) for \(z \geq 0\), also known as the table of standard normal <i>cumulative </i>probabilities. If graphed, you would see everything left of 0 also shaded. The relationship between the two tables is that the cumulative probabilities are \(0.5\) greater than the table included here since it includes all of the left half of the distribution. Pay attention to which distribution a table is based upon.<br></p>

<hr>
<h4 class="HELM_example">Example</h4>
<p>What is the probability that \(Z\) takes values between \(0\) and \(1.9\)? (Refer to the table of standard normal probabilities at the end of the quiz.)</p>
<h4 class="HELM_solution">Solution</h4>
<p>The row beginning ‘\(1.9\)’ and the column headed ‘\(0\)’ is the appropriate choice and its entry is \(4713\). This is to be read as \(0.4713\) (we omitted the ‘\(0.\)’ in each entry for clarity.) The interpretation is that the probability that \(Z\) takes values between \(0\) and \(1.9\) is \(0.4713\).</p>

<hr>
<h4 class="HELM_example">Example</h4>
<p>What is the probability that \(Z\) takes values between \(0\) and \(1.96\)?</p> 
<h4 class="HELM_solution">Solution</h4>
<p>This time we want the row beginning \(1.9\) and the column headed \('6'\).</p>
<p>The entry is \(4750\) so that the required probability is \(0.4750\).</p>

<hr>
<h4 class="HELM_example">Example</h4>
<p>What is the probability that \(Z\) takes values between \(0\) and \(1.965\)?</p> 
<h4 class="HELM_solution">Solution</h4>
<p>There is no entry corresponding to \(1.965\), so we take the average of the values for \(1.96\) and \(1.97\). (This linear interpolation is not strictly correct but is acceptable.)</p>
<p>The two values are \(4750\) and \(4756\) with an average of \(4753\). Hence the required probability is \(0.4753\).</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200728100933+C0Y5Fb</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200831132842+zH8Pn8</text></tag>
    </tags>
    
</question>
</quiz>