<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 33707 -->
  <question type="stack">
    <name>
				<text>19.2.9.7 (NOT IN QUIZ) LR circuit ODE</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>In an LR circuit with applied voltage \(E = {@g@}\), the current \(i\) is given by</p>
\[L\frac{\mathrm{d}i}{\mathrm{d}t} + Ri = {@g@}.\]
<p>If the initial current is \(i_0\), find \(i\) subsequently.</p>
<p>\(i(t) = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>This is a linear differential equation that is not in standard form.</p>
\[L\frac{\mathrm{d}i}{\mathrm{d}t} + Ri = {@g@}\]
<p>We must first convert it to standard form by dividing both sides by \(L\).
Note that here \(f(t) = R/L\) and \(\displaystyle g(t) = {@g / L@}\).</p>
\[\frac{\mathrm{d}i}{\mathrm{d}t} + \frac{R}{L}i = {@g / L@}.\]
<p>We use the key point to find the integrating factor:</p>
\[\text{I.F.} = \exp \left[ \int f(t) \,\mathrm{d}t \right] = {@IF@}\]
<p>Next we transform the equation into exact form by multiplying both sides with the
    integrating factor. It becomes:
</p>
\[{@IF@} \cdot \frac{\mathrm{d}i}{\mathrm{d}t} + \frac{R}{L} {@IF@} \cdot i = {@IF@} \cdot \left({@g / L@}\right)\]
<p>Note that the left-hand side is the derivative of \({@IF@} \cdot i\) with respect to \(t\).
    Our equation then becomes:
</p>
\[\frac{\mathrm{d}}{\mathrm{d}t} \left( {@IF@} \cdot i \right) = {@IF * g / L@}\]
<p>Therefore, we can find an expression for \({@IF@} \cdot i\) by integrating the right-hand side.
    The result is:
</p>
\[{@IF@} \cdot i = {@I@}\]
<p>Which then gives the following general solution:</p>
\[i = {@gsol@}\]
<p>We can now solve for \(C\) and use the given boundary conditions:</p>
\[C = {@Cexp@}\]
\[i(0) = i_0 \implies C = {@C0@}\]
<p>This gives the particular solution we seek:</p>
\[i(t) = {@ta@}\]]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021120900</text>
		</stackversion>
    <questionvariables>
				<text>/* Randomisation */
l: 0.1;
a: 10;

g: a * (1 - e^(-l * t));
IF: e^(int(R/L, t));
I: ratsimp(int(expand(IF * g / L), t)) + C;
gsol: expand(I / IF);
Cexp: rhs(solve(i = gsol, C)[1]);
C0: float(ev(Cexp, t=0, i=i0));
ta: ev(gsol, C=C0);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(\theta' + {@l@} \theta = {@g@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210617102557+K8KV5f</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220315150319+I7r0I1</text></tag>
    </tags>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>