<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26222 -->
  <question type="stack">
    <name>
				<text>7.2.4.3 3x3 identity and zero matrix</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>The \(3\times 3\) unit matrix is: \[ I = \begin{bmatrix} 1 &amp; 0 &amp; 0 \\ 0 &amp; 1 &amp; 0 \\ 0 &amp; 0 &amp; 1 \end{bmatrix} \]  and as in the \(2\times 2\) case this has the property that \( AI = IA = A \) for all matrices \(A\).<br></p>
<p>The \(3\times 3\) zero matrix is \[ \underline{0} = \begin{bmatrix} 0 &amp; 0 &amp; 0 \\ 0 &amp; 0 &amp; 0 \\ 0 &amp; 0 &amp; 0 \end{bmatrix} \]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200622095451+UqKi35</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200907142136+yQ5olt</text></tag>
    </tags>
    
</question>
</quiz>