<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24522 -->
  <question type="stack">
    <name>
				<text>13.3.3.1 Curves above and below x-axis</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span>The area bounded by a curve, not entirely above the x-axis</span><br></h3>

<p>The following figure shows a graph of \(y=-x^2+1\).</p>
<p style="text-align: center;"><img src="@@PLUGINFILE@@/13.3.3.1%20Area%20below%20axis.svg" alt="Area between x-axis and y = -x^2 + 1 between 1 and 2" width="371" height="181" class="img-responsive atto_image_button_text-bottom"><br></p>

<p>The shaded area is bounded by the \(x\)-axis and the curve, but lies entirely below the \(x\)-axis.&lt;/p&gt;</p>
<p> Let us evaluate the integral \({\displaystyle\int_1^2( -x^2+1)\ \mathrm{d}x}\).</p>

<p>\[\begin{aligned} {\displaystyle\int_1^2( -x^2+1)\ \mathrm{d}x} = \left[ -\dfrac{x^3}{3}+x\right]_1^2 &amp;= \left(-\dfrac{2^3}{3} + 2 \right) - \left(-\dfrac{1^3}{3} + 1 \right), \\ &amp;= -\dfrac{7}{3} + 1, \\ &amp;= -\dfrac{4}{3}.&nbsp; \end{aligned}\]</p>

<p>The evaluation of the area yields a negative quantity. There is, of course, no such thing as a negative area. The area is actually \(\dfrac{4}{3}\), and the negative sign is an indication that the area lies below the \(x\)-axis. (However, in applications
    of integration such as work/energy or distance travelled in a given direction negative values can be meaningful.)</p>

<p> If an area contains parts both above and below the horizontal axis, care must be taken when calculating this area. It is necessary to determine which parts of the graph lie above the horizontal axis and which lie below. Separate integrals need to be calculated
    for each ‘piece’ of the graph. This idea is illustrated in the next Example.</p>

<hr>
<h4 class="HELM_example">Example</h4>
<p>Find the total area enclosed by the curve \(y=x^3-5x^2+4x\) and the \(x\)-axis between \(x=0\) and \(x=3\).</p>
<h4 class="HELM_solution">Solution</h4>
<p>We need to determine which parts of the graph lie above and which lie below the \(x\)-axis. To do this it is helpful to consider where the graph cuts the \(x\)-axis. So we consider the function \(x^3-5x^2+4x\) and look for its zeros</p>
<p></p>
\[x^3-5x^2+4x = x(x^2-5x+4) = x(x-1)(x-4)\]<p></p>
<p>So the graph cuts the \(x\)-axis when \(x=0\), \(x=1\) and \(x=4\). Also, when \(x\) is large and positive, \(y\) is large and positive since the term involving \(x^3\) dominates. When \(x\) is large and negative, \(y\) is large and negative for the same
    reason. With this information we can sketch a graph showing the required area:</p>

<p style="text-align: center;"><img src="@@PLUGINFILE@@/13.3.3.1%20area%20above%20and%20below.svg" alt="the area between the x-axis and the function y = x^3 -5x^2+1 between zero and four" width="407" height="175" class="img-responsive atto_image_button_text-bottom"><br></p>

<p>From the graph we see that the required area lies partly above the \(x\)-axis (when \(0\le x\le 1\)) and partly below (when \(1\le x\le 3\)). So we evaluate the integral in two parts:</p>

<p>Firstly:</p>
\[{\displaystyle\int_0^1( x^3-5x^2+4x )\ \mathrm{d}x} =&nbsp; \left[\dfrac{x^4}{4} - \dfrac{5x^3}{3} + \dfrac{4x^2}{2}\right]_0^1 = \left(\dfrac{1}{4} - \dfrac{5}{3} + 2\right) - \left(0\right) = \dfrac{7}{12}\]<p></p>
<p>This is the part of the required area which lies above the \(x\)-axis.</p>
<p>Secondly:</p>

\[\begin{aligned} {\displaystyle\int_1^3( x^3-5x^2+4x )\ \mathrm{d}x} &amp;= \left[\dfrac{x^4}{4} - \dfrac{5x^3}{3} + \dfrac{4x^2}{2}\right]_1^3, \\ &amp;= \left(\dfrac{81}{4} - \dfrac{135}{3} + 18\right) - \left(\dfrac{1}{4} - \dfrac{5}{3} + 2\right), \\ &amp;= -\dfrac{22}{3}. \end{aligned}\]<p></p>
<p>This represents the part of the required area which lies below the \(x\)-axis.</p>
<p>The actual area is \(\dfrac{22}{3}\).</p>

<p>Combining the results of the two separate calculations we can find the total area bounded by the curve:</p>

<p> area = \(\dfrac{7}{12} + \dfrac{22}{3} = \dfrac{95}{12}\)</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200731093233+N9MqFO</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210713110953+XYjuya</text></tag>
    </tags>
    
</question>
</quiz>