<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24393 -->
  <question type="stack">
    <name>
				<text>16.4.4 General power series</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Power series centered at \(x_0\)&nbsp;</h4>
<p>A general power series has the form \[ b_0 + b_1(x-x_0) + b_2(x-x_0)^2 + \cdots = \sum_{p=0}^\infty b_p(x-x_0)^p\] </p>

<p> Exactly the same considerations apply to this general power series as apply to the ‘special’ series \(\sum_{p=0}^\infty b_p x^p\) except that the variable \(x\) is replaced by \((x-x_0).\) The radius of convergence of the general series is obtained in
    the same way: \[R = \lim_{p \to \infty} \left|\dfrac{b_p}{b_{p+1}}\right|\] and the interval of convergence is now shifted to have centre at \(x = x_0\) (see Figure below). The series is absolutely convergent if \(|x − x_0| &lt; R\), diverges if \(|x
    − x_0| &gt; R\) and may or may not converge if \(|x − x_0| = R\).</p>
<figure class="HELM_figure"><img src="@@PLUGINFILE@@/16-4-fig-11-Convergence.svg" alt="Interval of convergence" width="456" height="55" class="img-responsive atto_image_button_text-bottom"></figure>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+201021005546+ftjK7Y</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210610151542+3yD1Va</text></tag>
    </tags>
    
</question>
</quiz>