<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26725 -->
  <question type="stack">
    <name>
				<text>10.1.1.4 What is a Complex Number?</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<div class="HELM_keypoint">
<h4>Key Point </h4>
<p><span class="math display">\[\text{If}\qquad z=a+i b\qquad\text{then}\qquad 
\text{Re}(z)=a\quad \text{and}\quad \text{Im}(z)=b\]</span> Both the real and imaginary parts of a complex number are <span><strong>real</strong></span>.</p>
</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200625155227+sqvmyV</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210706164434+JY5KyI</text></tag>
    </tags>
    
</question>
</quiz>