<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 37767 -->
  <question type="stack">
    <name>
				<text>1.1.4.T3 BODMAS</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Evaluate the expression below<p>
<p>\(({@a1@} - {@a2@}) \times {@a3@} =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>The BODMAS rule tells us the bracketed quantity must be evaluated first. Thus</p>
<p>\(\qquad ({@a1@} - {@a2@}) \times {@a3@} = {@a1-a2@} \times {@a3@}\)</p>
<p>We then perform multiplication to give the final answer</p>
<p>\(\qquad {@a1-a2@} \times {@a3@} = {@ta@}\)</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021072200</text>
		</stackversion>
    <questionvariables>
				<text>a1: rand([4,5])
a2: rand([1,2,3])
a3: rand([6,7,8])

ta: (a1-a2)*a3</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(({@a1@} - {@a2@}) \times {@a3@} = {@ta@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>+,-,*,/</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210903120420+NMTzYe</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+211208114934+ir13Dv</text></tag>
    </tags>
    <deployedseed>969950068</deployedseed>
    <deployedseed>109304746</deployedseed>
    <deployedseed>264094293</deployedseed>
    <deployedseed>842409138</deployedseed>
    <deployedseed>1697729283</deployedseed>
    <deployedseed>1723946207</deployedseed>
    <deployedseed>49393869</deployedseed>
    <deployedseed>1994992272</deployedseed>
    <deployedseed>1971007240</deployedseed>
    <deployedseed>808783923</deployedseed>
    <deployedseed>29262661</deployedseed>
    <deployedseed>814608701</deployedseed>
    <deployedseed>1580506603</deployedseed>
    <deployedseed>260308716</deployedseed>
    <deployedseed>1122629017</deployedseed>
    <deployedseed>2134585183</deployedseed>
    <deployedseed>1067041953</deployedseed>
    <deployedseed>534519828</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>(a1-a2)*a3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore></expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>