<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24642 -->
  <question type="stack">
    <name>
				<text>18.2.2.2 Mixed partial derivatives 1</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Find the partial derivatives \(\dfrac{\partial^2f}{\partial x^2},\ \ \dfrac{\partial^2f}{\partial y^2},\ \ \dfrac{\partial^2f}{\partial x\partial y},\ \ \dfrac{\partial^2f}{\partial y\partial x}\) for the function \(f(x,y)={@exp1@}\).</p>
<ol class="HELM_parts">
<p><br></p>
<li>\(\dfrac{\partial^2f}{\partial x^2} =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</li>
<li>\(\dfrac{\partial^2f}{\partial y^2} =\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</li>
<li>\(\dfrac{\partial^2f}{\partial x\partial y} =\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</li>
<li>\(\dfrac{\partial^2f}{\partial y\partial x} =\) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</li></ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>We are given the function \[ f(x,y)={@exp1@} \]</p>
<p>and we are aske to find the partial derivatives \(\dfrac{\partial^2f}{\partial x^2},\ \ \dfrac{\partial^2f}{\partial y^2},\ \ \dfrac{\partial^2f}{\partial x\partial y},\ \ \dfrac{\partial^2f}{\partial y\partial x}.\)</p>
<p>First we find: \[\dfrac{\partial f}{\partial x} ={@diff(exp1,x)@} \quad \text{and} \quad \dfrac{\partial f}{\partial x} ={@diff(exp1,y)@}\]</p>
<p>So:</p>
<ol class="HELM_parts">
    <p><br></p>
    <li>\(\dfrac{\partial^2f}{\partial x^2} = \dfrac{\partial f}{\partial x} \left( \dfrac{\partial f}{\partial x} \right) = \dfrac{\partial f}{\partial x} \left( {@diff(exp1,x)@} \right) = {@ta1@}\)</li>
    <li>\(\dfrac{\partial^2f}{\partial y^2} = \dfrac{\partial f}{\partial y} \left( \dfrac{\partial f}{\partial y} \right) = \dfrac{\partial f}{\partial y} \left( {@diff(exp1,y)@} \right) = {@ta2@}\)</li>
    <li>\(\dfrac{\partial^2f}{\partial x\partial y} = \dfrac{\partial f}{\partial x} \left( \dfrac{\partial f}{\partial y} \right) = \dfrac{\partial f}{\partial x} \left( {@diff(exp1,y)@} \right) = {@ta3@}\)</li>
    <li>\(\dfrac{\partial^2f}{\partial y\partial x} = \dfrac{\partial f}{\partial y} \left( \dfrac{\partial f}{\partial x} \right) = \dfrac{\partial f}{\partial y} \left( {@diff(exp1,x)@} \right) = {@ta4@}\)</li>
</ol>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>1</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>aa: rand([3, 5, 7, 9]);
bb: rand([2, 4, 6, 8]);
cc: rand(6)+1;

exp1: x+bb*y+aa;

ta1: diff(diff(exp1,x), x);
ta2: diff(diff(exp1,y), y);
ta3: diff(diff(exp1,y), x);
ta4: diff(diff(exp1,x), y);
</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\[ f(x,y)={@exp1@} \]
\[ f_{xx}={@ta1@}, \quad f_{yy}={@ta2@}, \quad f_{yx}={@ta3@}, \quad f_{xy}={@ta4@}\]</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>1</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>0</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>y</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Diff</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>y</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200819143718+S8stkj</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+210302105950+BBFReP</text></tag>
    </tags>
    <deployedseed>1005171239</deployedseed>
    <deployedseed>1946734222</deployedseed>
    <deployedseed>1652720503</deployedseed>
    <deployedseed>198285632</deployedseed>
    <deployedseed>1840193599</deployedseed>
    <deployedseed>664855915</deployedseed>
    <deployedseed>1097876627</deployedseed>
    <deployedseed>134804661</deployedseed>
    <deployedseed>1956485934</deployedseed>
    <deployedseed>2131864177</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta2+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ev(ta3+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ev(ta4+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>