<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24599 -->
  <question type="stack">
    <name>
				<text>14.2.1.1 Average value of a function</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span>Average value of a function</span></h3><br>

<p>Suppose a time-varying function \(f(t)\) is defined on the interval \(a\le t\le b\). The area, \(A\), under the graph of \(f(t)\) is given by the integral \(\displaystyle A=\int_a^b f(t)\,\mathrm{d}t\). This is illustrated in the Figure below.</p><p style="text-align: center;"><img src="@@PLUGINFILE@@/14-2-fig-5-areaundercurve.svg" alt="Left: the area under the curve f(t) from x=a to x=b. Right: a rectangle from x=a to x=b with height m and with equal area." width="588" height="136" class="img-responsive atto_image_button_text-bottom"><br></p>

<p>On the Figure we have also drawn a rectangle with base spanning the interval \(a\le t\le b\) and which has the same area as that under the curve. Suppose the height of the rectangle is \(m\). Then \[\text{area of rectangle} = \text{area under curve} \Rightarrow\quad
    m(b-a) = \int_a^b f(t)\,\mathrm{d}t\quad \Rightarrow\quad m = \frac{1}{b-a}\int_a^b f(t)\,\mathrm{d}t\] The value of \(m\) is the <strong>mean value</strong> of the function across the interval \(a\le t\le b\).</p><br>

<div class="HELM_keypoint">
    <h4>Key Point</h4>
    <p>The <strong>mean value</strong> of a function \(f(t)\) in the intervall \(a\le t\le b\) is \(\displaystyle \frac{1}{b-a}\int_a^b f(t)\,\mathrm{d}t\).</p>
</div><br>

<p>The mean value depends upon the interval chosen. If the values of \(a\) or \(b\) are changed, then the mean value of the function across the interval from \(a\) to \(b\) will in general change as well.</p>
<p>&nbsp;</p>
<hr>
<h4 class="HELM_example">Example</h4>
<p>Find the mean value of \(f(t)=t^2\) over the interval \(1\le t\le 3\).</p><br>

<h4 class="HELM_solution">Solution</h4>
<p>Using the Key Point above with \(a=1\) and \(b=3\) and \(f(t)=t^2\) \[\text{mean value}~=~\frac{1}{b-a}\int_a^b f(t)\,\mathrm{d}t ~=~ \frac{1}{3-1}\int_1^3 t^2 \,\mathrm{d}t ~=~ \frac{1}{2} \left[\frac{t^3}{3}\right]_1^3 ~=~ \frac{13}{3}\]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201127181759+jlO1BQ</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210607140913+EzyDAY</text></tag>
    </tags>
    
</question>
</quiz>