<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62197 -->
  <question type="stack">
    <name>
				<text>23.6.2b Revision of the exponential form of a complex number</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[Rewriting in complex exponential form, we have
\[ a_n \cos n \theta + b_n \sin n\theta = \dfrac{a_n}{2} ( e^{in\theta} + e^{-in\theta} ) + \dfrac{b_n}{2i} ( e^{in\theta} - e^{-in\theta} ) \]
Now collecting the terms in \(e^{ in\theta}\) and in \(e^{ -in\theta}\)  and use the fact that \( \dfrac{1}{i} = -i \), we have
\[ \dfrac{1}{2} \left( a_n - ib_n \right) e^{ in\theta} + \dfrac{1}{2} \left( a_n + ib_n \right) e^{ -in\theta} \]
This can be expressed more concisely by defining 
<p></p>
\( c_n = \dfrac{1}{2} (a_n - ib_n)  \) which has complex conjugate \( \displaystyle{c_n^{*}} = \dfrac{1}{2} (a_n + ib_n)  \) giving us
\[ a_n \cos n \theta + b_n \sin n\theta = c_n e^{i n \theta} + c_n^{*} e^{-i n \theta} \]]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200919080621+cQ7woX</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200919081147+cVF4iT</text></tag>
    </tags>
    
</question>
</quiz>