<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25719 -->
  <question type="stack">
    <name>
				<text>2.4.1.3 Continuous and discontinuous functions and limits (cont.)</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>When defining a discontinuous function algebraically it is often necessary to give different function rules for different values of <span class="math inline">\(x\)</span>. Consider, for example, the function defined as: <span class="math inline">\(\qquad f(x)= 
\left\{
\begin{array}{ll}
3 &amp; x&lt;0\\
x^2 &amp; x\ge 0 \, . \\
\end{array}
\right.\)</span></p><p>Notice that there is one rule for when <span class="math inline">\(x\)</span> is less than 0 and another rule for when <span class="math inline">\(x\)</span> is greater than or equal to 0.</p><p>A graph of this function is shown in the figure below.</p><p style="text-align: center;"><span class="math display"><img src="@@PLUGINFILE@@/block6-fig2-1.svg" alt="A discontinuous function" class="img-responsive atto_image_button_text-bottom" width="220" height="191"></span></p>
<p style="text-align: center;"><span><strong>Figure</strong></span>: An example of a discontinuous function.</p>
<p>Suppose we ask ‘to what value does <span class="math inline">\(y\)</span> approach as <span class="math inline">\(x\)</span> approaches 0?’. From the graph we see that as <span class="math inline">\(x\)</span> gets nearer and nearer to 0, the value of <span class="math inline">\(y\)</span> gets nearer to 0, if we approach from the right-hand side. We write this formally as</p><p><span class="math inline">\(\qquad \lim_{x\rightarrow 0^+} f(x) = 0\)</span></p><p>and say ‘the limit of <span class="math inline">\(f(x)\)</span> as <span class="math inline">\(x\)</span> tends to 0 from above is 0.’</p><p>On the other hand if <span class="math inline">\(x\)</span> gets closer to zero, from the left-hand side, the value of <span class="math inline">\(y\)</span> remains at 3. In this case we write</p><p><span class="math inline">\(\qquad \lim_{x\rightarrow 0^-}f(x) = 3\)</span></p><p>and say ‘the limit of <span class="math inline">\(f(x)\)</span> as <span class="math inline">\(x\)</span> tends to 0 from below is 3.’</p><p>In this example the right-hand limit and the left-hand limit are not equal, and this is indicative of the fact that the function is discontinuous.</p>
<p>In general a function is continuous at a point <span class="math inline">\(x=a\)</span> if the left-hand and right-hand limits are the same there and are finite, and if both of these are equal to the value of the function at that point. That is</p>
<div class="HELM_keypoint">
<h4>Key Point</h4>
<p>A function \(f(x)\) is continuous at \(x=a\) if and only if: <span class="math display">\[\qquad \lim_{x\rightarrow a^+} f(x) = \lim_{x\rightarrow a^-} f(x) = f(a)\]</span></p>
</div>
<p>If the right-hand and left-hand limits are the same, we can simply describe this common limit as <span class="math inline">\(\lim_{x\rightarrow a} f(x)\)</span>. If the limits are not the same we say the limit of the function does not exist at <span class="math inline">\(x=a\)</span>.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200618154558+V9YaNX</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200812114618+APR7uF</text></tag>
    </tags>
    
</question>
</quiz>