<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62156 -->
  <question type="stack">
    <name>
				<text>43.1.3.4 Adjusted Coeffs of determination</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>The adjusted coefficient of determination</h3>

Denoted (often) by \(R^2_{adj}\), the Adjusted Coefficient of Determination is defined as
\[ R^2_{adj} = 1 - \dfrac{SS_E/(n-p)}{SS_T/(n-1)}, \]
where \(p\) is the number of variables in the regression equation. For the simple linear model, \(p = 2\) since we have two unknown parameters in the regression equation, the intercept \(c\) and the coefficient \(m\) of \(x\). It can be shown that:<p></p><p></p><ol><li> \(R^2_{adj}\) is a better indicator of the adequacy of predictive power than \(R^2\) since it takes into account the number of regressor variables used in the model;</li><li>  \(R^2_{adj}\) does not necessarily increase when a new regressor variable is added.</li></ol><p>Both coefficients claim to measure the adequacy of the predictive power of a regression model and their values indicate the proportion of variability explained by the model. For example a value of
\[ R^2 \text{   or   } R^2_{adj} = 0.9751 \]
may be interpreted as indicating that a model explains \(97.51\%\) of the variability it describes. For
example, the drum and cable example considered previously gives the results outlined below with
\[ R^2 = 96.2 \text{   and   } R^2_{adj} = 0.958. \]</p><p><br></p>

In general, \(R^2_{adj}\) is (perhaps) more useful than \(R^2\) for comparing alternative models. In the context of a simple linear model, \(R^2\) is easier to interpret. In the drum and cable example we would claim that the linear model explains some \(96.2\%\) of the variation it describes.<table style="height: 388px;" border="1" width="469">
<tbody>
<tr>
<td style="text-align: center; width: 59.3333px;"><b>Drum &amp; Cable \( (x) \)</b></td>
<td style="text-align: center; width: 59.3333px;"><b>\( (x^2) \)</b></td>
<td style="text-align: center; width: 60.6667px;"><b>Cable Length \( (y)\)</b></td>
<td style="text-align: center; width: 60.6667px;"><b>\( y^2 \)</b></td>
<td style="text-align: center; width: 60.6667px;"><b>\( xy \)</b></td>
<td style="text-align: center; width: 60.6667px;"><b>Predicted Values</b></td>
<td style="text-align: center; width: 60.6667px;"><b>Error Squares</b></td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">30</td>
<td style="text-align: center; width: 59.3333px;">900</td>
<td style="text-align: center; width: 60.6667px;">70</td>
<td style="text-align: center; width: 60.6667px;">4900</td>
<td style="text-align: center; width: 60.6667px;">2100</td>
<td style="text-align: center; width: 60.6667px;">70</td>
<td style="text-align: center; width: 60.6667px;">0.00</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">40&nbsp;</td>
<td style="text-align: center; width: 59.3333px;">1600&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">90&nbsp;&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">8100</td>
<td style="text-align: center; width: 60.6667px;">3600&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">100&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">100.00</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">40&nbsp;</td>
<td style="text-align: center; width: 59.3333px;">1600&nbsp; &nbsp;</td>
<td style="text-align: center; width: 60.6667px;">100</td>
<td style="text-align: center; width: 60.6667px;">10000</td>
<td style="text-align: center; width: 60.6667px;">4000&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">100&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">0.00</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">50&nbsp;</td>
<td style="text-align: center; width: 59.3333px;">2500&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">120&nbsp;&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">14400</td>
<td style="text-align: center; width: 60.6667px;">6000&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">130&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">100.00</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">50&nbsp;&nbsp;</td>
<td style="text-align: center; width: 59.3333px;">2500&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">130</td>
<td style="text-align: center; width: 60.6667px;">16900&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">6500&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">130&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">0.00</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">50&nbsp; &nbsp;</td>
<td style="text-align: center; width: 59.3333px;">2500&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">150</td>
<td style="text-align: center; width: 60.6667px;">22500</td>
<td style="text-align: center; width: 60.6667px;">7500&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">130&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">400.00</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">60&nbsp;</td>
<td style="text-align: center; width: 59.3333px;">3600&nbsp;&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">160</td>
<td style="text-align: center; width: 60.6667px;">25600&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">9600&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">160&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">0.00</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">70&nbsp; &nbsp;</td>
<td style="text-align: center; width: 59.3333px;">4900&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">190</td>
<td style="text-align: center; width: 60.6667px;">36100</td>
<td style="text-align: center; width: 60.6667px;">1330&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">190&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">0.00</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">70&nbsp;</td>
<td style="text-align: center; width: 59.3333px;">4900&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">200&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">40000&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">14000&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">190&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">100.00</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">80&nbsp;&nbsp;</td>
<td style="text-align: center; width: 59.3333px;">6400&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">200&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">40000</td>
<td style="text-align: center; width: 60.6667px;">16000&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">220&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">400.00</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">80&nbsp;</td>
<td style="text-align: center; width: 59.3333px;">6400&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">220&nbsp;&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">48400&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">17600</td>
<td style="text-align: center; width: 60.6667px;">220&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">0.00</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">80&nbsp;</td>
<td style="text-align: center; width: 59.3333px;">6400&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">230&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">52900&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">18400&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">220&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">100.00</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">Sum of \(x\)&nbsp;</td>
<td style="text-align: center; width: 59.3333px;">Sum of \(x^2\)&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">Sum of \(y\)&nbsp;&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">Sum of \(y^2\)</td>
<td style="text-align: center; width: 60.6667px;">Sum of \(xy\)&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">\(SS_E \)</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">\(= 700\)&nbsp; &nbsp;&nbsp;</td>
<td style="text-align: center; width: 59.3333px;">\(= 44200\)</td>
<td style="text-align: center; width: 60.6667px;">\(= 1860\)</td>
<td style="text-align: center; width: 60.6667px;">\(= 319800\)&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">\(= 118600\)</td>
<td style="text-align: center; width: 60.6667px;">&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">\( = 1200.00\)</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">&nbsp;</td>
<td style="text-align: center; width: 59.3333px;">&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">&nbsp;</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">\( b = 3\)</td>
<td style="text-align: center; width: 59.3333px;">\( a = -20\)&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">\(SS_E = \)&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">\( R^2 = \)&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">\(R^2_{adj} = \)</td>
</tr>
<tr>
<td style="text-align: center; width: 59.3333px;">&nbsp;</td>
<td style="text-align: center; width: 59.3333px;">&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">\(31500\)</td>
<td style="text-align: center; width: 60.6667px;">&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">\(0.962\)</td>
<td style="text-align: center; width: 60.6667px;">&nbsp;</td>
<td style="text-align: center; width: 60.6667px;">\(0.958\)</td>
</tr>
</tbody>
</table>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200704131331+pNcB6H</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200715053951+X5UgQY</text></tag>
    </tags>
    
</question>
</quiz>