<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24568 -->
  <question type="stack">
    <name>
				<text>15.2.1.5 Computing COM for multiple point masses</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_example">Example</h4>
<p>Obtain an equation for the location of the centre of mass of two objects of masses \(m_1\) and \(m_2\): </p>
<ol class="HELM_parts">
    <li>located at distances \(x_1\) and \(x_2\) respectively, as shown part (a) of the Figure below.</li>
    <li>positioned on opposite sides of the origin as shown in part (b) of the Figure below.</li></ol>
<figure class="HELM_figure"><img src="@@PLUGINFILE@@/15-2-fig-5-CoM.svg" alt="15-2-fig-5.svg" width="600" height="140" class="img-responsive atto_image_button_text-bottom"><br></figure>
    <p>Referring to part (a) of the above Figure, we first write down an expression for the sum of the individual moments: The total mass is \(m_1 + m_2\) and the moment of the total mass is \((m_1 + m_2) \times \bar{x}\)</p>
    <p> The moment of the total mass is equal to the sum of the moments of the individual masses. The equation satisfied by \(\bar{x}\) is </p>
    <p> \((m_1+m_2)\bar{x} = m_1x_1+m_2x_2\), so \(\bar{x}=\dfrac{m_1x_1+m_2x_2}{m_1+m_2}\)</p>
    <p> However, this is precisely what would have been obtained if, when working out the moment of a mass, we use its <strong> coordinate </strong> (which takes account of sign) rather than using its <strong>distance</strong> from the origin. </p>
    <p>For the second case, as depicted in part (b) of the figure, the mass <span class="math inline">\(m_1\)</span> positioned on the left-hand side has a turning effect <span><strong>opposite</strong></span> to that of the mass <span class="math inline">\(m_2\)</span>        positioned on the right-hand side. To take account of this difference we use a minus sign when determining the moment of <span class="math inline">\(m_1\)</span> about the origin. This gives a total moment </p>
    <p><span class="math inline">\(\qquad -(m_1x_1)+(m_2x_2)\)</span>
    </p>
    <p>leading to
    </p>
    <p><span class="math inline">\(\qquad (-m_1x_1+m_2x_2)=(m_1+m_2)\bar{x}\qquad\)</span>so <span class="math inline">\(\qquad\bar{x}=\dfrac{-m_1x_1+m_2x_2}{m_1+m_2}\)</span>
    </p>
    <p> However, this is precisely what would have been obtained if, when working out the moment of a mass, we use its <span><strong>coordinate</strong></span> (which takes account of sign) rather than using its <span><strong>distance</strong></span> from
        the origin.
    </p>
    <hr>
    <p> The above formulae can be generalised very easily to deal with the general situation of \(n\) masses, \(m_1, m_2, \dots, m_n\) located at <strong>coordinate positions </strong> \(x_1,x_2,\dots,x_n\) and is given in the Key Point below.</p>
    <div class="HELM_keypoint">
        <h4>Key Point</h4>
        <p>The <strong>centre of mass</strong> of individual masses \(m_1, m_2, \dots, m_n\) located at positions \(x_1,x_2,\dots,x_n\) is \[\bar{x}=\dfrac{\sum_{i=1}^n m_i x_i}{\sum_{i=1}^n m_i}\]</p>
    </div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200817104619+9E9xZv</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210610154726+H9ren7</text></tag>
    </tags>
    
</question>
</quiz>