<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26330 -->
  <question type="stack">
    <name>
				<text>2.3.3.3 Demonstrate one-to-one *****</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>By sketching a graph of \(y=3x-1\) show that this is a one-to-one
function.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200617155603+Aynsx8</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200622151517+HvVonp</text></tag>
    </tags>
    
</question>
</quiz>