<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25731 -->
  <question type="stack">
    <name>
				<text>2.4.5.3 Task: Comment on symmetries *****</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>The following diagrams shows&nbsp; graphs of several functions. They share a common property. Study the graphs and comment on any symmetry.</p><p><img src="@@PLUGINFILE@@/block6-fig7-1.svg" alt="block6-fig7" width="541" height="135" class="img-responsive atto_image_button_text-bottom"><br></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>There is rotational symmetry about the origin. That is, each curve, when rotated through \(180^{\circ}\), transforms into itself.</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200618182728+JDUoUS</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200618183904+3owLtc</text></tag>
    </tags>
    
</question>
</quiz>