<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27173 -->
  <question type="stack">
    <name>
				<text>38.3.3.3 Bus stop</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p><p></p><p>You arrive at the bus stop and the time (in minutes) it takes for the next bus to arrive follows an Exponential distribution with rate \(\displaystyle{@lam@}\). Calculate the following probabilities to \(4\) decimal places.</p><ol class="HELM_parts">

<li>  <p>You wait less than \(5\) minutes?</p>
<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p></li>

  <li><p>You wait longer than \(7\) minutes? <br></p>
 <p> [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p></li>

  <li><p>You wait for between \(5\) and \(7\) minutes?<br></p>
 <p>[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p></li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>Let \(T\) be the random variable for the wait time at the bus stop, then \(T \sim \text{Exp}(\lambda = {@lam@})\).</p>

<ol class="HELM_parts">
    <li><p>The probability of waiting less than \(5\) minutes is: \[\text{P}(T \leq 5) = F_T(5) = 1 - e^{-\lambda 5} = {@displaydp(ansA,4)@}~ (\text{to 4 d.p.}).\]</p>
  </li>
  
  <li><p>The probability of waiting more than \(7\) minutes is: \[\text{P}(T &gt; 7) = 1 - F_T(7) = e^{-\lambda 7} = {@displaydp(ansB,4)@}~ (\text{to 4 d.p.}).\]</p>
  </li>
  
  <li><p>The probability of waiting between \(5\) and \(7\) minutes is:</p>
    <p>\[\begin{align*}\text{P}(5 \leq T \leq 7) &amp;= F_T(7) - F_T(5)\\ &amp;= \left[1 - e^{-\lambda 7}\right] - \left[1 - e^{-\lambda 5}\right]\\ &amp;= e^{-\lambda 5} - e^{-\lambda 7}\\ &amp;= {@displaydp(ansC,4)@}~ (\text{to 4 d.p.}).\end{align*}\]</p>
    <p>An alternative approach to evaluating this probability can be achieved by noting that \(\text{P}(5 \leq T \leq 7) = 1 - \text{P}(T \leq 5) - \text{P}(T &gt; 7).\)</p>
  </li>
</ol>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021072200</text>
		</stackversion>
    <questionvariables>
				<text>lam: 1/(5+rand(5));
ansA: decimalplaces(float(cdf_exp(5,lam)),4);
ansB: decimalplaces(float(1 - cdf_exp(7,lam)),4);
ansC: decimalplaces(float(cdf_exp(7,lam) - cdf_exp(5,lam)),4);
mal_ansC: decimalplaces(float(1 - ansA - ansB),4);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(\lambda = {@lam@}\), \(\text{P}(T \leq 5) = {@displaydp(ansA,4)@}\),\(\text{P}(T > 7) = {@displaydp(ansB,4)@}\),\(\text{P}(5 \leq T \leq 7) = {@displaydp(ansC,4)@}\), round error?: \({@displaydp(mal_ansC,4)@}\)
</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>dispdp(ansA,4)</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>to 4 d.p.</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4, maxdp:4</options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>dispdp(ansB,4)</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>to 4 d.p.</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4, maxdp:4</options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>dispdp(ansC,4)</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>to 4 d.p.</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4, maxdp:4</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ansA</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>ansB</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumDecPlaces</answertest>
        <sans>ans3</sans>
        <tans>ansC</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>NumDecPlaces</answertest>
        <sans>ans3</sans>
        <tans>mal_ansC</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
					<text><p dir="ltr" style="text-align: left;">I suspect that you have incurred a rounding error. Do you calculations to a higher degree of accuracy before reporting to the required number of decimal places.</p></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200903131427+F1xDxZ</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+211018110913+xGqT9U</text></tag>
    </tags>
    <deployedseed>749319354</deployedseed>
    <deployedseed>441476864</deployedseed>
    <deployedseed>240228495</deployedseed>
    <deployedseed>1984113745</deployedseed>
    <deployedseed>1140930726</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>dispdp(ansA,4)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dispdp(ansB,4)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>dispdp(ansC,4)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>dispdp(ansB,4)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dispdp(ansC,4)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>dispdp(ansA,4)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>prt3-2-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>