<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25728 -->
  <question type="stack">
    <name>
				<text>2.4.4.4 Frequency from period</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>A periodic function has period \({\displaystyle{@float(pp)@}}\) seconds. How many times will the pattern in the graph repeat over an interval of {@tt@} seconds?</p>

<p>\( N = \) [[input:ans1]] [[validation:ans1]]<br></p><br>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>The periodic function has period \({\displaystyle{@float(pp)@}}\) seconds.</p><p>The frequency (the number of repetitions per time unit) is \( f = \frac{1}{T} = {@1/pp@}\, Hz.\)</p><p>So in the time interval of \( {@tt@} \,s\) the pattern repeats \( N = {@1/pp@} \cdot {@tt@} = {@Ans@}\) times.</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020070100</text>
		</stackversion>
    <questionvariables>
				<text>pp: 10^-(1+rand(3));
tt: rand([10,20,30,40,50,50,60,80]);

Ans: tt/pp;</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[<p>[[feedback:prt1]]<br></p>]]></text>
		</specificfeedback>
    <questionnote>
				<text>\(T ={@pp@}, t={@tt@}, N = {@Ans@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>Ans</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans1</sans>
        <tans>Ans</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200618174720+dsLSQf</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200812105855+xcDd5G</text></tag>
    </tags>
    <deployedseed>474454178</deployedseed>
    <deployedseed>2103682428</deployedseed>
    <deployedseed>26316965</deployedseed>
    <deployedseed>771985559</deployedseed>
    <deployedseed>748981128</deployedseed>
    <deployedseed>745917705</deployedseed>
    <deployedseed>398645102</deployedseed>
    <deployedseed>1867571400</deployedseed>
    <deployedseed>68311927</deployedseed>
    <deployedseed>1038670467</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>Ans</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>pp*tt</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>