<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26291 -->
  <question type="stack">
    <name>
				<text>22.3.2.14 Hermitian matrices</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Hermitian matrices</h4>
<p>In some applications, of which quantum mechanics is one, matrices with <strong>complex</strong> elements arise.</p>
<p>If \(A\) is such a matrix then the matrix \(\overline{A}^T\) is the <strong>conjugate transpose</strong> of \(A\), i.e. the complex conjugate of each element of \(A\) is taken as well as \(A\) being transposed. Thus if
\[ A=\left[\begin{array}{cc}2+\mathrm{i}&amp;2\\ 3\mathrm{i}&amp;5-2\mathrm{i}\end{array}\right]\]
then
\[ \overline{A}^T=\left[\begin{array}{cc}2-\mathrm{i}&amp;-3\mathrm{i}\\ 2&amp;5+2\mathrm{i}\end{array}\right].\]
An <strong>Hermitian</strong> matrix is one satisfying
\[ \overline{A}^T=A.\]
The matrix \(A\) above is clearly non-Hermitian. Indeed the most obvious feature of an Hermitian matrix is that its diagonal elements <strong>must</strong> be real. (Can you see why?) Thus
\[ A=\left[\begin{array}{cc}6&amp;4+\mathrm{i}\\ 4-\mathrm{i}&amp;-2\end{array}\right]\]
is Hermitian.</p>
<p>A \(3\times 3\) example of an Hermitian matrix is
\[ A=\left[\begin{array}{ccc}1&amp;\mathrm{i}&amp;5-2\mathrm{i}\\ -\mathrm{i}&amp;3&amp;0\\ 5+2\mathrm{i}&amp;0&amp;2\end{array}\right].\]
An Hermitian matrix is in fact a generalization of a symmetric matrix. The key property of an Hermitian matrix is the same as that of a real symmetric matrix – i.e. the eigenvalues are always <strong>real</strong>.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200624103226+EvDjTx</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200713171603+2t2zUb</text></tag>
    </tags>
    
</question>
</quiz>