<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24358 -->
  <question type="stack">
    <name>
				<text>16.1.2.2 Arithmetic sequence a and d</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Find \(a, \ d\) for the arithmetic sequence \({@a1@},\ {@a2@},\ {@a3@},\dots\)</p><br>
<p>\(a = \) [[input:ans_a]] [[validation:ans_a]] [[feedback:prt_a]]</p>
<p>\(d = \) [[input:ans_d]] [[validation:ans_d]] [[feedback:prt_d]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>We are given the arithmetic sequence \({@a1@},\ {@a2@},\ {@a3@},\dots\)</p>
<p>For an arithmetic sequence:</p>
<p>The value of \(a\) is always given by the first term of the sequence. Hence, \(a = {@aa@}\).</p>
<p>The value of \(d\) is given by the difference between any two consecutive terms in the sequence. Hence, \(d = {@dd@}\).</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>/*Constants*/
aa: rand(6)+1;
dd: rand(6)+1;

/*Sequence*/
a1: aa;
a2: aa+dd;
a3: aa+2*dd;</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>arithmeticc: \({@a1@},\ {@a2@},\ {@a3@},\dots \quad a = {@aa@}; d = {@dd@} \)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_a</name>
      <type>numerical</type>
      <tans>aa</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans_d</name>
      <type>numerical</type>
      <tans>dd</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt_a</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans_a</sans>
        <tans>aa</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_a-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_a-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_d</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>EqualComAss</answertest>
        <sans>ans_d</sans>
        <tans>dd</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_d-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt_d-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans_d</sans>
        <tans>dd</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_d-2-T</trueanswernote>
        <truefeedback format="html">
					<text><p dir="ltr" style="text-align: left;">Simplify your answer as far as possible.</p></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt_d-2-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans_d</sans>
        <tans>a2/a1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_d-3-T</trueanswernote>
        <truefeedback format="html">
					<text><p dir="ltr" style="text-align: left;">You are truying to find the common ratio of a geometric sequence. Here we have an arithmetic sequence.</p></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_d-3-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+201005155251+A2T6Y5</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210609143608+B2xHNc</text></tag>
    </tags>
    <deployedseed>1821905727</deployedseed>
    <deployedseed>251783026</deployedseed>
    <deployedseed>1542944653</deployedseed>
    <deployedseed>860921369</deployedseed>
    <deployedseed>814649881</deployedseed>
    <deployedseed>1871914094</deployedseed>
    <deployedseed>424005160</deployedseed>
    <deployedseed>1668651300</deployedseed>
    <deployedseed>188263919</deployedseed>
    <deployedseed>1915226398</deployedseed>
    <deployedseed>1416533407</deployedseed>
    <deployedseed>901161068</deployedseed>
    <deployedseed>1619021681</deployedseed>
    <deployedseed>174165223</deployedseed>
    <deployedseed>2142128525</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans_a</name>
        <value>aa</value>
      </testinput>
      <testinput>
        <name>ans_d</name>
        <value>dd</value>
      </testinput>
      <expected>
        <name>prt_a</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_a-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt_d</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_d-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans_a</name>
        <value>aa</value>
      </testinput>
      <testinput>
        <name>ans_d</name>
        <value>a2/a1</value>
      </testinput>
      <expected>
        <name>prt_a</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_a-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt_d</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_d-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans_a</name>
        <value>a1</value>
      </testinput>
      <testinput>
        <name>ans_d</name>
        <value>a2-a1</value>
      </testinput>
      <expected>
        <name>prt_a</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_a-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt_d</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_d-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <testinput>
        <name>ans_a</name>
        <value>ev(aa+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans_d</name>
        <value>ev(dd+2,simp)</value>
      </testinput>
      <expected>
        <name>prt_a</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_a-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt_d</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_d-3-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>