<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24627 -->
  <question type="stack">
    <name>
				<text>14.4.1.2 Exercise 2</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Find the length of \(y = {@m@}x^{\frac{3}{2}}\) between \(x = {@a@}\) and \(x = {@b@}\). Give your answer in <b>exact</b> form.</p>
<p>The length is [[input:ans1]]</p>
<p>[[validation:ans1]] [[feedback:prt1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p> Given \( y = {@y@} \) then \( \dfrac{dy}{dx} = {@dy@}.\)
<p>Using the formula \[\text{length of curve}=\int_a^b \sqrt{1+\left(\frac{\mathrm{d}y}{\mathrm{d}x}\right)^2} \ \mathrm{d}x,\] the length of the line is calculated to be: \[\ell = \int_{{@a@}}^{{@b@}} \sqrt{1+{@dy^2@}} \ \mathrm{d}x = \left[ {@I@} \right]_{{@a@}}^{{@b@}}={@ta1@}.\]</p>
For this integral you can use the substitution \( u= 1+{@dy^2@}.\)]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021071900</text>
		</stackversion>
    <questionvariables>
				<text>m: rand([2,4,6]);   /* I picked m=2 so that it cancels out when finding dy and makes calculations simpler */
y: m*x^(3/2);

a: 0          /*lower limit, pick 0 so it makes calculation easier*/
ps: rand([4,9,16,25,49,81]);   /* I want the quantity under the sqrt to give a perfect square so the result is a nice number, I pick this first and then find the uper limit I will use */
b: (ps-1)/(9*(m/2)^2);   /* upper limit*/

dy: diff(y, x);
I:  integrate(sqrt(1+(dy)^2), x);
ta1: integrate(sqrt(1+(dy)^2), x, a, b);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(y = {@m@}x^{\frac{3}{2}}, \,ps = {@ps@}, \,  x = [{@a@}, {@b@}], \, l = {@ta1@}\).</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.</span>]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>units</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201202062933+T9isB5</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210723093954+IC28OC</text></tag>
    </tags>
    <deployedseed>1725821938</deployedseed>
    <deployedseed>1912228845</deployedseed>
    <deployedseed>746452835</deployedseed>
    <deployedseed>1691250176</deployedseed>
    <deployedseed>499886202</deployedseed>
    <deployedseed>918642944</deployedseed>
    <deployedseed>1604755840</deployedseed>
    <deployedseed>74988268</deployedseed>
    <deployedseed>1726313280</deployedseed>
    <deployedseed>1424654372</deployedseed>
    <deployedseed>697257195</deployedseed>
    <deployedseed>88474474</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>