<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27192 -->
  <question type="stack">
    <name>
				<text>35.1.4.7 Complement of union, intersection</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Given 2 sets \(A, B\) and a universal set \(S\), recall that the complement of \(A\) is denoted by 
\(A'\) and the complement of \(B\) is denoted by \(B'\). Which of the following statements are true?</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>To see that these 2 statements are true, use definitions of complement, union, intersection. Alternatively, draw Venn diagrams representing the sets in question.</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200703224346+YY2Gcr</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200710174633+pZaujY</text></tag>
    </tags>
    
</question>
</quiz>