<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62181 -->
  <question type="stack">
    <name>
				<text>43.2 t distribution</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p></p><h3>Critical points of student’s \(t\) distribution</h3><br><p></p><p><img src="@@PLUGINFILE@@/43_2_t_Table_fig.svg" alt="" width="236" height="132" role="presentation" class="img-responsive atto_image_button_text-bottom"><br></p><h3><img src="@@PLUGINFILE@@/43_2_t_Table.svg" alt="" width="489" height="582" role="presentation" class="img-responsive atto_image_button_text-bottom"><br></h3><p><br></p><p><br></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200712061446+hs48Hm</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200712061726+PjT3el</text></tag>
    </tags>
    
</question>
</quiz>