<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26760 -->
  <question type="stack">
    <name>
				<text>10.3.2.4 Finding the exponential form</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
Express the complex number \(z = \displaystyle{@zed@}\)&nbsp; in exponential form. The argument of \(z\) should be given precisely in terms of \(\pi\), i.e. do not give your answer in decimal form:<p>\(z=\quad\)[[input:ans1]][[validation:ans1]]&nbsp;[[feedback:prt1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>We want to express \(z\) in the form \(re^{i\theta}\).</p><p>We calculate \(r = \sqrt{({@re@})^2 + ({@im@})^2} = {@arr@}\) and \(\theta = \tan^{-1}\left({@im@}/{@re@}\right) = {@th@}\).</p><p>Therefore: \[z = {@arr@}e^{{@th@}i}.\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021072200</text>
		</stackversion>
    <questionvariables>
				<text>th: rand([%pi/3,%pi/4,-%pi/3,-%pi/4])
arr: 2 + rand(10);
re: arr*cos(th);
im: arr*sin(th);
simp:false;
ta: arr*e^(th*%i);
zed: re + im*%i;</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\[ {@ta@} = {@zed@}\]</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color: green; --darkreader-inline-color:#7fff7f;" data-darkreader-inline-color=""><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color: orange; --darkreader-inline-color:#ffb733;" data-darkreader-inline-color=""><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color: red; --darkreader-inline-color:#ff3333;" data-darkreader-inline-color=""><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>13</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200618194224+7yFvM1</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210913221404+o3NePX</text></tag>
    </tags>
    <deployedseed>1901606859</deployedseed>
    <deployedseed>602818603</deployedseed>
    <deployedseed>894327509</deployedseed>
    <deployedseed>1479722686</deployedseed>
    <deployedseed>1994926504</deployedseed>
    <deployedseed>1023542860</deployedseed>
    <deployedseed>603625680</deployedseed>
    <deployedseed>1866674831</deployedseed>
    <deployedseed>1786551608</deployedseed>
    <deployedseed>1910411624</deployedseed>
    <deployedseed>1013862141</deployedseed>
    <deployedseed>583681724</deployedseed>
    <deployedseed>572819719</deployedseed>
    <deployedseed>686289728</deployedseed>
    <deployedseed>1394133457</deployedseed>
    <deployedseed>1380486129</deployedseed>
    <deployedseed>1750255028</deployedseed>
    <deployedseed>1893038075</deployedseed>
    <deployedseed>1857467541</deployedseed>
    <deployedseed>1886071763</deployedseed>
    <deployedseed>963949412</deployedseed>
    <deployedseed>1523582033</deployedseed>
    <deployedseed>1901542382</deployedseed>
    <deployedseed>1322232985</deployedseed>
    <deployedseed>1511949970</deployedseed>
    <deployedseed>233523471</deployedseed>
    <deployedseed>1198274324</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>0 + 2*%i</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>