<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26244 -->
  <question type="stack">
    <name>
				<text>7.3.2.1 Laplace expansion along the top row, minors</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Laplace expansion along the top row</h4>
<p>This is a technique which can be used to evaluate determinants of any order.
In principle, this method can use any row or any column as its starting point.  We
quote one example: using the top row.</p>

<p>Consider \(\Delta=\begin{vmatrix} 4 &amp; 1 &amp; 1 \\ 1 &amp; 2 &amp; 3 \\ 3 &amp; 1 &amp; 2 \end{vmatrix}\).</p>

<p>First we introduce the idea of a <b>minor</b>.  Each element in this array of numbers has
an associated number called its minor, calculated by removing the column and row in which the element
lies and taking the determinant of the remainder. For example consider entry \(a_{23}=3\). To calculate the minor associated with this entry, we strike out the second row and the third column:</p><p><img src="@@PLUGINFILE@@/7-3-fig-1-laplace-1.svg" alt="The matrix with rows [4, 1, 1], [1, 2, 3], [3, 1, 2], with a horizontal line crossing out the second row, and a vertical line crossing out column 3." class="img-responsive atto_image_button_text-bottom" width="68" height="63"><br></p><p>which leaves \(\begin{vmatrix} 4 &amp; 1 \\ 3 &amp; 1 \end{vmatrix} = 4 - 3 = 1.\) For the entry \(a_{31}=3\) we strike out the third row and first
column:</p><p><img src="@@PLUGINFILE@@/7-3-fig-2-laplace-2.svg" alt="The matrix with rows [4, 1, 1], [1, 2, 3], [3, 1, 2], with a horizontal line crossing out row 3, and a vertical line crossing out column 1." class="img-responsive atto_image_button_text-bottom" width="69" height="63"><br></p><p>which leaves \(\begin{vmatrix} 1 &amp; 1 \\ 2 &amp; 3 \end{vmatrix} = 3 - 2 = 1.\)</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200709145451+0Nuysz</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210702115240+gaELve</text></tag>
    </tags>
    
</question>
</quiz>