<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 32197 -->
  <question type="stack">
    <name>
				<text>19.2.2.3 Separation of Variables Ex1</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Solve the equation \(\dfrac{\mathrm{d}y}{\mathrm{d}x} = \displaystyle{@ff@}\) to obtain \(y = \pm\sqrt{A + D}\), where \(D = \displaystyle {@2/a2@}C \), \(C\) is the integration constant and \(A\) is an expression to be found.</p>
<p>Then \(A =&nbsp; \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>We are given the equation \(\dfrac{\mathrm{d}y}{\mathrm{d}x} = \displaystyle{@ff@}\).</p>
<p>We may rearrange this as \[\begin{aligned} {@f2@}\dfrac{\mathrm{d}y}{\mathrm{d}x} &amp;= {@f1@}, \\ \equiv \int {@f2@} \ \mathrm{d}y &amp;= \int {@f1@} \ \mathrm{d}x. \end{aligned}\]</p>
<p>Perform the integration to obtain</p>
<p>\[\begin{aligned} {@int(f2,y)@} &amp;= {@int(f1,x)@} + C, \\ \implies y^2 &amp;= - {@2/(a1*a2)*%e^(-a1*x)@} + {@2/a2@}C,&nbsp; \\ 
  \implies y &amp;= \pm \sqrt{- {@2/(a1*a2)*%e^(-a1*x)@} + {@2/a2@}C}
  \end{aligned}\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>a1: 1 + rand(4);
a2: 1 + rand(4);

f1: %e^(-a1*x);
f2: a2*y;
ff: (f1)/(f2);

f3: int(f2,y)-int(f1,x);


ta1: -2/(a1*a2)*(%e^(-a1*x))</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(\dfrac{\mathrm{d}y}{\mathrm{d}x} = {@ff@}\), \(D =  \displaystyle {@2/a2@}C\) </text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210614181228+h18A7p</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210615111759+1V3m2z</text></tag>
    </tags>
    <deployedseed>604858277</deployedseed>
    <deployedseed>2098210296</deployedseed>
    <deployedseed>1976295443</deployedseed>
    <deployedseed>1697973174</deployedseed>
    <deployedseed>1989390313</deployedseed>
    <deployedseed>583021017</deployedseed>
    <deployedseed>1023493569</deployedseed>
    <deployedseed>1636231051</deployedseed>
    <deployedseed>1223724923</deployedseed>
    <deployedseed>1532552658</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>