<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24559 -->
  <question type="stack">
    <name>
				<text>15.2.2.1 COM of plane lamina</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>Finding the centre of mass of a plane uniform lamina</h3>

<p>We are now interested in finding the centre of mass of a thin sheet of material, such as a plane sheet of metal, called a <b>lamina.&nbsp;</b>The mass is now not located at individual points. Rather, it is distributed continuously over the lamina. In
    what follows, we assume that the mass is distributed <b>uniformly</b> over the lamina and you will see how integration as the limit of a sum is used to find the centre of mass.
</p>
<p> The figure below shows a lamina where the centre of mass has been marked at point <em>G</em> with coordinates \((\bar{x}, \bar{y})\).&nbsp;If the total mass of the lamina is <em>M</em> then the moments about the \(y\)- and \(x\)-axes are respectively
    \(M\bar{x}\) and \(M\bar{y}\).&nbsp;Our approach to locating the position of <em>G</em>, i.e. finding \(\bar{x}\) and \(\bar{y}\), is to divide the lamina into many small pieces, find the mass of each piece, and calculate the moment of each piece
    about the axes.&nbsp;The sum of the moments of the individual pieces about the<em>&nbsp;\(y\)</em>-axis must then be equal to \(M\bar{x}\) and the sum of the moments of the individual pieces about the <em>x</em>-axis must equal \(M\bar{y}\).</p>
<figure class="HELM_figure"><img src="@@PLUGINFILE@@/15-2-fig-6-Lamina.svg" alt="Figure for section 15.2" width="611" height="350" class="img-responsive atto_image_button_text-bottom"><br>
    <figcaption>The centre of mass of the lamina is located at \(G(\bar{x}, \bar{y})\).</figcaption>
</figure>
<p>There are no formulae which can be memorized for finding the centre of mass of a <em>general</em> lamina because of the wide variety of possible shapes. Instead you should be familiar with the general technique for deriving the centre of mass. We will
    then see how in some cases, where the lamina can be described as the region bounded by simple curves, we can use the technique to find a standard formula.</p>
<p>An important preliminary concept is ‘mass per unit area’ which we now introduce.</p>
<h4>Mass per unit area</h4>
<p>Suppose we have a uniform lamina and select a piece of the lamina which has area equal to <b>one
</b><b>unit</b>. Let&nbsp;\(\rho\) stand for the mass of such a piece. Then \(\rho\) is called the <b>mass per unit area</b>. The mass of any other piece can be expressed in terms of \(\rho\). For example, an area of 2 units must have mass \(2\rho\),
    an area of 3 units must have mass \(3\rho\), and so on. Any portion of the lamina which has area \(A\) has mass \(\rho A\).</p>
<div class="HELM_keypoint">
    <h4>Key Point</h4>
    <p>If a lamina has mass per unit area, \(\rho\), then the mass of part of the lamina having area \(A\) is \(A \rho\).</p>
</div><br>In order to find the centre of mass of a lamina we will calculate \(M_x\) and \(M_y\), the moments of inertia about the \(x\)-axis and \(y\)-axis respectively. We do this in two ways. <br><br>First, acting as a point mass \(A\rho\) at its centre of
mass \((\bar{x},\bar{y})\) we have \(M_y=A\rho \bar{x}\) and&nbsp;\(M_x=A\rho \bar{y}\). In other words \(\bar{x}=\dfrac{M_y}{A\rho}\) and&nbsp;\(\bar{y}=\dfrac{M_x}{A\rho}\).<br><br>Secondly, we take thin rectangular slices of the lamina and sum the
contribution of each to the moments of inertia. Taking the limit as the width of these slices approaches \(0\) we will obtain an integral and solving this gives an expression for the moment of inertia which does not rely explicitly on&nbsp;\((\bar{x},\bar{y})\).&nbsp;<br><br>Finally
we combine this data to find \((\bar{x},\bar{y})\).<br>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200813161637+NCXphC</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210610161450+9ivWxG</text></tag>
    </tags>
    
</question>
</quiz>