<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24282 -->
  <question type="stack">
    <name>
				<text>Welcome constrained function</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>Where to find the theory</h3><br>See section 9.7.4 of Engineering Mathematics, or 8.7.4 of Mathematics for the Natural Sciences text for the necessary theory needed to complete these exercises.<br><br>Each question asks you to optimised (find the extreme points) of a given functions given some constrains. In solving this we construct an auxiliary function often called the Lagrangian. If we want to optimised \(f(x,y)=0\) subject the the constrain \( g(x,y)\) then we introduce the function<br>\[ f(x,y)-lg(x,y)\]<br>involving a new variable \(l\). The solution to the problem is given by the solutions to system of equations coming from the (three) first order partial derivatives of the new function. Please see the book for more details before attempting these exercises.<br>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:moodle2.maths.ed.ac.uk+210314224349+Ew9TLe</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+210314224349+dLOMBk</text></tag>
    </tags>
    
</question>
</quiz>