<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26215 -->
  <question type="stack">
    <name>
				<text>7.2.3.2 Between Task1 and Task2</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>It is clear that \(AB\) and \(BA\) are <b>not</b> in general the same. In fact it is the <b>exception</b> that \(AB=BA.\) In the special case in which \(AB=BA\) we say that the matrices \(A\) and \(B\) <b>commute</b>.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200619133424+iBRePa</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200701144031+HFuX6j</text></tag>
    </tags>
    
</question>
</quiz>