<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24289 -->
  <question type="stack">
    <name>
				<text>5.1.1.3 Linear functions</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Linear Functions</h4><p>Another workbook has introduced linear functions of the form \(y = ax + b\). Such functions give rise to straight-line graphs. The coefficient <em>a</em> is the slope. If <em>a</em> is positive the graph of <em>y</em> against <em>x</em> slopes upwards. If <em>a</em> is negative the graph slopes downwards. The coefficient <em>b</em> gives the intercept on the y-axis. The terms <em>a</em> and <em>b</em> may be called the <strong>parameters</strong> of the line. Note that this is a different use of the term ‘parameter’ than in the parametrisation of functions discussed in 2.</p>
<h5>Linear models for falling rocks</h5>
<p>In modelling it is wise to use a notation which fits in with the application. When modelling velocity under constant acceleration, we shall replace the dependent variable <em>y</em> by \(v\) (for velocity), and the independent variable \(x\) by \(t\) (for time). The acceleration will be denoted by the symbol \(a\). Consider the motion of a rock dislodged from the top of a cliff (35 m high) by a villain during the filming of a thriller. The film producer might be interested in how long the rock would take to fall to the ground below the cliff and how fast it would be travelling at ground impact. The rock may be assumed to have a constant downward acceleration of 9.8 m s\(^{-2}\) which the acceleration due to gravity. The velocity \((v\ \text{m s}^{-1})\) of a rock, falling from the top of a cliff 35 m high, can be modelled by the equation</p> <p> \(\qquad v = 9.8t \qquad (0\le t\le 2.7)\)</p> <p>where \(t\) is the time in seconds after the rock starts to fall. This follows from the fact that acceleration is the rate of change of velocity with time. If the acceleration is constant and the object starts from rest, then the velocity is given simply by the product of acceleration and time. The upper limit for \(t\) is the time at which the rock hits the ground measured with a stop-watch (about 2.7 s in this case).</p> <p>Figure 2 shows \(v\) as a function of \(t\). Velocity is a linearly increasing function of time and its graph is a straight line passing through \(t = 0, v = 0\). Note that various assumptions are needed to obtain the quoted result of a linear variation in speed with time: it is assumed that there is no air resistance, no spinning and no wind.</p> 
<p><img width="384" height="131" class="img-responsive atto_image_button_text-bottom" alt="velocity time graph y intercept = 0 and gradient is 9.8" src="@@PLUGINFILE@@/5.1%20linear%20functions%20figure.svg"></p>
<p><strong>Figure 2</strong>: Graph of \(v=9.8t\) for the falling rock</p> 
<p>In what way should the equation for \(v\) be altered if the villain were able to <strong>throw</strong> the rock downwards at 5 m s\(^{-1}\)? Provided we are measuring position or displacement downwards, a downwards velocity is positive. Now we have that \(v = 5\) when \(t = 0\). So a new model for \(v\) is </p><p>\(\qquad v = 9.8t + 5 \qquad (0\le t\le T_1)\) </p><p>Since they are both downwards, the initial velocity simply adds to the velocity at any time resulting from falling under gravity. Note that \(T_1\) is being used now for the upper limit on \(t\) (instead of 2.7) because 2.7 is (approximately) the time taken to fall 35 m <strong>from rest</strong> rather than with an initial downwards velocity. (Using the symbol \(T_1\) saves us trying to work out its value for the moment!) Note that a general form of the model for motion under constant acceleration of magnitude \({a}\) m s\(^{-2}\) given an initial speed \({b}\) m s\(^{-1}\) is \(v=at+b\). In the model just considered \(a=9.8\) and \(b=5\).</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200722160503+RLicAq</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220214100820+cBNcTG</text></tag>
    </tags>
    
</question>
</quiz>