<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26283 -->
  <question type="stack">
    <name>
				<text>22.3.1.4 Final key point</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<div class="HELM_keypoint">
<h4>Key Point</h4>
<p>An \(n\times n\) matrix with repeated eigenvalues can be diagonalized provided we can obtain \(n\) linearly independent eigenvectors for it. This will be the case if, for each repeated eigenvalue \(\lambda_i\), of multiplicity \(m_i&gt;1\), we can obtain \(m_i\) linearly independent eigenvectors.</p>
</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200624091812+yDfVZT</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200624091844+1FcfGK</text></tag>
    </tags>
    
</question>
</quiz>