<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26284 -->
  <question type="stack">
    <name>
				<text>22.3.2.1 Symmetric matrices</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span style="color: inherit; font-family: inherit;">Symmetric matrices</span></h3>
<p>Symmetric matrices have a number of useful properties which we will investigate in this Section.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200624092214+TU6vzB</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200914192450+thE5HG</text></tag>
    </tags>
    
</question>
</quiz>