<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24350 -->
  <question type="stack">
    <name>
				<text>16.5.4.2 Find the Taylor series</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>What is the Taylor series of <span class="math inline">\(\dfrac{1}{1-x}\)</span> about <span class="math inline">\(x=2\)</span>. (That is, find a power series in powers of <span class="math inline">\((x-2)\)</span>.)</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p dir="ltr" style="text-align: left;"><br></p>
<p>First, obtain the first three derivatives and the <span class="math inline">\(n^{\text{th}}\)</span> derivative of <span class="math inline">\(f(x)=\dfrac{1}{1-x}\)</span>: <span class="math display">\[f'(x)=\dfrac{1}{(1-x)^2},\qquad f''(x)=\dfrac{2}{(1-x)^3},\qquad f'''(x)=\dfrac{6}{(1-x)^4},\qquad
\cdots \quad f^{(n)}(x)=\dfrac{n!}{(1-x)^{n+1}}\]</span> Now evaluate these derivatives at <span class="math inline">\(x_0=2\)</span>: <span class="math display">\[f'(2)=1,\qquad f''(2)=-2,\qquad f'''(2)=6,\qquad f^{(n)}(2)=(-1)^{n+1}n!\]</span> Hence, write down the Taylor expansion of <span class="math inline">\(f(x)=\dfrac{1}{1-x}\)</span> about <span class="math inline">\(x=2\)</span>: <span class="math display">\[\dfrac{1}{1-x}=-1+(x-2)-(x-2)^2+(x-2)^3+\cdots + (-1)^{n+1}(x-2)^n+\cdots\]</span></p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201130143033+QKjqe2</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+210302092821+pUIntL</text></tag>
    </tags>
    
</question>
</quiz>