<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62140 -->
  <question type="stack">
    <name>
				<text>23.5.5 ex4 cosine series</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
Obtain the half-range Fourier Cosine series for the functions
\[f(t) = \sin t \hspace{1cm} 0 \leq t \leq \pi\]
<p></p>
Series is
<p>[[input:ans1]]  \(+ \displaystyle{ \sum_{n=2}^\infty} \) [[input:ans2]]</p>
[[validation:ans1]][[feedback:prt1]]
[[validation:ans2]][[feedback:prt2]]]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[We extend as an even function, so 
\[ f(t) = -\sin t \hspace{1cm} -\pi \leq t \leq 0 \]
Therefore we define 
\[ F(t) = 
\begin{cases} 
-\sin t \hspace{1cm} -\pi \leq t \leq 0 \\
\sin t \hspace{1.2cm} 0 \leq t \leq \pi
\end{cases}
\]
As \( F(t)\) is an even function, \(b_n = 0\). Then we have
\[
\begin{align}
a_n &amp;= \dfrac{1}{\pi} \int_{-\pi}^\pi F(t) \cos nt \hspace{0.1cm} dt = \dfrac{2}{\pi} \int_{0}^\pi \sin t \cos nt \hspace{0.1cm} dt  \\
&amp;= \dfrac{2}{\pi} \left[ - \dfrac{\cos \left( \left(1-n\right)  t \right)}{2 \left( 1-n \right)} - \dfrac{\cos \left( \left(1+n\right)  t \right)}{2 \left( 1+n \right)} \right]_0^\pi \\
&amp;= \dfrac{1}{\pi} \left[ \dfrac{1}{1-n} - \dfrac{ \cos \left( [1-n] \pi \right) }{1-n} + \dfrac{1}{1+n} - \dfrac{ \cos \left( [1+n] \pi \right) }{1+n} \right]  \\
&amp;= \dfrac{1}{\pi} \left[ \dfrac{1}{1-n} \left( 1 - \cos \left( [1-n] \pi \right) \right) + \dfrac{1}{1+n} \left( 1 - \cos \left( [1+n] \pi \right) \right) \right] 
\end{align}
\]
Also, 
\[ a_0 = \dfrac{1}{\pi} \left( 1 - \cos \pi + 1 - \cos \pi \right) = \dfrac{4}{\pi} \]
so \( \dfrac{a_0}{2} = \dfrac{2}{\pi} \).
<p></p>
We can see from \(a_n\) that the series is not defined for \(n = 1\), so the half range Fourier cosine series is given by
\[ F(t) = \dfrac{2}{\pi} + \sum_{n=2}^\infty \dfrac{1}{\pi} \left[ \dfrac{1}{1-n} \left( 1 - \cos \left( [1-n] \pi \right) \right) + \dfrac{1}{1+n} \left( 1 - \cos \left( [1+n] \pi \right) \right) \right] \cos nt \]]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021120900</text>
		</stackversion>
    <questionvariables>
				<text>ta1:2/pi;
ta2: 1/pi*( 1/(1-n)*(1 - cos((1-n)*pi)) + 1/(1+n)*(1 - cos((1+n)*pi)) )*cos(n*t);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text></text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>50</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200906103638+ZFgIBk</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220518110453+MZJgx1</text></tag>
    </tags>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>