<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26246 -->
  <question type="stack">
    <name>
				<text>7.3.2.3 Introducing cofactors</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4> Cofactors </h4>
<p>A cofactor is a minor with a sign attached. The appropriate sign comes from the pattern of signs appropriate to a 
\(3\times 3\) array: \[\begin{bmatrix}+&amp;-&amp;+\\ -&amp;+&amp;-\\ +&amp;-&amp;+\end{bmatrix}\]
(i.e. positive signs on the leading diagonal and the signs "alternate" everywhere else.)</p>

<p>Each entry of a matrix \(A\) has a cofactor associated with it. The cofactor of entry \(a_{11}\)
is denoted by \(A_{11}\), that of \(a_{23}\)  by \(A_{23}\) and so on.</p>

<p>To obtain the cofactor of an element of a \(3\times 3\)  matrix we simply multiply the minor of that element by the corresponding sign from the \(3\times 3\) array of signs.</p>

<p>Hence, for the matrix \(A = \begin{bmatrix} 4 &amp; 1 &amp; 1 \\ 1 &amp; 2 &amp; 3 \\ 3 &amp; 1 &amp; 2 \end{bmatrix}\), the cofactor corresponding to \(a_{23}\) is
\[A_{23} = - \begin{vmatrix} 4 &amp; 1 \\ 3 &amp; 1 \end{vmatrix} = -1\] 
and the cofactor corresponding to \(a_{31}\) is
\[A_{31} = + \begin{vmatrix} 1 &amp; 1 \\ 2 &amp; 3 \end{vmatrix}=1.\]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200725145227+zyxnSz</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200908091914+ii9way</text></tag>
    </tags>
    
</question>
</quiz>