<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24518 -->
  <question type="stack">
    <name>
				<text>13.3 Introduction</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span>The area bounded by a curve</span><br></h3>

<p>One of the important applications of integration is to find the area bounded by a curve. Often such an area can have a physical significance like the work done by a motor, or the distance travelled by a vehicle. In this Section we explain how such an area is calculated.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200730092231+2lJx57</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200730092501+KP3Qwa</text></tag>
    </tags>
    
</question>
</quiz>