<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 33831 -->
  <question type="stack">
    <name>
				<text>Exponential</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p></p>
<p>Graph of function should appear here:</p>
<p>[[jsxgraph width='300px' height='250px']] JXG.Options.axis.ticks.insertTicks = false; /*don't insert extra ticks between major ticks*/ JXG.Options.axis.ticks.majorHeight = 40; JXG.Options.text.fontSize=16;
JXG.Options.axis.ticks.ticksDistance=10000000;
var board = JXG.JSXGraph.initBoard(divid, {boundingbox: {#BB#}, axis: true, showCopyright: false, showNavigation:false});
var f = board.jc.snippet('{#fx#}', true, 'x', true); var ff = board.create('functiongraph',[function(x) {var n={#n1#}; return Math.exp(n*x);}]); var i1 = board.create('integral',[{#[a,b]#},ff], {fillColor:'#003399',fillOpacity:0.15});
i1.curveLeft.setAttribute({visible:false}); /* hide the gliders */ i1.curveRight.setAttribute({visible:false}); i1.label.setAttribute({visible:false});/*hide
    the integral calculation*/ var xaxis = board.create('axis', [[0,0],[1,0]]); board.create('ticks', [xaxis, {#[a1,b1]#}], {strokeColor: '#003399', majorHeight:15, drawLabels: true, label: { offset: [0, -5],anchorX: 'middle', anchorY: 'top' }}); [[/jsxgraph]]
</p>
<p></p>
<p>\( A = \) [[input:ans1]] [[validation:ans1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>/* Find area under curve */
n1: 2+rand(5);
fx:exp(n1*x);                       /* function*/
a1:1;                          /* lower limit*/
b1: 3;     /*upper limit*/
Area: integrate(fx, x, a1, b1);

/*jsxgraph settings */
[a,b]: [a1,b1];            /* matrix of limits for jsx graph*/
t_dist: exp(n1*1);      /*tick distance y axis */
BB: [-1,e^(n1*3.1),b1+1, -exp(n1*2.5)]     /* boundingbox:[left, top, right, bottom]*/

/* the boundingbox conrdinates will depend on the function, here is a simple hyperbola so I definded left-top-bottom by hand. In other functions it may need to evaluate the value of the function at a point to define the height of the box or something else that seems reasonable */</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>\(A=\int_{@a1@}^{@b1@}{@fx@}, dx = {@Area@} \)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>Area</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>Area</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210705103205+KGt4f0</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210705170039+QBC1Wx</text></tag>
    </tags>
    <deployedseed>645647693</deployedseed>
    <deployedseed>1046377108</deployedseed>
    <deployedseed>2138284967</deployedseed>
    <deployedseed>1815667557</deployedseed>
    <deployedseed>884782260</deployedseed>
    
</question>
</quiz>