<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24472 -->
  <question type="stack">
    <name>
				<text>11.7.1a imp to exp</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Implicit and explicit functions</h4>
Equations such as \(y = x^2\), \(y = \dfrac{1}{x}\), \(y = \sin x\) are said to define \(y\) <b>explicitly</b> as a function of \(x\) because the variable \(y\) appears alone on one side of the equation.
<p></p>
The equation
\[ yx + y + 1 = x\]
is not of the form \(y = f(x)\) but can be put into this form by simple algebra. Indeed we have
\(y(x + 1) = x - 1\) so
\[ y = \dfrac{x - 1}{x + 1}.\]
We say that \(y\) is defined <b>implicitly</b> as a function of \(x\) by means of \(yx+y+1 = x\), the actual function being given <b>explicitly</b> as
\[ y = \dfrac{x - 1}{x + 1}.\]
We note than an equation relating \(x\) and \(y\) can implicitly define <b>more than one</b> function of \(x\).
<p></p>
For example, if we solve
\[ x^2 + y^2 = 1 \]
we obtain \( y = \pm \sqrt{1 - x^2}\), so \( x^2 + y^2 = 1 \) defines implicitly two functions
\begin{align}
f_1(x) &amp;=  \sqrt{1 - x^2} \\
f_2(x) &amp;=  -\sqrt{1 - x^2} 
\end{align}]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200809091159+uw3cAP</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200809093749+FAVm9x</text></tag>
    </tags>
    
</question>
</quiz>