<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26239 -->
  <question type="stack">
    <name>
				<text>7.2.5.2 Non-square multiplication example</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_example">Example</h4>
<p>Find the product \(AB\) if \(A = \begin{bmatrix} 1 &amp; 2 &amp; 2 \\ 2 &amp; 3 &amp; 4 \end{bmatrix}\) and \(B = \begin{bmatrix} 2 &amp; 5 \\ 6 &amp; 1 \\ 4 &amp; 3 \end{bmatrix}\) </p>
<h4 class="HELM_solution">Solution</h4>
<p>Since \(A\) is a \(2\times3\) matrix and \(B\) is a \(3\times2\) the product \(AB\) can be found and results in a \(2\times2\) matrix. \begin{align*} AB &amp; = \begin{bmatrix} 1 &amp; 2 &amp; 2 \\ 2 &amp; 3 &amp; 4 \end{bmatrix}\begin{bmatrix} 2 &amp; 5 \\ 6 &amp; 1 \\ 4 &amp; 3 \end{bmatrix} \\
&amp; = \begin{bmatrix}
\begin{bmatrix} 1 &amp; 2 &amp; 2\end{bmatrix} \begin{bmatrix} 2 \\ 6 \\ 4 \end{bmatrix} &amp; \begin{bmatrix}
1 &amp; 2 &amp; 2 \end{bmatrix} \begin{bmatrix} 5 \\ 1 \\ 3 \end{bmatrix}
\\ \\
\begin{bmatrix} 2 &amp; 3 &amp; 4 \end{bmatrix} \begin{bmatrix} 2 \\ 6 \\ 4 \end{bmatrix} &amp; \begin{bmatrix}
2 &amp; 3 &amp; 4 \end{bmatrix} \begin{bmatrix} 5 \\ 1 \\ 3 \end{bmatrix}
\end{bmatrix} \\
&amp; = \begin{bmatrix} 22 &amp; 13 \\ 38 &amp; 25 \end{bmatrix}\end{align*}</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200814121443+orhWN4</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200819154214+lLj0d9</text></tag>
    </tags>
    
</question>
</quiz>