<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27286 -->
  <question type="stack">
    <name>
				<text>36.1.5.1 Location and spread</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>Location and spread</h3>
<p>The statistical properties of the observations in a sample are often described as a <b>frequency distribution</b>, referring to the fact that some values, in the case of discrete or categorical variables, or some ranges of values, in the case of continuous variables, may have greater frequencies than others. Very often the distribution of observations in a sample is summarized using two statistics, one of which measures the <b>location</b> or central value of the distribution and one which measures the <b>dispersion</b> or spread of the values about this central value.  As we shall see later, it is often important to convey information on the <b>shape</b> of a distribution as well. However, in many cases in practice, the shape of the distribution, as represented, for example, by a histogram, is reasonably symmetrical and roughly follows the "bell-shape'' illustrated below. In such cases the location-and-spread summary may suffice.</p>
<br>
<p style="text-align: center;"><img src="@@PLUGINFILE@@/36-1-fig-4-bell_shape.svg" alt="" width="397" height="84" class="img-responsive atto_image_button_text-bottom"></p>
<h5><b>Notation</b>:</h5><p>It is important to remember the distinction between a <b>population</b> and a <b>sample</b>. In later Workbooks we will look at how sample statistics can be used to give us information about summaries of a population. For now, however, we are concentrating on the sample statistics. To help keep the distinction clear, we usually use Greek letters, such as \( \mu\) or \( \sigma\), to represent population quantities, the values of which are, of course, usually unknown, and the equivalent Roman letters, such as \(m\) or \(s\), to represent sample statistics.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200914145654+4vxo0X</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200914145654+Dksufy</text></tag>
    </tags>
    
</question>
</quiz>