<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 36040 -->
  <question type="stack">
    <name>
				<text>1.5.2.16 Transposition</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise <br></p>
<p> When a ball is dropped from rest onto a horizontal surface it will bounce before eventually coming to rest after a time \(T\) where \[T=\dfrac{{@a1@}v}{g} (\dfrac{{@a2@}}{1-e})\] where \(v\) is the speed immediately after the first impact, and \(g\) is a constant called the acceleration due to gravity. Transpose this formula to make \(e\), the coefficient of restitution, the subject.
</p>

<p>\(e\) = [[input:ans1]] [[validation:ans1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p> We have the equation: \[T=\dfrac{{@a1@}v}{g} (\dfrac{{@a2@}}{1-e})\] First multiply both sides by g: \[\implies Tg = {@a1@}v (\dfrac{{@a2@}}{1-e})\] Then divide both sides by \({@a1@}v\): \[\implies \dfrac{Tg}{{@a1@}v} = \dfrac{{@a2@}}{1-e}\] Multiply both sides by (1-e) and then expand out the brackets: \[\implies \dfrac{Tg}{{@a1@}v}(1-e) = {@a2@} \implies \dfrac{Tg}{{@a1@}v} - \dfrac{Tg}{{@a1@}v}e = {@a2@}\] Add \(\dfrac{Tg}{{@a1@}v}e\) to both sides: \[\implies \dfrac{Tg}{{@a1@}v} = {@a2@} + \dfrac{Tg}{{@a1@}v}e\] Divide both sides by \(\dfrac{Tg}{{@a1@}v}\) and simplify the fraction: \[\implies 1 = \dfrac{{@a2@}}{\dfrac{Tg}{{@a1@}v}} + e \implies 1 = \dfrac{{@a2*a1@}v}{Tg} + e\] Finally subtracting \(\dfrac{{@a2*a1@}v}{Tg}\) from both of the equation makes \(e\) the subject of the formula: \[\implies e = 1 - \dfrac{{@a2*a1@}v}{Tg}\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021072200</text>
		</stackversion>
    <questionvariables>
				<text>a1: rand([2,3,4])
a2: rand([1,2,3])
ta: 1-(a1*a2*v)/(T*g)</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>\(T=\dfrac{{@a1@}v}{g} (\dfrac{{@a2@}}{1-e})\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>EqualComAssRules</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>[oneMul]</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
					<text>You have effectively transposed the formula to make \(e\) the subject but it should be fully simplified to: \(1-\dfrac{2v}{Tg}\)</text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210817113622+OiMkeu</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+211213113324+1utHgc</text></tag>
    </tags>
    <deployedseed>1460224079</deployedseed>
    <deployedseed>824269369</deployedseed>
    <deployedseed>1562797931</deployedseed>
    <deployedseed>1551295517</deployedseed>
    <deployedseed>705139204</deployedseed>
    <deployedseed>598488617</deployedseed>
    <deployedseed>1110880310</deployedseed>
    <deployedseed>2030398094</deployedseed>
    <deployedseed>446807908</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>((a1*a2*v)/(T*g))/((a1*a2*v)/(T*g))-(a1*a2*v)/(T*g)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>e=1-(a1*a2*v)/(T*g)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore></expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>