<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24572 -->
  <question type="stack">
    <name>
				<text>15.2.1.9 COM of 5 masses on line</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Find the \(x\) coordinate of the centre of mass of 5 identical masses placed at \(x={@x1@}, x={@x2@}, x={@x3@}, x={@x4@}, x={@x5@}.\)</p><p><b>Give your answer as a rational number (fraction).</b><br></p><p>\(\bar{x}= \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>We are asked to find the \(x\) coordinate of the centre of mass of 5 identical masses placed at \(x={@x1@}, x={@x2@}, x={@x3@}, x={@x4@}, x={@x5@}.\)<br></p>
We compute: \[\bar{x} = \dfrac{\sum_{i=1}^{5} x_i}{5} =&nbsp; {@ta@}\]]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>x1 : 1 + rand(3);
x2 : x1 + 1 + rand(3);
x3 : x2 +1+rand(3);
x4 : x3 + 1+rand(3);
x5 : x4 + 1+rand(3);
xt : x1 + x2 + x3 + x4 + x5;
ta : xt/5;</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>points: {@x1@},{@x2@},{@x3@},{@x4@},{@x5@}, so \( \bar{x} = {@ta@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>+</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200817134706+V5Sl7N</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210609153247+KMKFI8</text></tag>
    </tags>
    <deployedseed>1038343036</deployedseed>
    <deployedseed>1846931960</deployedseed>
    <deployedseed>1101824154</deployedseed>
    <deployedseed>1712544634</deployedseed>
    <deployedseed>1943211141</deployedseed>
    <deployedseed>469814823</deployedseed>
    <deployedseed>413827208</deployedseed>
    <deployedseed>150637372</deployedseed>
    <deployedseed>285376279</deployedseed>
    <deployedseed>1417114881</deployedseed>
    <deployedseed>1682095670</deployedseed>
    <deployedseed>1806508940</deployedseed>
    <deployedseed>649635468</deployedseed>
    <deployedseed>1720287244</deployedseed>
    <deployedseed>736271718</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(ta+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>