<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26193 -->
  <question type="stack">
    <name>
				<text>5.3.2.6 Water level</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Write down a function relating water level \(L\) (meters) in a harbour, to time \(t\) (hours), starting when the level is equal to the mean level of \({@m1@} \, \textrm{m}\), that has an amplitude of \({@a1@} \, \textrm{m}\) and has a period of twelvehours.</p>
<p>\(L= \) [[input:ans1]] [[validation:ans1]] [[feedback:ans1]]</p>
<p></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p></p>
<p>In the general form \(y=a\sin(bt+c)+d\), the phase is \(c=0\), the period is equal to \(12\), therefore \(b=\dfrac{2\pi}{12}=\dfrac{\pi}{6}\), the amplitude is \(a={@a1@}\), and the mean value is \(d={@m1@}\). Hence</p>
<p>\(L=\displaystyle {@ta@} \qquad (t \geq 0).\)</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021120900</text>
		</stackversion>
    <questionvariables>
				<text>m1: rand([2,3,4,5]);
a1: rand([6,7,8]);
T1: 12;
b: 2*%pi/T1;
ta:a1*sin(b*t)+m1
</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\[ L= {@ta@} \]</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>ans1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ans1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ans1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200821114127+FU6oTI</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220217133723+z005v1</text></tag>
    </tags>
    <deployedseed>1045286110</deployedseed>
    <deployedseed>750233390</deployedseed>
    <deployedseed>512656639</deployedseed>
    <deployedseed>513733147</deployedseed>
    <deployedseed>1234427712</deployedseed>
    <deployedseed>544562816</deployedseed>
    <deployedseed>1496032883</deployedseed>
    <deployedseed>179383654</deployedseed>
    <deployedseed>1362306131</deployedseed>
    <deployedseed>729353952</deployedseed>
    <deployedseed>199914419</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>ans1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>ans1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <expected>
        <name>ans1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>ans1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>