<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26035 -->
  <question type="stack">
    <name>
				<text>4.5.1.1 Applications of trigonometry to waves</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Two-dimensional motion</h4>
<p>Suppose that a wheel of radius \(R\) is rotating anticlockwise as shown in the figure below.</p>
<figure class="HELM_figure">
<img src="@@PLUGINFILE@@/4-5-fig38.svg" alt="Fig 38" class="atto_image_button_text-bottom" width="225" height="187">
</figure>
<p>Assume that the wheel is rotating with an angular velocity \(\omega\) radians per second about \(O\) so that, in a time \(t\) seconds, a point \((x,y)\) initially at position&nbsp; \(A\) on the rim of the wheel moves to a position \(B\) such that angle \(AOB=\omega t\) radians. <br></p><p style="text-align: left;"> Then the coordinates \((x,y)\) of&nbsp; \(B\) are given by \[ \quad \qquad x=OP=R\cos\omega t\] \[ \quad \qquad y=OQ=PB=R\sin\omega t\]</p>
<p> We know that both the standard sine and cosine functions have period \(2\pi\). Since the angular velocity is \(\omega\) radians per second the  wheel will make one complete revolution in \(\dfrac{2\pi}{\omega}\) seconds.</p>
<p>The time \(\dfrac{2\pi}{\omega}\) (measured in seconds in this case) for one complete revolution is called the <b>period</b> of rotation of the wheel.&nbsp; The number of complete revolutions per second is thus \(\dfrac{1}{T}=f\) say which is called the <b>frequency</b> of revolution.&nbsp; Clearly&nbsp;\(f =\dfrac{1}{T}=\dfrac{\omega}{2\pi}\) relates the three quantities introduced here.&nbsp; The angular velocity \(\omega=2\pi f\) is sometimes called the <b>angular frequency</b>.</p>

<h4>One-dimensional motion</h4>
<p>The situation we have just outlined&nbsp; is&nbsp; two-dimensional motion.&nbsp; More simply we might consider one-dimensional motion.</p><p>An example is the motion of the projection onto the \(x\)-axis of a point \(B\) which moves with uniform angular velocity \(\omega\) round a circle of radius \(R\) (see the figure below). As \(B\) moves round, its projection \(P\) moves to and fro across the diameter of the circle.</p>
<figure class="HELM_figure">
<img src="@@PLUGINFILE@@/4-5-fig39.svg" alt="Fig 39" class="atto_image_button_text-bottom" width="225" height="187">
</figure>
<p>The position of \(P\) is given by</p>
\[x=R\cos \omega t\]
<p>Clearly, from the known properties of the cosine function, we can deduce the following:</p>
<ol>
<li>\(x\) varies periodically with \(t\) with period \(T=\dfrac{2\pi}{\omega}\).</li>
<li>\(x\) will have maximum value \(+R\) and minimum value \(-R\). (This quantity \(R\) is called the <b>amplitude</b> of the motion.)</li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200619132211+RXBCXq</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210630122102+pGOJAB</text></tag>
    </tags>
    
</question>
</quiz>