<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26771 -->
  <question type="stack">
    <name>
				<text>10.4.2.1 Finding the roots of complex numbers</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>Finding the roots of complex numbers</h3>
<p></p>
<p>In this subsection we ask if we can obtain fractional powers of complex numbers; for example what are the values of <span class="math inline">\(8^{1/3}\),</span>&nbsp;<span class="math inline">\((-24)^{1/4}\)</span> or even <span class="math inline">\((1+i)^{1/2}\)</span>?</p>
<p>More precisely, for these three examples, we are asking for those values of <span class="math inline">\(z\)</span> which satisfy \[\qquad z^3-8=0\qquad\text{or}\qquad z^4+24=0\qquad\text{or}\qquad z^2-(1+i)=0\] Each of these problems involve finding roots
    of a complex number.</p>
<p>To solve problems such as these we need to be more careful with our interpretation of arg<span class="math inline">\((z)\)</span> for a given complex number <span class="math inline">\(z\)</span>.</p>
<h5><span style="">Revisiting arg\((z)\)</span></h5>
<p>By definition, arg<span class="math inline">\((z)\)</span> is the angle made by the line representing <span class="math inline">\(z\)</span> with the positive <span class="math inline">\(x\)</span>-axis. See (a) in the figure below. However, figure (b)
    below shows you can increase <span class="math inline">\(\theta\)</span> by <span class="math inline">\(2\pi\)</span> (<span class="math inline">\(360^{\circ}\)</span>) and still obtain the same line in the <span class="math inline">\(xy\)</span>    plane. In general, as indicated in figure (c), any integer multiple of <span class="math inline">\(2\pi\)</span> can be added to or subtracted from arg<span class="math inline">\((z)\)</span> without affecting the <strong>Cartesian form</strong> of
    the complex number.</p>
<p><br></p>
<figure class="HELM_figure"><img src="@@PLUGINFILE@@/fig1-1.svg" alt="" role="presentation">
<figcaption> The argument of a complex number.</figcaption>
</figure>
<p><br></p>
<div class="HELM_keypoint">
    <h4>Key Point</h4>
    arg\((z)\) is unique only up to an integer multiple of \(2\pi\).
</div>
<p><br></p>
<p>For example: <span class="math inline">\[\qquad z=1+i=\sqrt{2}\left(\cos{\pi \over 4}+i\sin{\pi\over 4}\right)\]</span></p>
<p>in polar form.</p>
<p>However, we could also write: <span class="math inline">\[\qquad z=1+i=\sqrt{2}\left(\cos\left({\pi \over 4}+2\pi\right)+i\sin\left({\pi\over 4}+2\pi\right)\right)\]</span>&nbsp;or in full generality:<span class="math inline">\[\qquad z=1+i=\sqrt{2}\left(\cos\left({\pi \over 4}+2k\pi\right)+i\sin\left({\pi\over
4}+2k\pi\right)\right)\qquad k=0,\pm 1,\pm 2,\ldots\]</span>This last expression shows that in the polar form of a complex number, the argument of <span class="math inline">\(z\)</span>, arg<span class="math inline">\((z)\)</span>, can take infinitely
    many different values, each one differing by an integer multiple of <span class="math inline">\(2\pi\)</span>. This is nothing more than a consequence of the well-known properties of the trigonometric functions: \[\qquad \cos(\theta +2k\pi )\equiv\cos\theta,\qquad
    \sin(\theta +2k\pi)\equiv\sin\theta\]for <em>any</em> integer \(k\).</p>
<p>We will now show how we can use this more general interpretation of arg<span class="math inline">\((z)\)</span> in the process of finding roots.</p>
<p></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200622112045+HAI55z</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210702151048+pjUg42</text></tag>
    </tags>
    
</question>
</quiz>