<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62102 -->
  <question type="stack">
    <name>
				<text>23.2.2.8 Orthogonality Relation 2</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[We summarise these results in the following Key Point:
<p></p>
<div class="HELM_keypoint">
<h4>Key Point</h4>
<p></p>
For any integers \(m\), \(n\),
\[ \displaystyle{ \int _{-\pi}^\pi } \sin nt \cos mt \hspace{0.15cm} dt = 0 \]
<p></p>
\[  \displaystyle{ \int _{-\pi}^\pi } \sin nt \cos mt \hspace{0.15cm} dt =
\begin{cases}
0 \hspace{1cm} n \neq m \\
\pi \hspace{1cm} n = m \neq 0 \\
2\pi \hspace{0.8cm} n = m = 0
\end{cases}
\]
<p></p>
\[  \displaystyle{ \int _{-\pi}^\pi } \sin nt \sin mt \hspace{0.15cm} dt =
\begin{cases}
0 \hspace{1cm} n \neq m, \hspace{0.2cm} n = m = 0 \\
\pi \hspace{1cm} n = m \neq 0
\end{cases}
\]
All these results hold for any integration range of width \(2\pi\).
</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200815205135+t6SZh6</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220518093136+J9n82u</text></tag>
    </tags>
    
</question>
</quiz>