<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24634 -->
  <question type="stack">
    <name>
				<text>18.2.1.1 First partial derivatives - x derivative</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span>First partial derivatives</span><br></h3>

<p><strong>The \(x\) partial derivative</strong></p>

<p>For a function of a single variable, \(y=f(x)\), changing the independent variable \(x\) leads to a corresponding change in the dependent variable \(y\). The <strong>rate of change</strong> of \(y\) with respect to \(x\) is given by the <strong>derivative</strong>,
    written \(\dfrac{\mathrm{d}f}{\mathrm{d}x}\). A similar situation occurs with functions of more than one variable. For clarity we shall concentrate on functions of just two variables. In the relation \(z=f(x,y)\), the <strong>independent variables</strong>    are \(x\) and \(y\), and the <strong>dependent variable</strong> is \(z\). We have seen in Section 18.1 that as \(x\) and \(y\) vary the \(z\)-value traces out a surface. Now both of the variables \(x\) and \(y\) may change <em>simultaneously</em>    inducing a change in \(z\). However, rather than consider this general situation, to begin with we shall hold one of the independent variables <strong>fixed</strong>. This is equivalent to moving along a curve obtained by intersecting the surface
    by one of the coordinate planes.</p>
<p>Consider \(f(x,y)=x^3+2x^2y+y^2+2x+1.\)</p>
<p>Suppose we keep \(y\) constant and vary \(x\); then what is the rate of change of the function \(f\)?</p>
<p> Suppose we hold \(y\) at the value \(3\) then</p>
\[\begin{aligned} f(x,3) &amp;= x^3+6x^2+9+2x+1 \\ &amp;= x^3+6x^2+2x+10 \end{aligned}\]

<p>In effect, we now have a function of \(x\) only. If we differentiate it with respect to \(x\) we obtain the expression: \(3x^2+12x+2\).</p>
<p>We say that \(f\) has been <strong>partially differentiated</strong> with respect to \(x\). We denote the partial derivative of \(f\) with respect to \(x\) by \(\dfrac{\partial f}{\partial x}\) (to be read as ‘partial dee \(f\) by dee \(x\)’ ). In this
    example, when \(y=3\):</p>
\[\dfrac{\partial f}{\partial x} = 3x^2+12x+2\]
<p>In the same way if \(y\) is held at the value 4 then \(f(x,4)=x^3+8x^2+16+2x+1=x^3+8x^2+2x+17\) and so, for this value of \(y\),</p>
\[\dfrac{\partial f}{\partial x} = 3x^2+16x+2.\]
<p></p>

<p>Now if we return to the original formulation \(f(x,y)=x^3+2x^2y+y^2+2x+1\) and treat \(y\) as a constant then the process of partial differentiation with respect to \(x\) gives</p>
\[ \begin{aligned} \frac{\partial f}{\partial x} &amp;= 3x^2+4xy+0+2+0 \\ &amp;= 3x^2+4xy+2. \end{aligned}\]

<div class="HELM_keypoint">
    <h4>Key Point</h4>
    <p><strong>The Partial Derivative of  \(f\) with respect to \(x\)</strong></p>
    <p> For a function of two variables \(z=f(x,y)\) the partial derivative of \(f\) with respect to \(x\) is denoted by \(\dfrac{\partial f}{\partial x}\) and is obtained by differentiating \(f(x,y)\) with respect to \(x\) in the usual way but treating the
        \(y\)-variable as if it were a constant.</p>
    <p>Alternative notations for \(\dfrac{\partial f}{\partial x}\) are \(f_x(x,y)\) or \(f_x\) or \(\dfrac{\partial z}{\partial x}\)
    </p>
</div>

<hr>
<h4 class="HELM_example">Example</h4>
<p>Find \(\dfrac{\partial f}{\partial x}\) for</p>
<ol class="HELM_parts_inline">
    <li> \(f(x,y)=x^3+x+y^2+y\), </li>
    <li> \(f(x,y)=x^2y+xy^3\).</li>
</ol>

<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
    <li>\(\dfrac{\partial f}{\partial x}=3x^2+1+0+0=3x^2+1\)</li>
    <li>\(\dfrac{\partial f}{\partial x}=2x\cdot y+1\cdot y^3=2xy+y^3\)</li>
</ol>
<p></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200818045608+iczQLI</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220304105455+2sQCEc</text></tag>
    </tags>
    
</question>
</quiz>