<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62171 -->
  <question type="stack">
    <name>
				<text>43.2.06 Comments on interpretation</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>Comments on interpretation</h3>

Some care should always be taken when interpreting results obtained from correlation coefficient
calculations.
<p></p>
(a) A high correlation does not necessarily imply that a causal relationship exists between
the variables considered. For example, it may be that a high degree of correlation exists
between the number of road accidents in a particular city and the number of late trains
arriving at a station in another city both over the same time period. In general one would
not expect to find a causal relation between the variables involved. Similar comments
apply to, for example, water hardness and average income for towns in the UK.

<p></p>
(b) When considering the behaviour of two variables, one should realize that it is possible
that both variables may change because of the influence of a third variable. An example
often quoted in this context is the Gas law
\[\dfrac{PV}{T} = \text{constant} \]
where say, pressure and volume may change because of a change in temperature.

<p></p>
(c) A low value of the correlation coefficient does not necessarily imply that no relationship
exists between the variables being considered. Remember that the correlation coefficient 
is indicative of a <i>linear relationship only</i> and that a low or zero value of \(r\) may indicate
that a non-linear relationship exists. For example a set of points lying on the curve \(y = x^2\)<i>
might</i> result in a zero value of \(r\).]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200707193307+ZdwSVe</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200708183850+VOWDNY</text></tag>
    </tags>
    
</question>
</quiz>