<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26699 -->
  <question type="stack">
    <name>
				<text>10.1.1.10 Complex Conjugate</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>If <span class="math inline">\(z=a+i b\)</span> is any complex number then the complex conjugate of <span class="math inline">\(z\)</span> is denoted by <span class="math inline">\(z^*\)</span> and is defined by <span class="math inline">\[z^*=a-i b.\]</span>(Sometimes the notation <span class="math inline">\(\bar z\)</span> is used instead of <span class="math inline">\(z^*\)</span> to denote the conjugate). </p>

<p>For example if <span class="math inline">\(z=2-3i\)</span> then <span class="math inline">\(z^*=2+3i\)</span>.
</p>

<p> If <span class="math inline">\(z\)</span> is entirely real then <span class="math inline">\(z^*=z\)</span> whereas if <span class="math inline">\(z\)</span> is entirely imaginary then <span class="math inline">\(z^*=-z\)</span>. E.g. if <span class="math inline">\(z=17i\)</span> then <span class="math inline">\(z^*=-17i\)</span>.
</p>

<p> In fact the following relationships are easily obtained: <span class="math inline">\[\qquad \text{Re}(z)={z+z^*\over 2}\qquad \text{and}\qquad \text{Im}(z)={i(z^*-z)\over {2}}\]</span></p>

<p>Complex numbers are added, subtracted, multiplied and divided in much the same way as these operations are carried out for real numbers.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200617091748+O4uBUR</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210706164639+iCAtzD</text></tag>
    </tags>
    
</question>
</quiz>