<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27338 -->
  <question type="stack">
    <name>
				<text>38.1.1.6 Cumulative Distribution Function</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p></p><p>For the pdf in the diagram below<br></p><p style="text-align: center;"><img src="@@PLUGINFILE@@/38.1%20figure%208.svg" alt="Graph of a function where: the line y=0 everywhere except for two line segments - one from (-2,0) to (0,1/2), and the second one from (0,1/2) to (2,0)." width="378" height="159" role="presentation" class="img-responsive atto_image_button_text-bottom"><br></p><p>obtain the cumulative distribution function:</p><p>\( F(x) = \begin{cases} 0 &amp; \text{for all $ x \leq -2$} \\ g(x) &amp; \text{for all $ -2&lt; x \leq 0$} \\ h(x) &amp; \text{for all $ 0 \leq x \leq 2$} \\&nbsp;1 &amp; \text{for all $ x \geq 2$} \end{cases}. \)<br></p><p>\(g(x) = \)&nbsp;[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p><p>\(h(x) = \)&nbsp;[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_solution">Solution</h4>
<p>For \(&nbsp;-2&lt; x \leq 0\), we have that</p><p>\[\begin{align*}F(x)&amp;=\displaystyle{\int_{-2}^x} \dfrac{1}{2}+\dfrac{u}{4} \; du\\&amp; = \left.\dfrac{u}{2}\right|_{-2}^x + \left.\dfrac{u^2}{8}\right|_{-2}^x\\&amp; = {@ta1@}.\end{align*}\]</p><p>Hence,&nbsp;<span style="font-size: 0.9375rem;">\(g(x) = \displaystyle{{@ta1@}}.\)</span></p><p>For \(&nbsp;0 \leq x &lt; 2\), we have that<br></p><p>\[\begin{align*}F(x) &amp; = F(0) + \displaystyle{\int_0^x} \dfrac{1}{2}-\dfrac{u}{4} \; du \\&amp; = \dfrac{1}{2} + \left.\dfrac{u}{2}\right|_{0}^x&nbsp; -&nbsp; \left.\dfrac{u^2}{8}\right|_{0}^x\\&amp; ={@ta2@}.\end{align*}\]<br></p><p></p><p>Hence,&nbsp;<span style="font-size: 0.9375rem;">\(h(x) = \displaystyle{{@ta2@}}.\)</span></p><br><p></p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021072200</text>
		</stackversion>
    <questionvariables>
				<text>ta1:1/2 +1/2*x + 1/8*x^2;
ta2:1/2 +1/2*x - 1/8*x^2;</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>ONLY ONE VARIANT!</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>Algebraic expression.</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords>integrate </forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>Algebraic expression.</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords>integrate</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200730110922+fQJ5Kv</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+211013104826+bgHOd7</text></tag>
    </tags>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>