<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26739 -->
  <question type="stack">
    <name>
				<text>10.2.2.4 Polar form</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Consider the complex number \[z={@z1@} \] Evaluate the modulus of \(z\) \((r = |z|)\) and the principal value of the argument of \(z\ (\theta = \mathrm{Arg}(z)\)). Enter your results in the correct order, i.e. \([r,\theta]\).</p>
<p>&nbsp;</p>
<p>\([r,\theta] = \) [[input:ans1]] [[validation:ans1]]&nbsp;[[feedback:Result]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>\[r=\sqrt{x^2+y^2}\] \[r=\sqrt{\left({@x@}\right)^2+\left({@y@}\right)^2}=\sqrt{{@x*x@}+{@y*y@}}\] \[r=\sqrt{@x*x+y*y@}={@rval@}\] \[\tan(\mathrm{arg}(z))=\frac{y}{x}=\frac{@y@}{@x@}={@y/x@}\] \[\mathrm{arg}(z)={@argsol@} + k\cdot\pi\] Considering the appropriate quadrant of the complex plane to fix \(k\), we find that \[\mathrm{arg}(z)={@principaltheta@}\].</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>argset:[Pi/6,Pi/4,Pi/3];
 n1:rand(3)+1;
 modsol:rand(5)+1;
 argsol:ev(((-1)^rand(2))*(rand(2)*Pi-argset[n1]),simp);
 z1:ev(modsol*cos(argsol) + modsol*%i*sin(argsol),simp);
 x:modsol*cos(argsol);
 y:modsol*sin(argsol);
 rval:sqrt(x*x+y*y);
 principaltheta:%pi - mod (%pi - argsol, 2*%pi);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(z={@z1@}\)  gives \([{@modsol@},{@principaltheta@}]\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>[modsol,principaltheta]</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>Result</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>[modsol,principaltheta]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-0-T </trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>Result-0-F </falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200629215234+1R2zbA</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210701151758+KMNXvg</text></tag>
    </tags>
    <deployedseed>1124288821</deployedseed>
    <deployedseed>726390547</deployedseed>
    <deployedseed>1676778895</deployedseed>
    <deployedseed>592831551</deployedseed>
    <deployedseed>937052114</deployedseed>
    <deployedseed>1815171070</deployedseed>
    <deployedseed>1314332238</deployedseed>
    <deployedseed>2041689355</deployedseed>
    <deployedseed>592139395</deployedseed>
    <deployedseed>1849464186</deployedseed>
    <deployedseed>1315253392</deployedseed>
    <deployedseed>907627411</deployedseed>
    <deployedseed>223248315</deployedseed>
    <deployedseed>444062909</deployedseed>
    <deployedseed>2133593970</deployedseed>
    <deployedseed>262730417</deployedseed>
    <deployedseed>1871523533</deployedseed>
    <deployedseed>578950768</deployedseed>
    <deployedseed>281850306</deployedseed>
    <deployedseed>1815742935</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>[modsol,principaltheta]</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>Result-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>[x,y]</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>Result-0-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>