<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26325 -->
  <question type="stack">
    <name>
				<text>2.3.1.1b Many-to-one and one-to-one</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Many-to-one and one-to-one functions</h4>
<p>Consider the function <span class="math inline">\(y(x)=x^2\)</span>. An input of <span class="math inline">\(x=3\)</span> produces an output of \(9\). Similarly, an input of <span class="math inline">\(-3\)</span> also produces an output of \(9\). In general, a function for which different inputs can produce the same output is called a <b>many-to-one function</b>. This is represented pictorially in the figure below from which it is clear why we call this a many-to-one function.</p>
<figure class="HELM_figure">
<img src="@@PLUGINFILE@@/block4-fig3-1.svg" alt="A many-to-one function" class="img-responsive atto_image_button_text-bottom" width="215" height="125">
<figcaption>This represents a many-to-one function</figcaption>
</figure>
<p>Note that whilst this is many-to-one it is still a function since any chosen input value has only one arrow emerging from it. Thus there is a single output for each input.</p>
<p>It is possible to decide if a function is many-to-one by examining its graph. Consider the graph of <span class="math inline">\(y=x^2\)</span> shown below.</p>
<figure class="HELM_figure">
<img src="@@PLUGINFILE@@/block4-fig4-1.svg" alt="y=x^2" class="img-responsive atto_image_button_text-bottom" width="209" height="132">
<figcaption>The function <span class="math inline">\(y=x^2\)</span> is a many-to-one function</figcaption>
</figure>
<p>We see that a horizontal line drawn on the graph cuts it more than once. This means that two (or more) different inputs have yielded the same output and so the function is many-to-one.</p>
<p>If a function is not many-to-one then it is said to be <b>one-to-one</b>. This means that each different input to the function yields a different output.</p>
<p>Consider the function <span class="math inline">\(y(x)=x^3\)</span> which is shown in the figure below. A horizontal line drawn on this graph will intersect the curve only once. This means that each input value of <span class="math inline">\(x\)</span> yields a different output value for <span class="math inline">\(y\)</span>.</p>
<figure class="HELM_figure">
<img src="@@PLUGINFILE@@/block4-fig5-1.svg" alt="y=x^3" class="img-responsive atto_image_button_text-bottom" width="224" height="148">
<figcaption>The function <span class="math inline">\(y(x)=x^3\)</span> is a one-to-one function</figcaption>
</figure>
<p><i>Try for yourself</i>: By sketching a graph of \(y=3x-1\) show that this is a one-to-one function.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200616143215+TFENpg</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210628123416+TI1KqF</text></tag>
    </tags>
    
</question>
</quiz>