<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24662 -->
  <question type="stack">
    <name>
				<text>18.1.1.1 Functions of Several Variables</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span>Functions of several variables</span><br></h3>

<p>We know that \(f(x)\) is used to represent a function of one variable: the input variable is \(x\) and the output is the value \(f(x)\). Here \(x\) is the <strong>independent</strong> variable and \(y=f(x)\) is the <strong>dependent variable</strong>.
    Suppose we consider a function with <strong>two independent</strong> input variables \(x\) and \(y\), for example \(f(x,y)=x+2y+3\). If we specify values for \(x\) and \(y\) then we have a <strong>single</strong> value \(f(x,y)\). For example, if
    \(x=3\) and \(y=1\) then \(f(x,y)=3+2+3=8\). We write \(f(3,1)=8\).</p>

<hr>
<h4 class="HELM_example">Example</h4>
<p>Find the values of \(f(2,1),\ f(-1,-3)\) and \(f(0,0)\) for the following functions.</p>
<ol class="HELM_parts">

    <li> \(f(x,y)=x^2+y^2+1\), </li>
    <li> \(f(x,y)=2x+xy+y^3\).</li>

</ol>
<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
    <li>\(f(2,1)=2^2+1^2+1=6\), \(f(-1,-3)=(-1)^2+(-3)^2+1=11\), \(f(0,0)=1\).</li>
    <li>\(f(2,1)=4+2+1=7\), \(f(-1,-3)=-2+3-27=-26\), \( f(0,0)=0\)</li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200814083729+SAdtsC</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+201215145033+CyusAu</text></tag>
    </tags>
    
</question>
</quiz>