<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26520 -->
  <question type="stack">
    <name>
				<text>6.2.4.2 Express sinh, cosh in terms of exp (exercise 1(b))</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Express the following in terms of \(e^{{@aa@}{@xx@}}\) and \(e^{-{@aa@}{@xx@}}\).</p>
<p>\(&nbsp; {@uu@} \sinh({@aa@}{@xx@}) - {@vv@} \cosh({@aa@}{@xx@}) = \) [[input:ans1]] [[validation:ans1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>\[\begin{align*} &amp; {@uu@} \sinh({@aa@}{@xx@}) - {@vv@} \cosh({@aa@}{@xx@})  \\ &amp; =
 {@uu@} \cdot \frac{1}{2} (e^{{@aa@}{@xx@}} - e^{-{@aa@}{@xx@}}) - {@vv@} \cdot \frac{1}{2} (e^{{@aa@}{@xx@}} + e^{-{@aa@}{@xx@}}) \\ &amp; =  {@ta@} \end{align*}\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>aa: rand(6)+2;
uu: rand(6)+2;
vv: rand_with_prohib(2,7,[uu]);
xx: rand([x,y,z,t]);
sinh: (e^(aa*xx)-e^(-aa*xx))/2;
cosh: (e^(aa*xx)+e^(-aa*xx))/2;
sum: uu*sinh-vv*cosh;
simp:true;
ta: expand(sum);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>\( {@uu@} \sinh({@aa@}{@xx@}) - {@vv@} \cosh({@aa@}{@xx@}) = {@ta@} \)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>25</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>sinh,cosh,[[BASIC-ALGEBRA]]</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200625185544+bklOIF</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210706153646+lpHnok</text></tag>
    </tags>
    <deployedseed>1438680074</deployedseed>
    <deployedseed>778343189</deployedseed>
    <deployedseed>1839598201</deployedseed>
    <deployedseed>2096539829</deployedseed>
    <deployedseed>1016869893</deployedseed>
    <deployedseed>1733530001</deployedseed>
    <deployedseed>1195960102</deployedseed>
    <deployedseed>654927053</deployedseed>
    <deployedseed>240186691</deployedseed>
    <deployedseed>1745639417</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(ta+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>