<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24690 -->
  <question type="stack">
    <name>
				<text>19.1.3.3  Classify the differential equation</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p></p><p>Classify the differential equation \[ \dfrac{d^2x}{ dt^2}+{@n1@} \dfrac{dx}{dt}= {@n2*x@}.\]&nbsp;</p><p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p dir="ltr" style="text-align: left;"><span style="font-size: 0.9375rem;">We are given the differential equation \[ \dfrac{d^2x}{ dt^2}+{@n1@} \dfrac{dx}{dt}= {@n2*x@}.\]&nbsp;</span><br></p><p dir="ltr" style="text-align: left;">The highest order of the derivatives is two and there is no non-linear term.<br></p><p dir="ltr" style="text-align: left;">So the differential equation is second order and non-linear.&nbsp;</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>n1: rand(6)+2;
n2: rand_with_prohib(2,8,[n1]);
ta: [[a, false, "First order and linear"], [b, true, "Second order and linear"], [c, false, "First order and non-linear"], [d, false, "Second order and non-linear"]];
ta:random_permutation(ta);
ta1: b;</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>\(  \dfrac{d^2x}{ dt^2}+{@n1@} \dfrac{dx}{dt}= {@n2*x@} \)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:moodle2.maths.ed.ac.uk+210319151347+nWukcu</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+210319154054+TfnBM2</text></tag>
    </tags>
    <deployedseed>426007004</deployedseed>
    <deployedseed>985794878</deployedseed>
    <deployedseed>1153137814</deployedseed>
    <deployedseed>753613890</deployedseed>
    <deployedseed>730905794</deployedseed>
    <deployedseed>1691980979</deployedseed>
    <deployedseed>1529385762</deployedseed>
    <deployedseed>1276755517</deployedseed>
    <deployedseed>1047242602</deployedseed>
    <deployedseed>1069706208</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_incorrect(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>