<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62098 -->
  <question type="stack">
    <name>
				<text>23.2.2.4 Orthogonality properties of sinusoids p2</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[The results of the integrals in the previous exercise are 
\[ \displaystyle{ \int _{-\pi}^\pi }\sin nt \hspace{0.15cm} dt = 0 \]
and 
\[ \displaystyle{ \int _{-\pi}^\pi }\cos nt \hspace{0.15cm} dt = 0 \]
where \(n\) is an integer. 
<p></p>
Note that any integration range \(\alpha\), \(\alpha + 2\pi\), would give these same (zero) answers.
<p></p>
These integrals enable enable us to calculate the constant term in the Fourier series given previously as 
\[ f(t) = \dfrac{a_0}{2} + \sum _{n=1}^{\infty} \left( a_n \cos nt + b_n \sin nt \right)\]
Integrating both sides of the Fourier series from \( -\pi\) to \(\pi\) gives the left hand side as
\[ \displaystyle{ \int _{-\pi}^\pi } f(t) \hspace{0.15cm} dt \]
(whose value clearly depends on the function \(f(t)\)).
Integrating the right-hand side term by term we get
\[ \dfrac{1}{2} \displaystyle{ \int _{-\pi}^\pi } a_0 \hspace{0.15cm} dt 
+  \sum _{n=1}^\infty \left\{ \int _{-\pi}^\pi  a_n \cos nt \hspace{0.15cm} dt +  \int _{-\pi}^\pi  b_n \sin nt \hspace{0.15cm} dt  \right\} = \dfrac{1}{2} \Big[ a_0 t \Big] _{-\pi}^\pi + \sum_{n=1}^\infty \left\{ 0 + 0 \right\} \]
Thus we get
\[  \displaystyle{ \int _{-\pi}^\pi } f(t) \hspace{0.15cm} dt = \dfrac{1}{2} (2 a_0 \pi ) \]
or
\[ a_0  = \dfrac{1}{\pi}\displaystyle{ \int _{-\pi}^\pi } f(t) \hspace{0.15cm} dt  \]
<p></p>
<div class="HELM_keypoint">
<h4>Key Point</h4>
The constant term in a trigonometric Fourier series for a function of period \(2\pi\) is
\[ \dfrac{a_0}{2} = \dfrac{1}{2\pi} \displaystyle{ \int _{-\pi}^\pi } f(t) \hspace{0.15cm} dt \]
which is the average value of \(f(t)\)  over 1 period. 
</div>
<p></p>
This result ties in with our earlier discussion on the significance of the constant term. Clearly a signal
whose average value is zero will have no constant term in its Fourier series. The following square
wave shown in the Figure below is an example.<p style="text-align: center;"><img src="@@PLUGINFILE@@/23-2-2-fig-9-square-pulse.svg" alt="square-pulse" width="494" height="156" class="img-responsive atto_image_button_text-bottom"><br></p>
<p></p>
We now obtain further integrals, known as orthogonality properties, which enable us to find the
remaining Fourier coefficients i.e. the amplitudes \(a_n\) and \(b_n\) \( (n = 1, 2, 3, \dots ) \) of the sinusoids.]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200815151614+t1L4FW</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220518093057+dlsE4n</text></tag>
    </tags>
    
</question>
</quiz>