<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24676 -->
  <question type="stack">
    <name>
				<text>19.1.3.1 Classifying differential equations</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Classifying differential equations</h4>
<p>When solving differential equations (either analytically or numerically) it is important to be able to recognise the various kinds that can arise. We therefore need to introduce some terminology which will help us to distinguish one kind of differential equation from another.</p>
<ol>
<li><p>An <span><strong>ordinary differential equation</strong></span> (ODE) is any relation between a function of a <span><strong>single variable</strong></span> and its derivatives. (All differential equations studied in this workbook are ordinary.)</p></li>
<li><p>The <strong>order</strong> of a differential equation is the order of the highest derivative in the equation.</p></li>
<li><p>A differential equation is <strong>linear</strong> if the dependent variable and its derivatives occur to the first power only and if there are no products involving the dependent variable or its derivatives.</p></li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201215164009+QYaF0T</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+201215164051+dbZrqC</text></tag>
    </tags>
    
</question>
</quiz>