<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24420 -->
  <question type="stack">
    <name>
				<text>11.2.2.1 Extending the table of derivatives</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Extending the table of derivatives</h4>
<p></p>
We now quote simple rules which enable us to extend the range of functions which we can differentiate. The first two rules are for differentiating sums or differences of functions. The reader should note that all of the rules quoted below can be obtained from first principles using the approach outlined previously.<p></p><p></p><div class="HELM_keypoint">
<h4>Key Point</h4>
<p></p>
Rule 1: The derivative of \(f(x) + g(x)\) is \( \dfrac{df}{dx} + \dfrac{dg}{dx} \) i.e. \[\dfrac{d}{dx} (f(x)+g(x))=\dfrac{df}{dx} + \dfrac{dg}{dx}. \]<p></p>
Rule 2: The derivative of \(f(x) - g(x)\) is \( \dfrac{df}{dx} - \dfrac{dg}{dx} \)&nbsp;i.e. \[\dfrac{d}{dx} (f(x)-g(x))=\dfrac{df}{dx} - \dfrac{dg}{dx} .\]</div><p></p>
These rules say that to find the derivative of the sum (or difference) of two functions, we simply
calculate the sum (or difference) of the derivatives of each function.
<hr>
<h4 class="HELM_example">Example</h4>
<p>Find the derivative of \(y = x^6 + x^4\).</p>
<h4 class="HELM_solution">Solution</h4>
<p>We simply calculate the sum of the derivatives of each separate function:
\[ \dfrac{dy}{dx} = 6x^5 + 4x^3\]</p><hr>

The third rule tells us how to differentiate a multiple of a function. We have already met and applied particular cases of this rule which appear in the Table of derivatives.

<p></p><div class="HELM_keypoint">
<h4>Key Point</h4>
<p></p>
Rule 3: If \(k\) is a constant, the derivative of \(kf(x)\) is \(k \dfrac{df}{dx}  \)&nbsp;i.e. \(\dfrac{d}{dx} (kf(x))=k\dfrac{df}{dx} \).</div><p></p>
This rule tells us that if a function is multiplied by a constant, \(k\), then the derivative is also multiplied by the same constant, \(k\).

<hr>
<h4 class="HELM_example">Example</h4>
<p>Find the derivative of \(y = 8e^{2x} \).</p>
<h4 class="HELM_solution">Solution</h4>
<p> Here we are interested in differentiating a multiple of the function \(e^{2x}\). We differentiate \(e^{2x}\), giving \(2e^{2x}\), and multiply the result by \(8\). Thus 
\[ \dfrac{dy}{dx} = 8 \times 2e^{2x} = 16 e^{2x} \]
</p><hr>

<hr>
<h4 class="HELM_example">Example</h4>
<p>Find the derivative of \(y = 6\sin 2x  + 3x^2 - 5e^{3x} \).</p>
<h4 class="HELM_solution">Solution</h4>
<p> We differentiate each part of the function in turn.
\begin{align}
y &amp;= 6\sin 2x  + 3x^2 - 5e^{3x} \\
\dfrac{dy}{dx} &amp;= 6(2\cos 2x) + 3(2x) - 5(3e^{3x}) \\
&amp;= 12\cos 2x + 6x - 15e^{3x}
\end{align}
</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200725062341+U3QM9p</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210901193354+oyi6nN</text></tag>
    </tags>
    
</question>
</quiz>