<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27182 -->
  <question type="stack">
    <name>
				<text>35.1.4.1 Complement, empty set exercise (exercise 1)</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Given a set \(A\), its complement \(A'\) and a universal set \(S\), which of the following expressions are true?</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p dir="ltr" style="text-align: left;">The complement set \(A'\) consists of all elements in the universal set that are not in set \(A\), so \(A \cup A' = S\) and \(A \cap A' = \emptyset\). Also, the subset sequence \(\emptyset \subset A \subset S\) means means that a union between any pair equates to the larger set, \(A \cup \emptyset = A\), and an intersection equates to the smaller set, \(A \cap S = A\).</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200702122328+M25Q9g</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+211004105034+XQV7nz</text></tag>
    </tags>
    
</question>
</quiz>