<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24606 -->
  <question type="stack">
    <name>
				<text>14.2.2.1 Root-mean-square value of a function</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>Root-mean-square value of a function</h3>

<p>If \(f(t)\) is defined on the interval \(a\le t\le b\), the <strong>mean-square</strong> value is given by the expression: \[ {\frac{1}{b-a}\int_a^b [f(t)]^2\,\mathrm{d}t}\] This is simply the mean value of \([f(t)]^2\) over the given interval. The related quantity:
    the <strong>root-mean square (r.m.s.)</strong> value is given by the following formula.</p><br>

<div class="HELM_keypoint">
    <h4>Key Point</h4>
    <p>\[\text{r.m.s. value } =~ \sqrt{\frac{1}{b-a} \int_a^b [f(t)]^2\ \mathrm{d}t}\]</p>
</div><br>

<p>The r.m.s. value depends upon the interval chosen. If the values of \(a\) or \(b\) are changed, then the r.m.s. value of the function across the interval from \(a\) to \(b\) will in general change as well. Note that when finding an r.m.s. value the function
    must be squared before it is integrated.</p>

<p>&nbsp;</p>
<hr>
<h4 class="HELM_example">Example</h4>
<p>Find the r.m.s. value of \(f(t)=t^2\) across the interval from \(t=1\) to \(t=3\).</p><br>
<h4 class="HELM_solution">Solution</h4>
<p>\(\displaystyle \text{r.m.s. }=~\sqrt{\frac{1}{b-a} \int_a^b [f(t)]^2\ \mathrm{d}t} ~=~ \sqrt{\frac{1}{3-1} \int_1^3 [t^2]^2\ \mathrm{d}t} ~=~ \sqrt{\frac{1}{2} \int_1^3 t^4\ \mathrm{d}t} ~=~ \sqrt{\frac{1}{2} \left[\frac{t^5}{5}\right]_1^3} \approx 4.92\)</p>

<p>&nbsp;</p>
<hr>
<h4 class="HELM_example">Example</h4>
<p>Calculate the r.m.s value of \(f(t)=\sin t\) across the interval \(0\le t\le 2\pi\).</p><br>
<h4 class="HELM_solution">Solution</h4>
<p>Here \(a=0\) and \(b=2\pi\) so r.m.s. \(\displaystyle =~\sqrt{\frac{1}{2\pi} \int_0^{2\pi} \sin^2 t ~\ \mathrm{d}t}\).</p>
<p>The integral of \(\sin^2 t\) is performed by using trigonometric identities to rewrite it in the alternative form \(\dfrac{1}{2}(1-\cos 2t)\). This technique was described before, as we hope you will remember.</p>
<p>\[\text{r.m.s. value }=~ \sqrt{\frac{1}{2\pi} \int_0^{2\pi} \frac{(1-\cos 2t)}{2}\ \mathrm{d}t}~=~ \sqrt{\frac{1}{4\pi} \left[t-\frac{(\sin 2t)}{2}\right]_0^{2\pi}} ~=~ \sqrt{\frac{1}{4\pi}(2\pi)} ~=~ \sqrt{\frac{1}{2}} \approx 0.707\]</p><br>
<hr>

<p>In the previous Example the amplitude of the sine wave was 1, and the r.m.s. value was 0.707. In general, if the amplitude of a sine wave is \(A\), its r.m.s value is 0.707\(A\).</p><br>

<div class="HELM_keypoint">
    <h4>Key Point</h4>
    <p>The r.m.s value of any sinusoidal waveform taken across an interval of width equal to one period is 0.707 \(\times\) amplitude of the waveform.</p>
</div><br>

<hr>
<h4 class="HELM_example">Engineering Example</h4>

<h5>Electrodynamic meters</h5>
<h6>Introduction</h6>

<p>A dynamometer or electrodynamic meter is an analogue instrument that can measure d.c. current or a.c. current up to a frequency of 2 kHz. A typical dynamometer is shown in the figure below.</p>
<p>It consists of a circular dynamic coil positioned in a magnetic field produced by two wound circular stator coils connected in series with each other. The torque \(T\) on the moving coil depends upon the mutual inductance between the coils given by: \[
    T=I_1I_2\frac{\mathrm{d}M}{\mathrm{d}\theta}\] where \(I_1\) is the current in the fixed coil, \(I_2\) the current in the moving coil and \(\theta\) is the angle between the coils. The torque is therefore proportional to the square of the current. If the current is
    alternating the moving coil is unable to follow the current and the pointer position is related to the mean value of the square of the current. The scale can be suitably graduated so that the pointer position shows the square root of this value, i.e.
    the r.m.s. current.</p><p style="text-align: center;"><img src="@@PLUGINFILE@@/14-2-fig-7-meter.svg" alt="Representation of an electrodynamic meter." width="297" height="220" class="img-responsive atto_image_button_text-bottom"><br></p><br>

<h6>Problem in words</h6>
<p>A dynamometer is in a circuit in series with a 400 \(\Omega\) resistor, a rectifying device and a 240 V r.m.s alternating sinusoidal power supply. The rectifier resists current with a resistance of 200 \(\Omega\) in one direction and a resistance of 1
    k\(\Omega\) in the opposite direction. Calculate the reading indicated on the meter.</p><br>

<h6>Mathematical Statement of the problem</h6>
<p>We know from the Key Point in the text that the r.m.s. value of any sinusoidal waveform taken across an interval equal to one period is 0.707 \(\times\) amplitude of the waveform. Where 0.707 is an approximation of \(\dfrac{1}{\sqrt{2}}\). This allows
    us to state that the amplitude of the sinusoidal power supply will be: \[ V_{\text{peak}}=\frac{V_{\text{rms}}}{\frac{1}{\sqrt{2}}}=\sqrt{2}V_{\text{rms}}\] In this case the r.m.s power supply is 240 V so we have \[ V_{\text{peak}}=240\times \sqrt{2}=339.4\
    V\] During the part of the cycle where the voltage of the power supply is positive the rectifier behaves as a resistor with resistance of 200 \(\Omega\) and this is combined with the 400 \(\Omega\) resistance to give a resistance of 600 \(\Omega\)
    in total. Using Ohm’s law \[ V=IR\Rightarrow I=\frac{V}{R}\] As \(V=V_{\text{peak}}\sin(\theta)\) where \(\theta=\omega t\) where \(\omega\) is the angular frequency and \(t\) is time we find that during the positive part of the cycle \[ I_{\text{rms}}^2=\frac{1}{2\pi}\int_0^\pi\left(\frac{339.4\sin(\theta)}{600}\right)^2\ \mathrm{d}\theta\] During the part of the cycle where the voltage of the power supply is negative the rectifier behaves as a resistor with resistance of 1 k\(\Omega\) and this is combined with the 400 \(\Omega\) resistance to give 1400 \(\Omega\) in total.
    So we find that during the negative part of the cycle \[ I_{\text{rms}}^2=\frac{1}{2\pi}\int_\pi^{2\pi}\left(\frac{339.4\sin(\theta)}{1400}\right)^2\ \mathrm{d}\theta\] Therefore over an entire cycle \[ I_{\text{rms}}^2=\frac{1}{2\pi}\int_0^\pi \left(\frac{339.4\sin(\theta)}{600}\right)^2\ \mathrm{d} \theta+\frac{1}{2\pi}\int_\pi^{2\pi}\left(\frac{339.4\sin(\theta)}{1400}\right)^2\ \mathrm{d}\theta\] We can calculate this value to find \(I_{\text{rms}}^2\) and therefore \(I_{\text{rms}}\).</p><br>

<h6>Mathematical analysis</h6>
<p>\[ I_{\text{rms}}^2=\frac{1}{2\pi}\int_0^\pi\left(\frac{339.4\sin(\theta)}{600}\right)^2\ \mathrm{d}\theta+\frac{1}{2\pi}\int_\pi^{2\pi}\left(\frac{339.4\sin(\theta)}{1400}\right)^2\ \mathrm{d}\theta \]\[ I_{\text{rms}}^2=\frac{339.4^2}{2\pi\times 10000} \left(\int_0^\pi\frac{\sin^2(\theta)}{36}\ \mathrm{d}\theta+ \int_\pi^{2\pi}\frac{\sin^2(\theta)}{196}\ \mathrm{d}\theta\right)\] Substituting the trigonometric identity \(\sin^2(\theta)\equiv \dfrac{1-\cos(2\theta)}{2}\) we get \[\begin{aligned} I_{\text{rms}}^2&amp;=\frac{339.4^2}{4\pi\times 10000}\left(\int_0^\pi
    \frac{1-\cos(2\theta)}{36}\ \mathrm{d}\theta+\int_\pi^{2\pi}\frac{1-\cos(2\theta)}{196}\ \mathrm{d}\theta\right) \\ &amp;=\frac{339.4^2}{4\pi\times 10000}\left(\left[\frac{\theta}{36}-\frac{\sin(2\theta)}{72} \right]_0^\pi+\left[\frac{\theta}{196}-\frac{\sin(2\theta)}{392}\right]_\pi^{2\pi}\right)
    \\ &amp;=\frac{339.4^2}{4\pi\times 10000}\left(\frac{\pi}{36}+\frac{\pi}{196}\right)=0.0946875\ A^2 \\ I_{\text{rms}}&amp;=0.31\ A\ \text{to}\ 2 \ \text{d.p.}\end{aligned}\]</p><br>

<h6>Interpretation</h6>
<p>The reading on the meter would be 0.31 A.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201128174844+a9Mk2n</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210607154843+6HAh7C</text></tag>
    </tags>
    
</question>
</quiz>