<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25742 -->
  <question type="stack">
    <name>
				<text>2.4.2.* Piecewise continuity</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Study the graphs of the following functions for \(x \in \mathbb{R}\) and fill in the blanks:</p><p>\(f(x)= x^3-1\) for \(x&lt;1\) and \(=x^2-1\) for \(x \geq 1\) is [[1]] and</p><p></p>\(f(x)= x^3-1\) for \(x&lt;-1\) and \(=x^2-1\) for \(x \geq -1\) is [[2]].]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200812184807+4Bfzer</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200910132643+Oclj8u</text></tag>
    </tags>
    
</question>
</quiz>