<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27126 -->
  <question type="stack">
    <name>
				<text>36.2.3.1 (for Edinburgh) Skewness, gaps and multiple peaks</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Skewness, gaps and multiple peaks</h4>
<br>
<p>When exploring a data set, four properties worth looking for are outliers, skewness, gaps and multiple
peaks. Outliers have been dealt with in some detail above so the comments below briefly address skewness, gaps and multiple peaks.</p>
<br>
<p><b>Skewness</b></p>
<p>If a skewed distribution is represented purely by two numbers, say the mean and standard deviation,
then the representation will be inadequate. Remember that the term ‘skewness’ refers to the location
of the ‘tail’ of a distribution.</p><p style="text-align: center;"><img src="@@PLUGINFILE@@/36.2.3%20Skewness.svg" alt="In the picture two distributions characterised by right skew and left skew respectively" class="img-responsive atto_image_button_text-bottom" width="345" height="69"></p><p style="text-align: left;">

As an example, the data set below gives the current required to burn out a component under test.<br></p>
<table>
<caption></caption>
<thead>
<tr>
<th style="text-align: right;" scope="col"><span style="font-weight: normal;">\(9.5\) &nbsp;&nbsp; &nbsp; &nbsp; </span><br></th>
<th style="text-align: right;" scope="col"><span style="font-weight: normal;">\(11.9\)&nbsp; &nbsp;&nbsp; </span><br></th>
<th style="text-align: right;" scope="col"><span style="font-weight: normal;">\(20.0\) &nbsp;&nbsp; </span><br></th>
<th style="text-align: right;" scope="col"><span style="font-weight: normal;">\(33.4\) &nbsp;&nbsp; </span><br></th>
<th style="text-align: right;" scope="col"><span style="font-weight: normal;">\(40.1\) &nbsp;&nbsp; </span><br></th>
<th style="text-align: right;" scope="col"><span style="font-weight: normal;">\(50.0\) &nbsp;&nbsp; </span><br></th>
<th style="text-align: right;" scope="col"><span style="font-weight: normal;">\(12.7\) &nbsp;&nbsp; </span><br></th>
<th style="text-align: right;" scope="col"><span style="font-weight: normal;">\(21.0\) &nbsp;&nbsp; </span><br></th>
<th style="text-align: right;" scope="col"><span style="font-weight: normal;">\(33.6\) &nbsp;&nbsp; </span><br></th>
<th style="text-align: right;" scope="col"><span style="font-weight: normal;">\(40.6\) &nbsp;&nbsp; </span><br></th>
</tr>
</thead>
<tbody>
<tr>
<td style="text-align: right;">\(50.0\) &nbsp; &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(15.5\) &nbsp;&nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(26.4\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(35.4\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(41.1\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(50.0\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(17.7\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(37.9\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(41.3\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(50.0\) &nbsp;&nbsp; <br></td>
</tr>
<tr>
<td style="text-align: right;">\(41.9\) &nbsp; &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(50.4\)&nbsp; &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(43.0\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(43.3\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(43.6\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(43.7\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(43.8\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(44.7\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(44.9\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(45.0\) &nbsp;&nbsp; <br></td>
</tr>
<tr>
<td style="text-align: right;">\(45.1\)&nbsp; &nbsp; &nbsp; <br></td>
<td style="text-align: right;">\(45.2\) &nbsp;&nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(45.3\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(45.5\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(46.1\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(46.5\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(46.6\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(47.1\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(48.0\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(48.2\) &nbsp;&nbsp; <br></td>
</tr>
<tr>
<td style="text-align: right;">\(48.5\)&nbsp; &nbsp; &nbsp; <br></td>
<td style="text-align: right;">\(48.4\) &nbsp;&nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(48.6\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(48.7\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(48.8\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(48.9\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(49.4\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(49.5\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(49.6\) &nbsp;&nbsp; <br></td>
<td style="text-align: right;">\(49.8\) &nbsp;&nbsp; <br></td>
</tr>
</tbody>
</table>
<br>
<p>The data were obtained by measuring the current in mA applied to an electronic component under conditions of destructive testing give the following values for the mean, standard deviation, median and IQR:
\[ \bar{x} = 40.72, s_{n-1} = 11.49; \text{median} = 45.05; \text{and} \text{ IQR} = 8.5.\]</p>
<p>The values of \(\bar{x}\) and \(s_{n-1}\) indicate that a lower average current with a greater spread will result in the destruction of the component than that indicated by the median and IQR. Clearly, further
investigations are necessary to resolve this situation.</p>
<br>
<p><b>Gaps and multiple peaks</b></p>
<p>Distributions with gaps and multiple peaks can be very difficult to summarise easily. The table below summarise some 1972 data concerning adult literacy.</p>
<table class="HELM_table">
<thead>
    <tr>
      <th>Interval</th>
      <th>Frequency</th>
    </tr>
</thead>
<tbody>
<tr><td>\([0,10]\)</td><td>\(23\)</td></tr>
<tr><td>\((10,20]\)</td><td>\(11\)</td></tr>
<tr><td>\((20,30]\)</td><td>\(8\)</td></tr>
<tr><td>\((30,40]\)</td><td>\(9\)</td></tr>
<tr><td>\((40,50]\)</td><td>\(11\)</td></tr>
<tr><td>\((50,60]\)</td><td>\(3\)</td></tr>
<tr><td>\((60,70]\)</td><td>\(10\)</td></tr>
<tr><td>\((70,80]\)</td><td>\(14\)</td></tr>
<tr><td>\((80,90]\)</td><td>\(11\)</td></tr>
<tr><td>\((90,100]\)</td><td>\(28\)</td></tr>
</tbody>
</table><br>
<p>The virtual lack of data between \(50\) and \(60\) indicates a gap and suggest that we are in fact dealing with two separate features. The usual summary statistics that you might be tempted to calculate like mean, standard deviation, median and IQR are not of much use since neither set indicates the gap or skewness. Without visual representation, such as a histogram, a single-peaked distribution tends to be assumed which is not true in this case.</p>
<p><b>Final comments on data representations</b></p>
<ol>
<li>You should not rely on summary statistics such as the mean and standard deviation or median and IQR alone to represent a data set. Remember that if a distribution has outliers, gaps, skewness or multiple peaks, then the shape is probably more important than location and spread.</li>
<li>The shape of a distribution is better shown visually than numerically. Remember that a histogram and box-and-whisker plots emphasises the detail contained in the tails of a distribution.</li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200915102258+ghevrJ</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200915102258+jPdR9V</text></tag>
    </tags>
    
</question>
</quiz>