<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 38396 -->
  <question type="stack">
    <name>
				<text>1.4.1.7 Cancelling Common Factors</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_example">Example</h4>
<p>Show that the algebraic fraction \(\dfrac{3}{x+1}\) and \(\dfrac{3(x+4)}{x^2+5x+4}\) are equivalent.</p>

<h4 class="HELM_solution">Solution</h4>
<p>The denominator \(x^2+5x+4\) can be factorised as \((x+4)(x+1)\) so that</p>
<p>\(\qquad \dfrac{3(x+4)}{x^2+5x+4} = \dfrac{3(x+4)}{(x+4)(x+1)}\)</p>
<p>Note that \((x + 4)\) is a factor common to both the numerator and the denominator and can be cancelled to leave \(\dfrac{3}{x+1}\). Thus \(\dfrac{3}{x+1}\) and \(\dfrac{3(x+4)}{x^2+5x+4}\) are equivalent fractions</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210914131021+8N34XI</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220117172415+Jt0lB9</text></tag>
    </tags>
    
</question>
</quiz>