<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26756 -->
  <question type="stack">
    <name>
				<text>10.2.1.4 Argand Diagram(Algebraic Operations)</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Given that <span class="math inline">\(z={@z@}\)</span>, <span class="math inline">\(w={@w@}\)</span> , which of the points below represent \({@s@}z+{@t@}w\) on an Argand diagram.</p>
<figure class="HELM_figure"><img src="@@PLUGINFILE@@/Complex%20Numbers%20Argand%20plot.svg" alt="Argand diagram with points A, B, C, D" width="300" height="238" class="img-responsive atto_image_button_text-bottom"><br></figure>
<p><br></p>
<p>\({@s@}z+{@t@}w=\)[[input:ans1]] [[validation:ans1]]&nbsp;[[feedback:prt1]]<br></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>We are given that \(z={@z@}\) and&nbsp;\(w={@w@}\). Then,&nbsp;&nbsp;\({@s@}z+{@t@}w={@s@}({@z@})+{@t@}({@w@}).\) This is the point \({@cor1@}\) on the Argand diagram.</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>a:rand_with_prohib(-6,6,[-2,-1,0,1,2]);
b:rand([a-1,a,a+1,-a-1,-a,-a+1]);
absa:abs(a);
absb:abs(b);
z:a+%i*b;
c:rand_with_prohib(-7,7,[-2,-1,0,1,2]);
d:rand([c-1,c,c+1,-c-1,-c,-c+1]);
absc:abs(c);
absd:abs(d);
w:c+%i*d;
s:rand(4)+1;
t:rand(4)+4;
exp1: ev(s*z+t*w,simp);
absreal:abs(s*a+t*c);
abscomp:abs(s*b+t*d);
pts:[ [[absreal,abscomp],"A"],[[absreal,-abscomp],"B"],[[-absreal,-abscomp],"C"],[[-absreal,abscomp],"D"] ];

ta:map(lambda([ex],[ex[1],is(ex[1]=[s*a+t*c,s*b+t*d]),ex[2]]),pts);
cor1:[s*a+t*c,s*b+t*d];
letter1:maplist(third, sublist(ta, lambda([ex], second(ex))))[1];
</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(z={@z@}\), \(w={@w@}\),\({@s@}z+{@t@}w\), So, \({@s@}z+{@t@}w={@cor1@}\).</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>cor1</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200709101035+3JrsCC</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210701151856+syZO26</text></tag>
    </tags>
    <deployedseed>968378954</deployedseed>
    <deployedseed>14024368</deployedseed>
    <deployedseed>1898290021</deployedseed>
    <deployedseed>1737355458</deployedseed>
    <deployedseed>390064873</deployedseed>
    <deployedseed>1523810932</deployedseed>
    <deployedseed>1504689373</deployedseed>
    <deployedseed>613332400</deployedseed>
    <deployedseed>1158218488</deployedseed>
    <deployedseed>415584686</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev([s*a+t*c,s*b+t*d],simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev([-(s*a+t*c),s*b+t*d],simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev([s*a+t*c,-(s*b+t*d)],simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev([-(s*a+t*c),-(s*b+t*d)],simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>