<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 38394 -->
  <question type="stack">
    <name>
				<text>1.4.1.5 Cancelling Common Factors</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_example">Example</h4>
<p>Simplify \(\dfrac{3x+6}{6x+12}\)</p>

<h4 class="HELM_solution">Solution</h4>
<p>First we factorise the numerator and the denominator to see if there are any common factors.</p>
<p>\(\qquad \dfrac{3x+6}{6x+12} = \dfrac{3(x+2)}{6(x+2)} = \dfrac{3}{6} = \dfrac{1}{2}\)</p>
<p>The factors \(x + 2\) and 3 have been cancelled.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210914130050+6F0APe</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220117172350+38ppJb</text></tag>
    </tags>
    
</question>
</quiz>