<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27310 -->
  <question type="stack">
    <name>
				<text>37.1.1.6 Example permutation take two at a time</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<hr>
<h4 class="HELM_example">Example</h4>
<p>Find the number of permutations of the four letters \(A, B, C\) and \(D\) taken two at a time.</p>
<h4 class="HELM_solution">Solution</h4>
<p>We may choose the first letter in 4 ways and the second letter in 3 ways giving us 
\(\qquad
4\times 3=\dfrac{4\times 3\times 2\times 1}{1\times 2}=\dfrac{4!}{2!}=12\quad\) permutations]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200714145557+AQWCyi</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200715155756+4nxN3L</text></tag>
    </tags>
    
</question>
</quiz>