<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 33791 -->
  <question type="stack">
    <name>
				<text>AdrianTest</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<figure class="HELM_figure">
[[jsxgraph width='250px' height='250px']]
JXG.Options.axis.ticks.insertTicks = false;
JXG.Options.axis.ticks.minorTicks = 0;
JXG.Options.axis.ticks.majorHeight = 10;
JXG.Options.axis.ticks.ticksDistance = 1;
JXG.Options.axis.ticks.drawLabels = false;
JXG.Options.text.fontSize = 16; 
const board = JXG.JSXGraph.initBoard(divid, {boundingbox: [-8, 8, 8, -8], axis: true,
    showCopyright: false, showNavigation:false});
const s = board.create('slider', [[1, -5], [7, -5], [-10, 1, 10]]);
const F1 = board.create('point', [0, Math.sqrt(5)], {name: 'F_1'});
const F2 = board.create('point', [0, -Math.sqrt(5)], {name: 'F_2'});
const el = board.create('ellipse', [F2, F1, function(){ return s.Value(); }]);[[/jsxgraph]]
</figure>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text></text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text></text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210622150842+KObw6p</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210622155434+Uqvgng</text></tag>
    </tags>
    <deployedseed>2127394329</deployedseed>
    <deployedseed>1274090329</deployedseed>
    <deployedseed>2023039359</deployedseed>
    <deployedseed>455744874</deployedseed>
    <deployedseed>350754509</deployedseed>
    <deployedseed>1229334746</deployedseed>
    <deployedseed>1119075273</deployedseed>
    <deployedseed>127964942</deployedseed>
    <deployedseed>633357327</deployedseed>
    <deployedseed>1202052179</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
    </qtest>
    <qtest>
      <testcase>2</testcase>
    </qtest>
</question>
</quiz>