<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62180 -->
  <question type="stack">
    <name>
				<text>43.2 F distribution</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p><h3>Upper \(5\%\) points of the \(F\) distribution</h3><br></p><p><br></p><img src="@@PLUGINFILE@@/43_2_F_Table_fig.svg" alt="" width="236" height="136" role="presentation" class="img-responsive atto_image_button_text-bottom"><h3><br></h3><p><br></p><p><img src="@@PLUGINFILE@@/43_2_F_Table.svg" alt="" width="647" height="572" role="presentation" class="img-responsive atto_image_button_text-bottom"><br></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200712061019+EpL6x1</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200712061035+mCF8Bo</text></tag>
    </tags>
    
</question>
</quiz>