<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26329 -->
  <question type="stack">
    <name>
				<text>2.3.3.1 One-to-many not a function *****</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Explain why a one-to-many rule cannot be a function.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200617155253+Ty0oAC</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200622151525+zQ3bG3</text></tag>
    </tags>
    
</question>
</quiz>