<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62133 -->
  <question type="stack">
    <name>
				<text>23.4.2.3 FS expression 3</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
Obtain the Fourier series of \[ f(t) = t  \hspace{1cm} \pi \leq t \leq \pi \hspace{1cm} f(t + 2\pi) = f(t) \] and use it to find the answer to \( 1 - \dfrac{1}{3} + \dfrac{1}{5} - \dfrac{1}{7} + \dots\)<p>
</p><p></p>
\( 1 - \dfrac{1}{3} + \dfrac{1}{5} - \dfrac{1}{7} + \dots = \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]<p></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[As \( f(t) \) is an odd function, \( a_n = 0 \). For \( b_n\) we have
\[
\begin{align}
b_n &amp;= \dfrac{1}{\pi}  \displaystyle{ \int _{-\pi}^\pi }  t \sin nt \hspace{0.1cm} dt \\
&amp;= \dfrac{1}{\pi} \left( \left[ -\dfrac{t}{n} \cos nt \right]_{-\pi}^\pi + \dfrac{1}{n} \displaystyle{ \int _{-\pi}^\pi }  \cos nt \hspace{0.1cm} dt  \right) \\
&amp;= \dfrac{1}{\pi} \left( -\dfrac{\pi}{n} \cos n \pi -\dfrac{\pi}{n} \cos n \pi + \dfrac{1}{n^2} \left[ \sin n\pi \right] _{-\pi}^\pi \right) \\
&amp;= -\dfrac{2\cos n \pi}{n} 
\end{align}
\]
\[ =
\begin{cases}
-\dfrac{2}{n} \hspace{1cm} n = 2, 4, 6, \dots \\
\dfrac{2}{ni} \hspace{1cm} n = 1, 3, 5, \dots
\end{cases}
\]
which can be expressed as \( \dfrac{-2(-1)^n}{n} \).
<p></p>
So the Fourier series is given by
\[ f(t) =-2 \sum _{n=1}^\infty \dfrac{(-1)^n}{n} \sin nt  \]
Now setting \(t = \dfrac{\pi}{2} \) we have \(f(t) = \dfrac{\pi}{2}\). The Fourier series is
\[ -2 \sum _{n=1}^\infty \dfrac{(-1)^n}{n} \sin \dfrac{n\pi}{2} \]
which is equal to zero for even values of \(n\). Thus we have
which can be rearranged to give
\[ 
\begin{align}
\dfrac{\pi}{2} &amp;= -2 \sum _{n=1}^\infty \dfrac{(-1)^n}{n}\sin \dfrac{n\pi}{2} \\
&amp;= 2 \left( 1 - \dfrac{1}{3} + \dfrac{1}{5} - \dfrac{1}{7} + \dots \right) 
\end{align}
\]
Thus 
\[ 1 - \dfrac{1}{3} + \dfrac{1}{5} - \dfrac{1}{7} + \dots&nbsp; = \dfrac{\pi}{4} \]]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021120900</text>
		</stackversion>
    <questionvariables>
				<text>ta1: pi/4</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text></text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200905121818+Hwp8cL</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220518104324+X6NZMX</text></tag>
    </tags>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>