<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26596 -->
  <question type="stack">
    <name>
				<text>(not in quiz) 8.3.1.5 Remark about non-integers before Gaussian Elimination</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p> Although we have worked so far with integers this will not always be the case and fractions will enter the solution process. We must then take care and it is always wise to check that the equations balance using the calculated solution. </p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200626171435+REqpMa</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200813095925+GuRjaF</text></tag>
    </tags>
    
</question>
</quiz>