<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27339 -->
  <question type="stack">
    <name>
				<text>38.1.1.7 Example 2</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<br>
<h4 class="HELM_example">Example</h4>
<p>Traditional electric light bulbs are known to have a mean lifetime to failure of \(2000\) <b>hours</b>. It is also known that the distribution function \(p(t)\) of the time to failure
takes the form</p><p>&nbsp; &nbsp; &nbsp; &nbsp; \(p(t)=1-e^{-t/\mu}\)</p><p>where \(\mu\) is the mean time to failure. You will see if you study the topic of reliability
in more detail that this is a realistic distribution function. The reliability function \(r(t)\), giving the probability that the light bulb is still working at time \(t\), is defined
as<br></p><p>&nbsp; &nbsp; &nbsp; &nbsp;\(r(t)=1-p(t)=e^{-t/\mu}.\)</p><p>Find the proportion of light bulbs that you would expect to fail before \(1500\) <b>hours
</b>and the proportion you would expect to last longer than \(2500\) <b>hours</b>.<br></p>
<h4 class="HELM_solution">Solution</h4>
<p>Let \(T\) be the random variable ‘time to failure’.</p><p>The proportion of bulbs expected to fail before \(1500\) <b>hours </b>is given as<br></p><p>\(P(T&lt;1500)= 1 - e^{-1500/2000}=1-e^{1-3/4}=1-0.4724=0.5276.\)</p><p>The proportion of bulbs expected to last longer that \(2500\) <b>hours </b>is given as<br></p><p>\( \displaystyle{P(T&gt;2500) = 1 - P(T \leq 2500) = e^{-2500/2000}=e^{-5/4}=0.2865.}\)</p><p>Using \(r(t)=1-p(t)\) we have \(r(2500)=0.2865.\)</p><p>Hence we expect just under \(53 \%\) of light bulbs to fail before \(1500\) <b>hours </b>service and just under \(29\%\) of light bulbs to give over <b>2500 </b>hours service.<br></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200730113819+PEhHYF</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200804224632+h8YCxt</text></tag>
    </tags>
    
</question>
</quiz>