<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25727 -->
  <question type="stack">
    <name>
				<text>2.4.4.3 Sketch discontinuous periodic *****</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>

<p>Sketch a graph of a periodic function which has discontinuities.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200618173823+z5m6WH</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200622153125+XgTrJx</text></tag>
    </tags>
    
</question>
</quiz>