<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26720 -->
  <question type="stack">
    <name>
				<text>10.1.2.16 Dividing complex numbers</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Convert \[\frac{{@n1@} + {@n2@}\,{@%i@}}{{@n3@}\,{@%i@}}\]</p>
<p>into the form \(p+q\,{@%i@}\), where \(p\) and \(q\) are real. Input your answer, in the form <span style="font-family: 'courier new', courier, monospace;">p + q*i</span>, in the box below:</p>
<p>\(\displaystyle \frac{{@n1@} + {@n2@}\,{@%i@}}{{@n3@}\,{@%i@}} =\) [[input:ans1]] [[validation:ans1]]</p>
<p> [[feedback:Result]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>Multiplying the top and bottom by the conjugate of the denominator gives</p><p>\[ \frac{\ {@n1@} + {@n2@}\,{@%i@}\ }{\ {@n3@}\,{@%i@}\ } \times \frac{\ - {@n3@}\,{@%i@}\ }{\ - {@n3@}\,{@%i@}\ } = \frac{- {@n1*n3@}\,{@%i@} - {@n2*n3@}\,{@%i@}^2\ }{\&nbsp; - {@n3*n3@}\,{@%i@}^2\ },\]&nbsp;</p>
<p>and because \({@%i@}^2= -1\) we obtain</p>
<p>\[ \frac{\ {@n2*n3@} - {@n1*n3@}\,{@%i@}\ }{\ {@n3*n3@} \ } = {@correctanswer@}.\]&nbsp;</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>n1: ev(rand(8)+1,simp);
n2: ev(rand(7)+2, simp);
n3: ev(rand(6)+1,simp);

tempc:ev((n2*n3)+(-n1*n3)*%i, simp);
pval: ev(n2/n3, simp);
qval: ev(-n1/n3, simp);

correctanswer: ev(pval+qval*%i,simp);
correctanswer2: ev(pval,simp)+ev(qval,simp)*%i;
pvalwr:(n2*n3);
qvalwr:(-n1*n3);
wrong1:ev(pvalwr+qvalwr*%i,simp)/(n3*n3);
simp:true</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(\dfrac{{@n1@} + {@n2@}\,{@%i@}}{{@n3@}\,{@%i@}} = \displaystyle {@correctanswer@}\)</text>
		</questionnote>
    <questionsimplify>0</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correctanswer</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>Result</name>
      <value>1.0000000</value>
      <autosimplify>0</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text>sansr:ev(realpart(ans1),simp);
  sansi:ev(imagpart(ans1),simp);</text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>correctanswer</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>CMP-QUOT-TRUE</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>CMP-QUOT-FALSE-0</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>sansr</sans>
        <tans>pval</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>REAL_PART_CORRECT</trueanswernote>
        <truefeedback format="html">
					<text><p>The real part of your answer is correct.</p></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>CMP-QUOT-FALSE-2</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>sansi</sans>
        <tans>qval</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>IM_PART_CORRECT</trueanswernote>
        <truefeedback format="html">
					<text><p>The imaginary part of your answer is correct.</p></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>CMP-QUOT-FALSE-3</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>3</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>NOT_SIMPLIFIED_ENOUGH</trueanswernote>
        <truefeedback format="html">
					<text><p>Your answer is equivalent to the correct answer, but not in the required form \(p + q\,i\).</p></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>CMP-QUOT-FALSE-1</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>4</name>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>correctanswer2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>Result-5-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>Result-5-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200618152022+JWWWxl</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210706165753+KICf16</text></tag>
    </tags>
    <deployedseed>1130459301</deployedseed>
    <deployedseed>1088572197</deployedseed>
    <deployedseed>1655466154</deployedseed>
    <deployedseed>1560522889</deployedseed>
    <deployedseed>2049069101</deployedseed>
    <deployedseed>947482087</deployedseed>
    <deployedseed>1860206311</deployedseed>
    <deployedseed>796287237</deployedseed>
    <deployedseed>1560273457</deployedseed>
    <deployedseed>1143737572</deployedseed>
    <deployedseed>24827407</deployedseed>
    <deployedseed>1969373247</deployedseed>
    <deployedseed>1332916957</deployedseed>
    <deployedseed>783930167</deployedseed>
    <deployedseed>1370653549</deployedseed>
    <deployedseed>354825038</deployedseed>
    <deployedseed>1583047228</deployedseed>
    <deployedseed>1183999513</deployedseed>
    <deployedseed>2068823685</deployedseed>
    <deployedseed>1109246487</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>correctanswer</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>CMP-QUOT-TRUE</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>wrong1</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>NOT_SIMPLIFIED_ENOUGH</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(correctanswer + 1,simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>IM_PART_CORRECT</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(correctanswer + %i,simp)</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>REAL_PART_CORRECT</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <testinput>
        <name>ans1</name>
        <value>0</value>
      </testinput>
      <expected>
        <name>Result</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>CMP-QUOT-FALSE-3</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>