<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26600 -->
  <question type="stack">
    <name>
				<text>(not in quiz) 8.3.2.3 Remarks after describing Gauss Elimination</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>You should always check your solution in the original equations to ensure that the equations balance.</p>
<p>Again we emphasise that we chose a particular set of procedures in Stage 1. This was chosen mainly to keep the arithmetic simple by delaying the introduction of fractions. Sometimes we are courageous and take fewer, harder steps.</p>
<p>An important point to note is that when in Stage 2 we wrote <span class="math inline">\(R2-2R1\)</span> against row 2; what we meant is that row 2 is replaced by the combination (row <span class="math inline">\(2)-2\times\)</span>(row 1).</p>
<p>In general, the operation <span class="math display">\[\text{ row } i-\alpha \times \text{ row }\ j\]</span> means replace <span><strong>row <span class="math inline">\(i\)</span></strong></span> by the combination <span class="math display">\[\text{ row } i-\alpha \times \text{ row } j.\]</span></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200630111956+f2DT3f</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200813095951+FUByef</text></tag>
    </tags>
    
</question>
</quiz>