<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26242 -->
  <question type="stack">
    <name>
				<text>7.3.1.1 Determinant of a 2x2 matrix</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Determinant of a \(2\times2\) matrix</h4>
<p>The <b>determinant</b> of the matrix \(A=\begin{bmatrix} a &amp; b \\ c &amp; d \end{bmatrix}\) is denoted by \(\begin{vmatrix} a &amp; b \\ c &amp; d \end{vmatrix}\) (note the change from square brackets to vertical lines) and is defined to be the number \(ad-bc\). That is: \[\begin{vmatrix} a &amp; b \\ c &amp; d \end{vmatrix} = ad-bc.\] </p>

<p>We can use the notation \(\det(A)\) or \(\lvert A \rvert\) to denote the determinant of \(A\). A common symbol used to represent the determinant of a given matrix is \(\Delta.\)</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200709133222+dM0uBI</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200908091949+i9p9DD</text></tag>
    </tags>
    
</question>
</quiz>