<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24290 -->
  <question type="stack">
    <name>
				<text>5.1.1.5 Falling rocks - initial velocity</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>Consider now how the function for \(v\) will change if the villain is even mightier than we previously thought and throws the rock <strong>upwards</strong> with an initial speed of \(5\ \text{m\ s}^{-1}\) instead of simply dislodging it or throwing it
    downwards. In this circumstance, the initial velocity is directed upwards, and since position is being measured downwards, the initial velocity is negative. We can use the equation \(v = 9.8t + b \) again. This time \(v = -5\) when \(t = 0\), leading
    to \(b = -5\) and </p>
<p>\(\qquad v = 9.8t -5 \) \((0\le t\le T_2)\)</p>
<p>The new time at which the rock hits the ground is denoted by \(T_2\). The rock will rise before falling to the gound this time so \(T_2\) will be larger than \(T_1\).</p>
<p>From the modelling point of view, there is one other significant time before the rock hits the ground. Figure 3 shows the new graph of \(v\) against \(t\). Notice that there is a time at which \(v\) (which starts at \(-5)\) is zero. What does this mean?</p>

<p><img width="376" height="136" class="img-responsive atto_image_button_text-bottom" alt="velocity time graph y = 9.8t - 5" src="@@PLUGINFILE@@/5.1%20linear%20functions%20example.svg"></p>
<p><strong>Figure</strong>: Graph of \(v=9.8t-5\) for the falling rock</p>

<p>As time goes by, the fact that gravity is causing the rock to accelerate downwards means that the rock’s upward motion will slow. Its velocity will decrease in magnitude until it reaches zero. At this particular instant the rock will be at its highest
    point and its velocity will change from upwards (negative) to downwards (positive) passing instantaneously through zero in the process.</p>
<p>We can calculate this time by substituting zero for \(v\) and working out the corresponding \(t\).</p>
<p>\(\qquad 0 = 9.8t -5\), so \(\qquad t = \dfrac{5}{9.8} = 0.51\).</p>
<p>This means that the rock is stationary about a half second after being thrown upwards. Subsequently the rock will fall until it hits the ground. But there is yet one more time that may be significant in the modelling context chosen here. During its journey
    to the ground 35 m beneath the cliff-top, the rock will pass the top of the cliff again. Note that we are modelling the motion of a particular point, say the lowest point, on the rock. A real rock, with appreciable size, will only pass the top of
    the cliff, without landing on it or hitting it, if it is thrown a little forward as well as up. Anyway, in principle we could use the function that we started with, representing the velocity of an object falling from rest under gravity, to work out
    how long the rock will take to pass the top of the cliff having reached the highest point in its path. A simpler method is to argue that, as long as the rock is thrown from the cliff top level (this requires the villain to be lying down!), the rock
    should take exactly the same time (approximately 0.5 s) to return to the level of the cliff top as it took to rise above the cliff top to the highest point in its path. So we simply double 0.5 s to deduce that the rock passes the cliff top again about
    1 s after being thrown.</p>

<p><br></p>
<hr>
<h4 class="HELM_example">Example</h4>
<p>This Example concerns the falling rock model just discussed.</p>
<ol class="HELM_parts">
    <li>
        <p>Add lines to a sketch version of previous figure to represent velocity as a function of time if the rock is: </p>
        <p>(i) dislodged, (ii) thrown with velocity \(3 \, \textrm{ms}^{−1}\) downwards, (iii) thrown with velocity \(-2 \, \textrm{ms}^{−1}.\)</p>
    </li>

    <li>
        <p>What do you deduce about the effect of the initial velocity on the graph of velocity against time?</p>
    </li>

    <li>Imagine that the filming was on the Moon with roughly one-sixth the gravitational pull of Earth. Find a linear function that would describe the velocity of a dislodged rock.</li>

    <li>What do you deduce about the effect of changing the acceleration due to gravity on the graph of velocity against time?</li>
</ol>

<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
    <li><img src="@@PLUGINFILE@@/Block7_Fig4a.svg" alt="Graph of velocity vs time of the falling rock" width="369" height="151" class="img-responsive atto_image_button_text-bottom"></li>
    <li>The effect of changing the initial velocity (in size or in sign) is simply to displace the straight line upwards or downwards without changing its slope.</li>
    <li>In Moon the gravitational pull is roughly one-sixth of Earth's so \(v= \dfrac{9.8}{6}t \approx 1.6t\)</li>
    <li>The graph of velocity agains time is still linear but the change in the acceleration due to gravity changes the slope.</li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200724074315+OGiDNL</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220407205435+7wyaYF</text></tag>
    </tags>
    
</question>
</quiz>