<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27306 -->
  <question type="stack">
    <name>
				<text>37.1.1.2 Permutations of ABCD</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<hr>
<h4 class="HELM_example">Example</h4>
<p> Write out the possible permutations of the letters \(A, B, C\) and \(D\). </p>
<h4 class="HELM_solution">Solution</h4>
The possible permutations are 
\[\begin{matrix}ABCD&	ABDC&	ADBC&	ADCB	&ACBD&	ACDB\\
BADC&	BACD&	BCDA&	BCAD&	BDAC&	BDCA\\
CABD&	CADB&	CDBA&	CDAB&	CBAD&	CBDA\\
DABC&	DACB&	DCAB&	DCBA&	DBAC&	DBCA
\end{matrix}\]
There are \(4!=24\) permutations of the four letters \(A, B, C\) and \(D\).]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200714140853+zepFID</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200715155720+dfA5Y4</text></tag>
    </tags>
    
</question>
</quiz>