<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 84678 -->
  <question type="stack">
    <name>
				<text>18.3 Introduction (copy)</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>The calculation of the optimum value of a function of two variables is a common requirement in many areas of engineering, for example in thermodynamics. Unlike the case of a function of one variable we have to use more complicated criteria to distinguish
    between the various types of stationary point.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+220307124632+OLddEI</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220307124632+7yi4kd</text></tag>
    </tags>
    
</question>
</quiz>