<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 33708 -->
  <question type="stack">
    <name>
				<text>19.2.9.2 First Order DE Ex1</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Solve the equation \(x^2\dfrac{\mathrm{d}y}{\mathrm{d}x} + xy = {@aa@}\).</p>
<p>\(y =\) [[input:ans1]] [[validation:ans1]] \( + \dfrac{C}{x}\). [[feedback:prt1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>We are given the equation \(x^2\dfrac{\mathrm{d}y}{\mathrm{d}x} + xy = {@aa@}\).</p>
<p>Divide through by \(x^2\) to obtain the standard form of the equation \[ \dfrac{\mathrm{d}y}{\mathrm{d}x} + \dfrac{1}{x}y = \dfrac{{@aa@}}{x^2}.\] The coefficient in front of \(y\) is \(\dfrac{1}{x}\). So the integrating factor is \[e^{\int \frac{1}{x}\ \mathrm{d}x} = e^{\log(x)} = x.\]</p>

<p>Multiply the standard form of the equation by the integrating factor to get \[ x\dfrac{\mathrm{d}y}{\mathrm{d}x} + y = \dfrac{{@aa@}}{x}.\] Recognise that this is now exact and construct an integral \[\frac{\mathrm{d}}{\mathrm{d}x}(xy) = \dfrac{{@aa@}}{x} \iff xy = \int \dfrac{{@aa@}}{x}\ \mathrm{d}x.\]</p>

<p>Integrating gives \[xy = {@aa*log(x)@} + C\] which implies \[y = \dfrac{{@aa@}}{x}\log(x) + \dfrac{C}{x}\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>aa: 1+ rand(10);

ta1: (aa/x)*log(x);

</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(x^2\dfrac{\mathrm{d}y}{\mathrm{d}x} + xy = {@aa@}, y = {@ta1@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210617104907+LBXXNi</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210617114705+i7h1Mr</text></tag>
    </tags>
    <deployedseed>1637073094</deployedseed>
    <deployedseed>321574568</deployedseed>
    <deployedseed>257950607</deployedseed>
    <deployedseed>54555136</deployedseed>
    <deployedseed>2109550951</deployedseed>
    <deployedseed>1567919944</deployedseed>
    <deployedseed>1910671046</deployedseed>
    <deployedseed>1468801973</deployedseed>
    <deployedseed>1907207895</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(2*ta1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>