<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24681 -->
  <question type="stack">
    <name>
				<text>19.1.2.3 Arbitrary Constants in a Differential Equation</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>We have seen that an expression including one arbitrary constant required one differentiation to obtain a differential equation which eliminated the arbitrary constant. Where two constants were present, two differentiations were required. Is the converse true? For example, would a differential equation involving <span class="math inline">\(\dfrac{d^2y}{dx^2}\)</span> as the only derivative have a general solution with one arbitrary constant and would a differential equation which had <span class="math inline">\(\dfrac{d^2y}{dx^2}\)</span> as the highest derivative produce a general solution with two arbitrary constants? The answer is, usually, yes.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:moodle2.maths.ed.ac.uk+210305180307+eBC2Hg</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+210319100827+fKX12u</text></tag>
    </tags>
    
</question>
</quiz>