<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24626 -->
  <question type="stack">
    <name>
				<text>14.4.1.1 Exercise 1</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Find the length of the line \(y = {@y@}\) between \(x = {@a@}\) and \(x = {@b@}\) using the technique of this Section. Verify your result from your knowledge of the straight line. Give your answer in <b>exact</b> form.</p>
<p>The length is [[input:ans1]] [[validation:ans1]]. </p>
<p>[[feedback:prt1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>The distance is from \(({@a@},{@y_a@})\) to \(({@b@},{@y_b@})\) along the line. This is given using Pythagoras' theorem as \(\sqrt{({@b@}-{@a@})^2+({@y_b@}-{@y_a@})^2}={@ta1@} \)</p>
<p>Using the formula \[\text{length of curve}=\int_a^b \sqrt{1+\left(\frac{dy}{dx}\right)^2}\&nbsp; \mathrm{d}x,\] the length of the line is calculated to be \[\int_{{@a@}}^{{@b@}} \sqrt{1+({@dy@})^2} \ \mathrm{d}x={@ta1@}.\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>a: rand_with_step(0,3,1);
b: a+2;
m: rand_with_step(2,5,1);
c: rand_with_prohib(-5,5,[0]);
y: m*x + c;
y_a: m*a+c;
y_b: m*b+c;
dy: diff(y, x)
ta1: integrate(sqrt(1+(dy)^2), x, a, b);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\( y={@y@}, \, x=[{@a@}, {@b@}], \, l={@ta1@} \)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.</span>]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>units</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201202062607+4qbL87</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210608110214+TCmeWN</text></tag>
    </tags>
    <deployedseed>1353566626</deployedseed>
    <deployedseed>922363332</deployedseed>
    <deployedseed>733795246</deployedseed>
    <deployedseed>2053700472</deployedseed>
    <deployedseed>136601151</deployedseed>
    <deployedseed>337263064</deployedseed>
    <deployedseed>813504560</deployedseed>
    <deployedseed>1335482690</deployedseed>
    <deployedseed>112380295</deployedseed>
    <deployedseed>1586231821</deployedseed>
    <deployedseed>1237756843</deployedseed>
    <deployedseed>1036858917</deployedseed>
    <deployedseed>1346648312</deployedseed>
    <deployedseed>1269866539</deployedseed>
    <deployedseed>1357371083</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>