<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27191 -->
  <question type="stack">
    <name>
				<text>35.1.2.2 Venn diagram example</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<hr>
<h4 class="HELM_example">Example</h4>
<p>Let \(A = \{0,1\}\) and let \(B = \{2,3,4\}\).</p>
<p>Note that \(A\) and \(B\) have no elements in common. This is represented pictorially in the Venn diagram by circles which are totally separate from each other as shown in the figure below. Here, \(S\) denotes the universal set.</p>
<img src="@@PLUGINFILE@@/35-1-2-Venndiagram1.svg" alt="" width="189" height="115" class="img-responsive atto_image_button_text-bottom">
<p>Figure: The sets \(A\) and \(B\) have no elements in common.</p><br>
<p>The complement of \(A\), denoted by \(A'\), is the set whose elements do not belong to \(A\). The set \(A'\) is shown in the figure below (the shaded area).</p>
<img src="@@PLUGINFILE@@/35-1-2-Venndiagram2.svg" alt="" width="188" height="115" class="img-responsive atto_image_button_text-bottom">
<p>Figure: The complement of \(A\) contains elements which are not in \(A\).</p><br>
<p>Similarly, the set \(B'\) is shown in the figure below (the shaded area).</p>
<img src="@@PLUGINFILE@@/35-1-2-Venndiagram3.svg" alt="" width="188" height="115" class="img-responsive atto_image_button_text-bottom">
<p>Figure: Everything outside of \(B\) (but within \(S\)) is \(B'\).</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200703213449+Kr4fOR</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200727224603+O6T1oZ</text></tag>
    </tags>
    
</question>
</quiz>