<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27343 -->
  <question type="stack">
    <name>
				<text>Introduction</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>This Section introduces the simplest type of continuous probability distribution which features acontinuous random variableXwith probability density functionf(x)which assumes a constantvalue over a finite interval.<br></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200708144210+a6CjHH</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200708144210+GVhveJ</text></tag>
    </tags>
    
</question>
</quiz>