<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 37797 -->
  <question type="stack">
    <name>
				<text>1.1.6.2 Combining Numbers Together Using +,−,×,÷</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_example">Example</h4>
<p>Write out explicitly </p>
<ol class="HELM_parts_inline">
<li>\(\sum_{k=1}^{3}1\)</li>
<li>\(\sum_{k=0}^{4}2\)</li>
</ol>
<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
<li>Here k does not appear explicitly in the terms to be added. This means add the constant 1, three times.
<p>\(\qquad \sum_{k=1}^{3}1 = 1 + 1 + 1 = 3\)<p>
<p>In general \(\sum_{k=1}^{n}1 = n\)
</li>
<li> Here \(k\) starts at zero so there are \(n + 1\) terms where \(n = 4\):
<p>\(\qquad \sum_{k=0}^{4}2 = 2 + 2 + 2 + 2 + 2 = 10\)<p>
</li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210906111329+WDmBMX</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210906112142+eOhkth</text></tag>
    </tags>
    
</question>
</quiz>