<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25057 -->
  <question type="stack">
    <name>
				<text>Improper integrals</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 style="text-align: left;">Improper integrals</h4>
<p dir="ltr" style="text-align: left;">So far most of the integrals we have looked at have been of continuous functions \(f(x)\) defined on a closed and bounded interval \([a,b]\). The definite integral is then \[\int_a^b f(x) \mathrm{d}x = F(b)-F(a)\] where \(F(x)\) is any function with \(F'(x)=f(x)\).<br><br>An
    <em>improper integral</em> is where either the function is unbounded, e.g. \(\displaystyle{ \int_0^1 \dfrac1x \mathrm{d}x}\) or the domain of integration is unbounded, e.g.&nbsp;\(\displaystyle{ \int_1^\infty \dfrac1x \mathrm{d}x}\).</p>
<p dir="ltr" style="text-align: left;">In either case the usual approach will not make sense, make sure you see why,</p>
<h5 style="text-align: left;">Solving improper integrals</h5>
<p><strong>Step 1:</strong> Identify where the integrand is unbounded or where the limit(s) of integration is unbounded. <br><br>Note that there may be several points such as in&nbsp;\(\displaystyle{ \int_1^\infty \dfrac{2x}{(x^2-4)} \mathrm{d}x}\), in this
    example the integrand is unbounded with a vertical asymptote at \(2\), and one of the limits of integration is infinity. Note that the though the integrand has another asymptote at \(-2\) this is not a problem as it falls outside the limits of integration.<br><br><strong>Step 2:</strong>    Split the integral into a sum of several pieces so that each piece has at most one improper point.<br><br>Using our example we could have</p>
<p dir="ltr" style="text-align: left;">\[\int_1^\infty \dfrac{2x}{(x^2-4)} \mathrm{d}x = \int_1^2 \dfrac{2x}{(x^2-4)} \mathrm{d}x + \int_2^3 \dfrac{2x}{(x^2-4)} \mathrm{d}x+\int_3^\infty \dfrac{1}{(x^2-4)} \mathrm{d}x.\]</p>
<p dir="ltr" style="text-align: left;">The choice of \(3\) here was arbitrary, any value greater than \(2\) would be appropriate.</p>
<p dir="ltr" style="text-align: left;"><strong>Step 3:</strong>&nbsp;Consider each improper integral (with a single unbounded point) separately. There are two types.<br><br>If the integrand is unbounded, such as in \(\displaystyle{\int_1^2 \dfrac{2x}{(x^2-4)} \mathrm{d}x}\), we integrate over
    a smaller region and take a limit, i.e.</p>
<p dir="ltr" style="text-align: left;">\[\displaystyle{\int_1^2 \dfrac{2x}{(x^2-4)} \mathrm{d}x}= \lim_{t \rightarrow 2-}&nbsp;\displaystyle{\int_1^t \dfrac{2x}{(x^2-4)} \mathrm{d}x}<br>.\]<br>In this example we are the taking the limit as \(t\) approaches \(2\) from below. On this region
    the integral is well defined so we have<br>\[\lim_{t \rightarrow 2-}&nbsp;\displaystyle{\int_1^t \dfrac{2x}{(x^2-4)} \mathrm{d}x} =&nbsp;\lim_{t \rightarrow 2-} \left( \ln|t^2-4|-\ln|-4|\right)\]<br>which diverges, in this example.<br><br>
</p>
<p dir="ltr">If the domain of integration is unbounded, such as in \(\displaystyle{\int_3^\infty \dfrac{2x}{(x^2-4)} \mathrm{d}x}\), we integrate over a smaller region and take a limit, i.e.</p>
<p dir="ltr">\[\displaystyle{\int_3^\infty \dfrac{2x}{(x^2-4)} \mathrm{d}x}= \lim_{t \rightarrow \infty}&nbsp;\displaystyle{\int_3^t \dfrac{2x}{(x^2-4)} \mathrm{d}x}<br>.\]<br>On this region the integral is well defined so we have<br>\[\lim_{t \rightarrow \infty}&nbsp;\displaystyle{\int_3^t
    \dfrac{2x}{(x^2-4)} \mathrm{d}x} =&nbsp;\lim_{t \rightarrow 2-} \left( \ln|t^2-4|-\ln|9-4|\right)\]<br>which diverges again, in this example.</p><strong>Step 4: </strong>If any of the of the integrals with one unbounded point diverge then we say the
original integral&nbsp; diverges. Otherwise the integral converges only if all of the integrals with one unbounded point converges.<br><br>

<hr>
<h4 class="HELM_example">Example</h4>
<br>For example, the integral \( \displaystyle{ \int_1^\infty \frac{1}{x^2} \mathrm{d}x\) is improper as the upper limit is infinity. This is the only unbounded point as \(0\) is not within the region of integration. To determine if this integral converges
we consider the limit<br><br>\[ \lim_{t \rightarrow \infty} \int_1^t \frac{1}{x^2} \mathrm{d}x = \lim_{t \rightarrow \infty} \left(\frac{1}{-t}-&nbsp;\frac{1}{-1}\right)=0+\frac11=1.\]<br>Since this converges we say the integral converges, and the value of the integral
is \(1\), in this case.<br>
<p></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>1</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:moodle2.maths.ed.ac.uk+210211005728+92ccZB</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+210216120624+u5DBw3</text></tag>
    </tags>
    
</question>
</quiz>