<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26327 -->
  <question type="stack">
    <name>
				<text>2.3.2.1 Inverse of a function</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>Inverse of a function </h3>
<p>We have seen that a function can be regarded as taking an input, <span class="math inline">\(x\)</span>, and processing it in some way to produce a single output <span class="math inline">\(f(x)\)</span> as shown in part (a) of the figure below.</p>
<figure class="HELM_figure">
<img src="@@PLUGINFILE@@/block4-fig7-1.svg" alt="Graphical representation of a function and its inverse" width="447" height="183" class="img-responsive atto_image_button_text-bottom">
<figcaption>The second block reverses the process in the first</figcaption>
</figure>
<p>A natural question to ask is whether we can find another function that will reverse the process. In other words, can we find a function that will start with <span class="math inline">\(f(x)\)</span> and process it to produce <span class="math inline">\(x\)</span> again? This idea is also shown in part (b) of the figure. If we can find such a function it is called the <b>inverse function</b> to <span class="math inline">\(f(x)\)</span> and is given the symbol <span class="math inline">\(f^{-1}(x)\)</span>. <br></p><p><br></p><p>Do not confuse the ‘<span class="math inline">\(-1\)</span>’ with an index, or power. Here the superscript is used purely as the notation for the inverse function. Note that the composite function <span class="math inline">\(f^{-1}(f(x))=x\)</span> as shown in figure below.</p>
<figure class="HELM_figure">
<img src="@@PLUGINFILE@@/block4-fig8-1.svg" alt="Graphical representation of a function and its inverse composed" width="495" height="85" class="img-responsive atto_image_button_text-bottom">
<figcaption>\(f^{-1}\) reverses the process in \(f\)</figcaption>
</figure>

<hr>
<h4 class="HELM_example">Example</h4><p>Find the inverse function to \(f(x)=3x-8\).</p>
<h4 class="HELM_solution">Solution</h4>
<p>The given function takes an input, <span class="math inline">\(x\)</span> and produces an output <span class="math inline">\(3x-8\)</span>. The inverse function, <span class="math inline">\(f^{-1}\)</span>, must take an input <span class="math inline">\(3x-8\)</span> and give an output <span class="math inline">\(x\)</span>. That is</p><p><span class="math inline">\(\qquad f^{-1}(3x-8) = x \, .\)</span></p><p>If we introduce a new variable <span class="math inline">\(z=3x-8\)</span>, and transpose this for <span class="math inline">\(x\)</span> to give</p><p><span class="math inline">\(\qquad x=\dfrac{z+8}{3}
\qquad
\text{then}\qquad f^{-1}(z)=\dfrac{z+8}{3} \, .\)</span></p><p>So the rule for <span class="math inline">\(f^{-1}\)</span> is add 8 to the input and divide the result by 3. Writing <span class="math inline">\(f^{-1}\)</span> with <span class="math inline">\(x\)</span> as its argument gives</p><p><span class="math inline">\(\qquad f^{-1}(x)= \dfrac{x+8}{3} \, .\)</span></p><p>This is the inverse function.</p>

<hr>
<p>Not all functions possess an inverse function. In fact, only one-to-one functions do so. If a function is many-to-one the process to reverse it would require many outputs from one input contradicting the definition of a function.</p><p><br></p><p><i>Try for yourself</i>:&nbsp;Explain why a many-to-one function does not have an inverse function. Give an example.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200617145152+4j2OX3</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210628123752+EsktMy</text></tag>
    </tags>
    
</question>
</quiz>