<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26578 -->
  <question type="stack">
    <name>
				<text>8.2.1.3 Example 1 inverse matrix method</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_example">Example</h4>
<p>Solve the system <span class="math inline">\(AX=B\)</span> where <span class="math display">\[A=\left[\begin{array}{rr}2&amp;3\\
1&amp;-2\end{array}\right], \hspace{2mm}
B = \left[\begin{array}{r}5\\ -1\end{array}\right].\]</span></p>

<h4 class="HELM_solution">Solution</h4>

<p> In the previous example, we checked that <span class="math inline">\(A\)</span> has an inverse and computed it. Therefore,
    <span class="math display">\[X=A^{-1}B=\dfrac{1}{7}\left[\begin{array}{rr}2&amp;3\\
1&amp;-2\end{array}\right]\left[\begin{array}{r} 5\\
-1\end{array}\right]=\dfrac{1}{7}\left[\begin{array}{c}7\\
7\end{array}\right]=\left[\begin{array}{c}1\\ 1\end{array}\right].\]</span> </p>
<p> Hence <span class="math inline">\(x_1=1,\  x_2=1.\)</span></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200617113201+Zx81IM</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210630130958+4az6zC</text></tag>
    </tags>
    
</question>
</quiz>