<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24685 -->
  <question type="stack">
    <name>
				<text>19.1.3.4</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p dir="ltr" style="text-align: left;">Classify the differential equation \[\dfrac{d^3y}{dx^3}=\left(\dfrac{dy}{dx}\right)^2+\dfrac{dy}{dx}.\]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p dir="ltr" style="text-align: left;">The highest derivative is of order three, and the squared term is non-linear.</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:moodle2.maths.ed.ac.uk+210315125940+4w2vKg</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+210315131854+q2h2ys</text></tag>
    </tags>
    
</question>
</quiz>