<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27304 -->
  <question type="stack">
    <name>
				<text>36.1.6.4 Exercise 3 - Average over 2 samples</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>A lecturer gives a science test to two classes and calculates the results as follows:</p>
<ul><li>Class \(A\) - average mark \({@x@}\%\)</li>
<li>Class \(B\) - average mark \({@y@}\%\)</li></ul>
<ol class="HELM_parts">
<li><p>Suppose Class \(A\) has \({@a@}\) students and Class \(B\) has \({@a@}\) students. 
</p><p>The average mark over the two classes = [[input:ans1]] [[validation:ans1]] \(\%\) [[feedback:prt1]]</p>
<p>Syntax hint: If the answer is, say \(30 \%\), then enter \(30\).</p></li>
<li><p>Suppose Class \(A\) has \({@a@}\) students and Class \(B\) has \({@b@}\) students. 
</p><p>The average mark over the two classes = [[input:ans2]] [[validation:ans2]] \(\%\) [[feedback:prt2]]</p>
<p>Syntax hint: Enter 2 decimal places. If the answer has less than 2 decimal places, then enter zero(s) at the end.</p></li>
<li><p>Suppose Class \(A\) and Class \(B\) have the same number of students. 
</p><p>The average mark over the two classes = [[input:ans3]] [[validation:ans3]] \(\%\) [[feedback:prt3]]</p></li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>(a) The sum of the marks of all students in Class \(A\) is \({@a@} \times {@x@} \%\). Similarly, the sum of the marks of all students in Class \(B\) is \({@a@} \times {@y@} \%\). Hence, the average mark over the two classes is
\[ \dfrac{{@a@} \times {@x@} \% + {@a@} \times {@y@} \%}{{@a@}+{@a@}} = {@ta1@} \%.\]</p>
<p>(b) The sum of the marks of all students in Class \(A\) is \({@a@} \times {@x@} \%\). Similarly, the sum of the marks of all students in Class \(B\) is \({@b@} \times {@y@} \%\). Hence, the average mark over the two classes is
\[ \dfrac{{@a@} \times {@x@} \% + {@b@} \times {@y@} \%}{{@a@}+{@b@}} = {@ta2@} \%.\]</p>
<p>(c) Suppose Class \(A\) has \(n\) students and so does Class \(B\). The sum of the marks of all students in Class \(A\) is \(n \times {@x@} \%\). Similarly, the sum of the marks of all students in Class \(B\) is \(n \times {@y@} \%\). Hence, the average mark over the two classes is
\[ \dfrac{n \times {@x@} \% + n \times {@y@} \%}{n+n} = \dfrac{n \times ({@x@} \%+{@y@} \%)}{2n} =\dfrac{{@x@} \%+{@y@} \%}{2} = {@ta1@} \%.\]</p>
<p>Note that the answers for (a) and (c) are the same. That is, when Class \(A\) and Class \(B\) have the same number of students, the answer is independent of the class sizes, and is given by \[ \dfrac{{@x@} \%+{@y@} \%}{2} = {@ta1@} \%.\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020070100</text>
		</stackversion>
    <questionvariables>
				<text>a:rand([20,25,30]);
b:rand([50,55,60]);
x:rand([34,36,38]);
y:rand([40,42,44]);
av:(a*x+b*y)/(a+b);
ta1:(x+y)/2;
ta2:decimalplaces(float(av),2);


</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\[A = {@x@}\%, \quad B ={@y@}\% \]
\[ \frac{x+y}{2} = {@ta1@}, \quad \dfrac{{@a@} \times {@x@} + {@b@} \times {@y@}}{{@a@}+{@b@}} = {@ta2@}\]</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>dispdp(ta2,2)</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>dispdp(ta2,2)</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans3</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200914150622+jJOvir</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200914150623+NuEXpH</text></tag>
    </tags>
    <deployedseed>273223100</deployedseed>
    <deployedseed>2023152754</deployedseed>
    <deployedseed>271480207</deployedseed>
    <deployedseed>484160731</deployedseed>
    <deployedseed>2004032693</deployedseed>
    <deployedseed>1224351379</deployedseed>
    <deployedseed>338769291</deployedseed>
    <deployedseed>165171661</deployedseed>
    <deployedseed>2134611134</deployedseed>
    <deployedseed>642217134</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta2+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dispdp(ta2,2)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>