<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27204 -->
  <question type="stack">
    <name>
				<text>35.2.2.1 Definitions of Probability</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>Definitions of Probability </h3>
<h4> Relative frequency applied to probability </h4>
<p> Consider the experiment of throwing a single coin many times. </p>
<p> Suppose we throw a coin 10 times and obtain six Heads and four Tails; does this suggest that the coin is biased? Clearly not! What about the case when we obtain 9 Heads and 1 Tail? We conducted an experiment in which a coin was thrown 100 times and the result recorded each
time as 1 if a Head appeared face up and 0 if a Tail appeared. In the figure below we have plotted the
average score \(r/n\), where \(r\) is the number of Heads and \(n\) is the total number of throws, against \(n\) for \(n=10,20,\dots,100\). The quantity \(r/n\) is called the <strong>relative frequency</strong> of Heads.</p><p><strong><img src="@@PLUGINFILE@@/3-5-2-relativefrequency.svg" alt="" role="presentation"><br></strong></p><p><strong>Figure:</strong> Plot of the relative frequency, \(r/n\).</p>
<br>
<p> As \(n\) increases the relative frequency settles down near the value \(1/2\). This is an experimental estimate of the probability of throwing a Head with this particular coin. Note that when \(n = 50\) this estimate was \(0.49\) and when \(n = 100\) this estimate was \(0.51\). When we repeated the whole experiment again, the value of \(r/n\) when \(n = 100\) was 0.46. Hence the use of the word <strong>estimate</strong>. Normally, as the number of trials is increased the estimate tends to settle down but this is not certain to occur. Theoretically, the probability of obtaining a Head when a fair coin is thrown is \(1/2\). Experimentally, we <strong>expect</strong> the relative frequency to approach \(1/2\)
as \(n\) increases. </p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200701084053+lX6D1d</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200820135259+7I7PK8</text></tag>
    </tags>
    
</question>
</quiz>