<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27110 -->
  <question type="stack">
    <name>
				<text>36.2.2.3 Criteria for rejecting outliers</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Criteria for rejecting outliers</h4><br>
<p>As you already know, outliers may be taken to be observations which lie well outside the range of
most of a sample. They are important for several reasons:</p>
<ol>
<li>they can have a misleading effect on statistics such as the mean and standard deviation;</li>
<li>their occurrence may be due to incorrect observation, measurement or recording. In this case, it is often possible to correct the data;</li>
<li>their presence can induce a false skewness in a data set;</li>
<li>they may actually be members of a population not under consideration. For example, a study
of urban families may involve recording the number of children in a family, say between \(0\) and
\(4\) for the sake of discussion. An outlier might be caused by a rural family with, say, \(10\) children,
living in temporary urban accommodation. This family is part of a different population.</li>
</ol><br>

<p>Simple criteria exist which facilitate the detection of outliers. These criteria should be used with some
caution and never automatically used simply to reject an outlier. You should always ask why such
a value occurred in the first place and work to answer such a question sensibly before considering
rejection. Two criteria for the detection of outliers are given below. Criterion \(1\) may be applied to
data sets that are known to be normal in shape. Criterion \(2\) uses the five-number summary discussed
above and may be applied to any data sets.</p>
<br>
<p><b>Criterion \(1\)</b></p>
<p>Knowing that some \(99.7%\) of a normal population lies within \(3\) standard deviations of the mean, we
could treat any value further than say \(3.3\) standard deviations from the mean as on outlier. This
choice essentially implies that a value has less than \(1\) in a \(1000\) of chance of occurring naturally
outside the range defined by \(3.3\) standard deviations from the mean. Using standardized scores with
as the potential outlier we can state the criterion</p>
<p>&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp; Accept \(x_0\) if \(\left|\dfrac{x_0-\bar{x}}{s_{n-1}}\right| \leq 3.3\)&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp; Investigate if \(\left|\dfrac{x_0-\bar{x}}{s_{n-1}}\right|&gt; 3.3\).</p>
<p>Note that \(\bar{x}\) and \(s_{n-1}\) are sample estimates of the mean and standard deviation of the population.</p>

<br>
<p><b>Criterion \(2\)</b></p>
<p>Using a five-number summary of a data set one can easily set up a criterion which may be used to
classify outliers as either ‘moderate’ or ‘extreme’. The following diagram illustrates the situation where IQR is the Inter-Quartile Range.</p><p><br></p><p style="text-align: center;"><img src="@@PLUGINFILE@@/36.2.2.2%20IQR%20is%20the%20Inter-Quartile%20Range.%20%281%29.svg" alt="diagram showing the situation of our distribution, we use the five-number summary of the data set in terms of the Inter-Quartile Range" role="presentation" class="atto_image_button_text-bottom" width="550" height="103"></p><p style="text-align: center;"><br></p>

<p>While all values classified as outliers should be investigated, this is particularly true of those classified as extreme outliers.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200725204305+mUeo36</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+201102162408+BFvFbF</text></tag>
    </tags>
    
</question>
</quiz>