<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62117 -->
  <question type="stack">
    <name>
				<text>23.3.1.1 Even and odd functions</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Even and odd functions</h4>
We have shown in a previous Section how to calculate, by integration, the coefficients \(a_n\) (\(n =
0, 1, 2, 3, \dots\)) and \(b_n\) (\(n = 1, 2, 3, \dots\)) in a Fourier series. Clearly this is a somewhat tedious process and it is advantageous if we can obtain as much information as possible without recourse to
integration. Also in a previous Section we showed that the square wave (one period of which shown in
the Figure below) has a Fourier series containing a constant term and cosine terms only (i.e. all the Fourier
coefficients \(b_n\) are zero).<p><br></p><p style="text-align: center;"><img src="@@PLUGINFILE@@/23-3-1-fig-12-square-pulse.svg" alt="square-pulse" width="330" height="113" class="img-responsive atto_image_button_text-bottom"><br></p><p><br></p><p>&nbsp;Alternatively the function shown below has a more complicated Fourier series containing both cosine and sine terms as well as a constant.</p><p style="text-align: center;"><img src="@@PLUGINFILE@@/23-3-1-fig-13-saw-tooth.svg" alt="saw tooth" width="501" height="115" class="img-responsive atto_image_button_text-bottom"><br></p>
One can see that the square (shown first) wave has a graph which is symmetrical about the \(y\)-axis and is called an <b>even</b> function. The saw-tooth wave (shown second) has no particular symmetry.
<p></p>
In general a function is called <b>even</b> if its graph is unchanged under reflection in the \(y\)-axis. This is equivalent to
\[ f(-t) = f(t) \hspace{0.5cm} \text{for all} \hspace{0.3cm} t\]
Obvious examples of even functions are \(t^2\), \(t^4\), \(\lvert t \rvert \), \(\cos t\), \(\cos^2 t\), \(\sin^2 t\), \(\cos nt\).
<p></p>
A function is said to be <b>odd</b> if its graph is symmetrical about the origin (i.e. it has rotational
symmetry about the origin). This is equivalent to the condition
\[ f(-t) = -f(t)\]
The Figure below shows an example of an odd function.<p><br></p><p style="text-align: center;"><img src="@@PLUGINFILE@@/23-3-1-fig-14-odd-function.svg" alt="odd-function" width="225" height="148" class="img-responsive atto_image_button_text-bottom"><br></p>
Examples of odd functions are \(t\), \(t^3\), \(\sin t\), \(\sin nt\). A periodic function which is odd is the saw-tooth wave in the Figure below.<p><br></p><p style="text-align: center;"><img src="@@PLUGINFILE@@/23-3-1-fig-15-saw-tooth.svg" alt="saw-tooth-pulse" width="401" height="160" class="img-responsive atto_image_button_text-bottom"><br></p>
Some functions are <b>neither even nor odd</b>. The periodic saw-tooth wave (second Figure from top) is an example; another is the exponential function \(e^t\).]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200825192512+QaER1o</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220518103506+NSIcEX</text></tag>
    </tags>
    
</question>
</quiz>