<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24661 -->
  <question type="stack">
    <name>
				<text>18.1 Introduction</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span>Functions of Several Variables</span><br></h3>

<p>A function of a single variable \(y=f(x)\) is interpreted graphically as a planar curve. In this Section we generalise the concept to functions of more than one variable. We shall see that a function of two variables \(z=f(x,y)\) can be interpreted as a surface. Functions of two or more variables often arise in engineering and in science and it is important to be able to deal with such functions with confidence and skill. We see in this Section how to sketch simple surfaces. In later Sections we shall examine how to determine the rate of change of \(f(x,y)\) with respect to \(x\) and \(y\) and also how to obtain the optimum values of functions of several variables.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200814083616+7gM8ZQ</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+201215125327+m8fGla</text></tag>
    </tags>
    
</question>
</quiz>