<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24616 -->
  <question type="stack">
    <name>
				<text>14.3.1.5 Volume of a solid of revolution</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<div class="HELM_keypoint">
<h4>Key Point</h4>
<p>If the graph of \(y(x)\), between \(x = a\) and \(x = b\), is rotated about the \(x\)-axis the volume of the solid formed is \[\int_a^b \pi y^2 \, \mathrm{d}x\]</p>
</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201201161249+OU3UEh</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210607114524+iXj2Xr</text></tag>
    </tags>
    
</question>
</quiz>