<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26248 -->
  <question type="stack">
    <name>
				<text>7.3.2.5 Evaluating general determinants</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4> Evaluating general determinants</h4>
<p>Cofactors are important as it can be shown that the value of the determinant \(\Delta\) of a \(3\times 3\) matrix \(A\) can be found from the formula \(\Delta=a_{11}A_{11}+a_{12}A_{12}+a_{13}A_{13}.\)</p>

<p>In words, "the determinant of a \(3\times 3\) matrix is obtained by multiplying each element of the first row by its corresponding cofactor and then adding the three together". (In fact this rule can be extended to apply to <b>any</b> row or <b>any</b> column and to <b>any</b> order square matrix.)</p>

<div class="HELM_keypoint">
<h4>Key Point</h4>
<h5>Evaluating General Determinants</h5>
<p>If \(A\) is an \(n\times n\) square matrix then: \(\displaystyle{\operatorname{det}(A) = \sum_{j=1}^{n}a_{ij}A_{ij}}\).</p>
In words: The determinant of a square matrix is obtained by multiplying each element
of row \(i\) by its corresponding cofactor and then adding these products together.<p></p></div>

<hr>
<h4 class="HELM_example">Example</h4>
<p>In the case of \(\Delta=\begin{vmatrix} 4 & 1 & 1 \\ 1 & 2 & 3 \\ 3 & 1 & 2 \end{vmatrix}\) we have \(a_{11}=4, a_{12}=1, a_{13}=1,\) and \begin{align*} A_{11} & = + \begin{vmatrix} 2 & 3 \\ 1 & 2 \end{vmatrix} = 4 - 3 = 1, \\ A_{12} & = - \begin{vmatrix} 1 & 3 \\ 3 & 2 \end{vmatrix} = - ( 2 - 9 ) = 7, \\ A_{13} & = + \begin{vmatrix} 1 & 2 \\ 3 & 1 \end{vmatrix} = 1 - 6 = - 5.\end{align*}</p> 

<p>Hence, \(\Delta=4\times 1 + 1\times 7 + 1\times -5 = 6\).</p>

<p>Alternatively, choosing to expand along the second row:
\begin{align*}\Delta & = a_{21} A_{21} + a_{22} A_{22} + a_{23} A_{23} \\ & = 1 \left( - \begin{vmatrix} 1 & 1 \\ 1 & 2 \end{vmatrix} \right) + 2 \left( \begin{vmatrix} 4 & 1 \\ 3 & 2 \end{vmatrix} \right) + 3 \left( - \begin{vmatrix} 4 & 1 \\ 3 & 1 \end{vmatrix} \right) \\ & = 6\end{align*}
as before.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200725152051+BdhG8R</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210702140448+qSQYWl</text></tag>
    </tags>
    
</question>
</quiz>