<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62107 -->
  <question type="stack">
    <name>
				<text>23.2.4.3 Examples of Fourier series</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[Clearly, in finding the Fourier series of square waves, the integration is particularly simple because \(f(t)\) takes on piecewise constant values. For other functions, such as saw-tooth waves this will
not be the case. Before we tackle such functions however we shall generalise our formulae for the
Fourier coeffcients \(a_n\), \(b_n\) to the case of a periodic function of arbitrary period, rather than confining
ourselves to period \(2\pi\).]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200816153347+qQvMth</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220518094639+HIXov3</text></tag>
    </tags>
    
</question>
</quiz>