<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26261 -->
  <question type="stack">
    <name>
				<text>7.3.2.6 Evaluating general determinants Key Point</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<div class="HELM_keypoint">
<h4>Key Point</h4>
<h5>Evaluating General Determinants</h5>
<p>If \(A\) is an \(n\times n\) square matrix then: \(\displaystyle{\operatorname{det}(A) = \sum_{j=1}^{n}a_{ij}A_{ij}}\).</p>
In words: The determinant of a square matrix is obtained by multiplying each element
of row \(i\) by its corresponding cofactor and then adding these products together.<p></p></div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200812103024+1EoaeB</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200908085713+P66OfQ</text></tag>
    </tags>
    
</question>
</quiz>