<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62085 -->
  <question type="stack">
    <name>
				<text>5.4.1.1 Engineering Case Study - Sound Intensity</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p></p><h4>Engineering Case Study - Sound Intensity</h4><p><b>Introduction to the problem</b></p><p>Many aspects of physics and engineering involve inverse square law dependence. For example gravitational forces and electrostatic forces vary with the inverse square of distance from the mass or charge. The following short case study illustrates this and concerns the dependence of sound intensity on distance from a source.</p><p>For a single source of sound power \(W\) (watts) the dependence of sound intensity magnitude \(I\) (W m\(^{-2}\)) on distance \(r\) (m) from a source is expressed as</p>
<p>\(\qquad I=\dfrac{W}{4\pi r^2}\)</p>
<p>The way in which sounds from different sources are added depends on whether or not there is a phase relationship between them. There will be a phase relationship between two loudspeakers connected to the same amplifier. A stereo system will sound best if the loudspeakers are in phase. The loudspeaker sources are said to be <b>coherent</b>&nbsp;sources. Between such sources there can be reinforcement or cancellation depending on position. Usually there is no phase relationship between two separate items of industrial equipment. Such sources are called <b>incoherent</b>. For two  such incoherent sources \(A\) and \(B\)  the combined sound intensity magnitude (\(I_C\)  W m\(^{-2}\)) at a specific point is given by the sum of the magnitudes of the intensities due to each source at that point. So</p>
<p>\(\qquad I_C=I_A+I_B=\dfrac{W_A}{4\pi r^2_A}+\dfrac{W_B}{4\pi r^2_B}\)</p><p>where \(W_A\) and \(W_B\) are the respective sound powers of the sources; \(r_A\) and \(r_B\) are the respective distances from the point of interest. Note that sound intensity is directional. So if \(A\) and \(B\) are on opposite sides of the receiver's position their intensity contributions will have opposite directions.</p><p><b>Problem in words</b></p><p><span></span></p><p style="">With reference to the situation shown in the figure below, given incoherent point sources \(A\) and \(B\), with sound powers 1.9 W and 4.1 W respectively, 6 m apart, find the sound intensity magnitude at points \(C\) and \(D\) at distances \(p\) and \(q\) from the line joining  \(A\) and \(B\)  and  find the locations of \(C, D\) and \(E\) that correspond to sound intensity magnitudes of 0.02, 0.06 and 0.015  W m\(^{-2}\)        respectively.</p><img src="@@PLUGINFILE@@/5-4-fig-1-problem_in_words.svg" alt="" width="364" height="148" role="presentation" class="img-responsive atto_image_button_text-bottom"><p><br></p><p></p><b>Mathematical statement of problem</b><br><p></p><p></p><p><span></span></p><p style=""><b>(a)</b>\(\quad\) Write down an expression for the sound intensity magnitudes at point \(C\) due to the independent sources \(A\) and \(B\) with powers \(W_A\) and \(W_B\), taking advantage of the symmetry of their locations about the line through \(C\) at right-angles to the line joining \(A\) and \(B\).</p>
<p style=""><b>(b)</b>\(\quad\) Find the expression for \(p\) in terms of\(I_C, W_A, W_B\) and \(m\).</p>
<p style=""><b>(c)</b>\(\quad\)  If \(W_A\) = 1.9 W, \(W_B\) = 4.1 W and \(m\) = 6m calculate the distance \(p\) at which the sound intensity is 0.02 W m\(^{-2}\)?</p>
<p style=""><b>(d)</b>\(\quad\)  Find an expression for the intensity magnitude at point \(D\).</p>
<p style=""><b>(e)</b>\(\quad\) Find the value for \(q\) such that the intensity magnitude at \(D\) is 0.06  W m\(^{-2}\)        and the other values are as in part <b>(c)</b>.</p>
<p style=""><b>(f)</b>\(\quad\) Find an equation in powers of \(n\) relating \(I_E,\) (intensity magnitude at point \(E\)) \(W_A, W_B, n\) and \(m\).</p>
<p style=""><b>(g)</b>\(\quad\)  By plotting this function for \(I_E\) = 0.015  W m\(^{-2}\), \(m\) = 6 m, \(W_A\) = 1.9 W, \(W_B\) = 4.1 W, find the corresponding values for \(n\).</p><p style=""><b>Mathematical analysis</b><br></p><p style=""><span></span></p><p style=""><span style=""></span></p><p><b>(a)</b>\(\quad\) The combined sound intensity magnitude \(I_C\)  W m\(^{-2}\)        is given by the sum of the intensity magnitudes due to each source at \(C\). Because of symmetry of the position of \(C\) with respect to \(A\) and \(B\), write \(|\overrightarrow{AC}|=|\overrightarrow {BC}|=r\), then</p><p></p><p>\(\qquad I_C=I_A+I_B=\dfrac{W_A}{4\pi r_A^2}+\dfrac{W_B}{4\pi r_B^2}=\dfrac{W_A+W_B}{4\pi r^2}\).&nbsp;</p><p>Using Pythagoras' theorem,</p><p>\(\qquad r^2=\left(\dfrac{m}{2}\right)^2+p^2\quad\) hence \(\quad I_C=\dfrac{W_A+W_B}{4\pi ((m/2)^2+p^2)}=\dfrac{W_A+W_B}{\pi (m^2+4p^2)}\)</p><p></p><p><b>(b)</b>\(\quad\) Making \(p\) the subject of the last formula,</p><p></p><p>\(\qquad p=\pm \dfrac{1}{2}\sqrt{(\dfrac{W_A+W_B}{\pi I_C})-m^2}\)</p>
<p>The result that there are two possible values of \(p\) is a consequence of the symmetry of the sound field about the line joining the two sources. The positive value gives the required location of \(C\) above the line joining \(A\) and \(B\) in the figure above. The negative value gives a symmetrical location 'below' the line. Note also that if \(0=\dfrac{W_A+W_B}{\pi I_C}-m^2\) or \(I_C=\dfrac{(W_A+W_B)}{\pi m^2}\), then \(p=0\), i.e. \(C\) would be on the line joining \(A\) and \(B\).</p><p></p><p><b>(c)</b>\(\quad\) Using the given values, \(p=3.86\) m.</p><p><b>(d)</b>\(\quad\) Using Pythagoras' theorem again, the distance from \(A\) to \(D\) is given by \(\sqrt{q^2+m^2}\). So</p><p>\(\qquad I_D=I_A+I_B=\dfrac{W_A}{4\pi r_A^2}+\dfrac{W_B}{4\pi r_B^2}=\dfrac{W_A}{4\pi (q^2+m^2)}+\dfrac{W_B}{4\pi q^2}\)</p><p></p><p><b>(e)</b>\(\quad\) Multiplying through by \(4\pi q^2(q^2+m^2)\)&nbsp; and collecting together like powers of \(q\)&nbsp; produces a quartic equation,</p><p></p><p>\(\qquad 4\pi I_Dq^4+[4\pi m^2I_D-(W_A+W_B])q^2-W_Bm^2=0.\)&nbsp;</p><p>Since the quartic equation contains only even powers of \(q\), it can be regarded as a quadratic equation in \(q^2\)&nbsp; and this can be solved by the standard formula. Hence</p>
<p>\(\qquad q^2=\dfrac{-[4\pi m^2I_D-(W_A+W_B)]\pm \sqrt{[4\pi m^2I_D-(W_A+W_B)]^2+16\pi I_DW_Bm^2}}{8\pi I_D}\)&nbsp;</p>
<p>Using the given values,</p>
<p>\(\qquad q^2=\dfrac{-21.14\pm29.87}{1.51}\)&nbsp;</p>
<p>Since \(q\)&nbsp; must be real, the negative result can be ignored. Hence \(q\approx 2.40\) m.</p><p></p><p><b>(f)</b>\(\quad\) Using the same procedure as in <b>(d)</b> and <b>(e)</b>,</p><p></p><p>\(\qquad I_E=I_A+I_B=\dfrac{W_A}{4\pi r_A^2}+\dfrac{W_B}{4\pi r_B^2}=\dfrac{W_A}{4\pi (m+n)^2}+\dfrac{W_B}{4\pi n^2}\)</p>
<p>\(\qquad 4\pi I_Dn^2(m+n)^2I_E=W_An^2+(m+n)^2W_B=0\)</p>
<p>A general expression for the distance \(n\) at which the intensity at point \(E\) is \(I_E\) is given by collecting like powers of \(n$ and is another quartic equation, i.e.</p>
<p>\(\qquad 4\pi I_En^4+8\pi I_Emn^3+[4\pi I_Em^2-(W_A+W_B)]n^2-2mW_Bn-m^2W_B=0\)</p>
<p>Unfortunately this cannot be treated simply as a quadratic equation in \(n^2\) since there are terms in odd powers of \(n\). One way forward is to plot the curve corresponding to the equation after substituting the given values, another is to use a numerical method such as Newton-Raphson.</p><p></p><p><b>(g)</b>\(\quad\) Substitution of the given values produces the equation</p><p></p><p>\(\qquad 0.1885n^4+2.2619n^2+0.7858n^2-49.2n-147.6=0.\)</p>
<p>The plot of the quartic equation in the figure below shows that there are two roots of interest. Use of a numerical method for finding the roots of polynomials gives values of the roots to any desired accuracy i.e. \(n\approx 4.876\) m and \(n\approx -9.628\) m.</p><p><img src="@@PLUGINFILE@@/5-4-fig-2-solution.svg" alt="" width="372" height="177" role="presentation" class="img-responsive atto_image_button_text-bottom"><br></p><p></p><p style=""></p><span><b></b></span><p><b>Interpretation</b></p><p>The result for part <b>(g)</b> implies that there are two locations for \(E\) along the line joining the two sources where the intensity magnitude will have the given value. One position is about 3.6 m to the left of source \(A\)  and the other is  about 4.9 m to the right of source \(B\).</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200810092537+oSWQMA</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200810110259+g01Fui</text></tag>
    </tags>
    
</question>
</quiz>