<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26750 -->
  <question type="stack">
    <name>
				<text>10.2.3.9 Find z1/z2 in polar form (NOT IN QUIZ)</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<ol class="HELM_parts">
<p>If \( \displaystyle{z_1} = {@z1@}\) and&nbsp;\( \displaystyle{z_2} = {@z2@}\) find:</p>

<li><p> \(\displaystyle{ \left|\dfrac{z_1}{z_2}\right| =} \) [[input:ans1]] [[validation:ans1]]&nbsp;[[feedback:prt1]]</p></li>

<li><p>\( \displaystyle{\text{arg}\left(\dfrac{z_1}{z_2}\right) =} \)&nbsp;[[input:ans2]] [[validation:ans2]]&nbsp;[[feedback:prt2]]</p></li>
<p>Give your answers to \(3\) d.p.<br></p></ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<ol class="HELM_parts">
<p>\( z_1 = {@z1@} \) so \( |z_1| = {@r1@} \) and&nbsp;&nbsp;\( \text{arg}(z_1) = {@theta1@}\)&nbsp;</p>
<p>\( z_2 = {@z2@} \) so \( |z_2| = {@r2@} \) and&nbsp;&nbsp;\( \text{arg}(z_2) = {@theta2@}\)&nbsp;<br></p>
  <li>
<p> \( \left|\dfrac{z_1}{z_2}\right| = \dfrac{|z_1|}{|z_2|} = {@ta1@} \)<br></p>
  </li>
  <li>
    <p>\( \text{arg}\left(\dfrac{z_1}{z_2}\right) = \text{arg}(z_1) - \text{arg}(z_2) = {@ta2@} \)</p> </li>
<p>Have in mind that \( -\pi &lt; \text{arg}(z) \leq \pi \) so you may need to add or subtract a \( \pi\) from your \( \text{arg} \) results.</p>
</ol>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>a1: rand_with_prohib(-15,15,[0]);
b1: rand_with_prohib(-12,12,[0]);
a2: rand_with_prohib(-10,10,[0]);
b2: rand_with_prohib(-8,8,[0]);
z1: a1+(b1)*%i;
z2: a2+(b2)*%i;
r1: sqrt((a1)^2+(b1)^2);    /*mod z1*/
r2: sqrt((a2)^2+(b2)^2);    /*mod z2*/
r: decimalplaces(float(r1/r2),3);              /*mod z1/z2 */
theta1: decimalplaces(float(atan2(b1, a1)),4);
theta2: decimalplaces(float(atan2(b2, a2)),4);

x: (a1*a2+b1*b2)/(a2*a2+b2*b2)            /*Real part of z */
y: (a2*b1-a1*b2)/(a2*a2+b2*b2)           /*Imaginary part of z*/
theta: decimalplaces(float(atan2(y,x)),3);  /* Arg of z */

ta1: dispdp(r,3);
ta2: dispdp(theta,3);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\( z_1 = {@z1@}\), \( z_2 = {@z2@}\) so \( |z_1 z_2| = {@ta1@}\) and \( \text{arg}(z_1) = {@theta1@}\), \( \text{arg}(z_2) = {@theta2@}\)  so \( \text{arg}(z_1 z_2) = {@ta2@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>units</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>units</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>r</tans>
        <testoptions>3</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>r</tans>
        <testoptions>0.005</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>theta</tans>
        <testoptions>3</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>theta</tans>
        <testoptions>0.005</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200629221628+pvShBn</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210702132246+RMzxMR</text></tag>
    </tags>
    <deployedseed>17268628</deployedseed>
    <deployedseed>733335117</deployedseed>
    <deployedseed>1411664107</deployedseed>
    <deployedseed>24324902</deployedseed>
    <deployedseed>48626672</deployedseed>
    <deployedseed>2094651466</deployedseed>
    <deployedseed>452720233</deployedseed>
    <deployedseed>1737094112</deployedseed>
    <deployedseed>1405500824</deployedseed>
    <deployedseed>1032873779</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(r+0.006,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(theta-0.007,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-2-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>