<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62110 -->
  <question type="stack">
    <name>
				<text>23.2.5.3 FS saw-tooth example</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<hr>
<h4 class="HELM_example">Example</h4>
<p>Find the Fourier series of the function shown in the Figure below which is a saw-tooth
wave with alternate portions removed.</p><p><br></p><p style="text-align: center;"><img src="@@PLUGINFILE@@/23-2-2-fig-11-saw-tooth-pulse.svg" alt="saw-tooth-pulse" width="501" height="115" class="img-responsive atto_image_button_text-bottom"><br></p>


<p></p>
<h4 class="HELM_solution">Solution</h4>
Here the period \(T = 2l = 4\) so \(l = 2\). The Fourier series will have the form
\[ f(t) = \dfrac{a_0}{2} + \sum_{n=1}^\infty \left\{ a_n \cos \left( \dfrac{n\pi t}{2} \right) + b_n \sin \left( \dfrac{n\pi t}{2} \right) \right\} \]
The coefficients \(a_n\) are given by
\[ a_n = \dfrac{1}{2} \int _{-2}^2 f(t) \cos  \left( \dfrac{n\pi t}{2} \right) \hspace{0.15cm} dt \]
where
\[
f(t) = 
\begin{cases}
0 \hspace{1cm} -2 &lt; t &lt; 0\\
t \hspace{1cm}  &lt; t &lt; 2
\end{cases}
\]
\[f(t+4) = f(t)\]
Hence \( a_n = \dfrac{1}{2} = \displaystyle{\int_{0}^2} t \cos  \left( \dfrac{n \pi t}{2} \right) \hspace{0.15cm} dt \).
<p></p>
The integration is readily performed using integration by parts:
\[
\begin{align}
 \displaystyle{\int_{0}^2} t \cos  \left( \dfrac{n \pi t}{2} \right) \hspace{0.15cm} dt  &amp;= \left[ t \dfrac{2}{n\pi}\sin \left( \dfrac{n\pi t}{2} \right) \right]_0^2 - \dfrac{2}{n\pi}  \displaystyle{\int_{0}^2} \sin \left( \dfrac{n\pi t}{2} \right)  \hspace{0.15cm} dt \\
&amp;= \dfrac{4}{n^2 \pi^2} \left[ \cos \left( \dfrac{n\pi t}{2} \right) \right]_0^2 \hspace{1cm} n \neq 0 \\
&amp;= \dfrac{4}{n^2 \pi^2} (\cos n\pi - 1).
\end{align}
\]
Hence, since \( a_n = \dfrac{1}{2} = \displaystyle{\int_{0}^2} t \cos  \left( \dfrac{n \pi t}{2} \right) \hspace{0.15cm} dt \), we have
\[
a_n = 
\begin{cases}
0 \hspace{2.2cm} n = 2, 4, 6, \dots \\
-\dfrac{4}{n^2\pi^2} \hspace{1cm} n = 1, 3, 5, \dots
\end{cases}
\]
The constant term is \( \dfrac{a_0}{2}\) where \( a_0 = \dfrac{1}{2} \displaystyle{\int_{0}^2} t \hspace{0.15cm}dt = 1\).
<p></p>
Similarly,
\[  b_n = \dfrac{1}{2} = \displaystyle{\int_{0}^2} t \sin  \left( \dfrac{n \pi t}{2} \right) \hspace{0.15cm} dt \]
where
\[ \displaystyle{\int_{0}^2} t \sin  \left( \dfrac{n \pi t}{2} \right) \hspace{0.15cm} dt  = \left[- t \dfrac{2}{n\pi}\cos \left( \dfrac{n\pi t}{2} \right) \right]_0^2 + \dfrac{2}{n\pi}  \displaystyle{\int_{0}^2} \cos \left( \dfrac{n\pi t}{2} \right)  \hspace{0.15cm} dt \]
The second integral gives zero. Hence
\[ 
b_n = - \dfrac{2}{n\pi}\cos n\pi = 
\begin{cases}
-\dfrac{2}{n\pi} \hspace{1cm} n = 2, 4, 6, \dots \\
+\dfrac{2}{n\pi} \hspace{1cm} n = 1, 3, 5, \dots
\end{cases}
\]
Hence, using all these results for the Fourier coefficients, the required Fourier series is
\[
\begin{align}
f(t) = \dfrac{1}{2} &amp;- \dfrac{4}{\pi^2} \left\{ \cos \left( \dfrac{\pi t}{2} \right) + \dfrac{1}{9} \cos \left( \dfrac{3 \pi t}{2} \right) + + \dfrac{1}{25} \cos \left( \dfrac{5 \pi t}{2} \right) + \dots \right\} \\
&amp;+ \dfrac{2}{\pi}\left\{ \sin \left( \dfrac{\pi t}{2} \right) - \dfrac{1}{2} \sin \left( \dfrac{2 \pi t}{2} \right)  + \dfrac{1}{3} \sin \left( \dfrac{3 \pi t}{2} \right) + \dots \right\} 
\end{align}
\]
Notice that because the Fourier coefficients depend on \( \dfrac{1}{n^2}\) (rather than \( \dfrac{1}{n}\) as was the case for the square wave) the sinusoidal components in the Fourier series have quite rapidly decreasing amplitudes. We would therefore expect to be able to approximate the original saw-tooth function using only a quite small number of terms in the series.
<hr>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200821200352+r3Cyhn</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220518095225+DOCSkl</text></tag>
    </tags>
    
</question>
</quiz>