<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27216 -->
  <question type="stack">
    <name>
				<text>35.3.1.2 Mutually Exclusive or Not?</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p></p><p>Decide which of the following pairs of events (\(A\) and \(B\)) arising from the experiments described are mutually exclusive.<br></p><ol><li>Two cards are drawn from a pack:<br>

\(A=\) a red card is drawn.<br>

\(B=\) a picture card is drawn.<br>
\(A\) and \(B\) are [[2]]</li>
<li>The daily traffic accidents in Edinburgh involving pedal cyclists and motorcyclists are counted.<br>

\(A=\) three motorcyclists are injured in collisions with cars.<br>

\(B=\) one pedal cyclist is injured when hit by a bus.<br>
\(A\) and \(B\) are [[1]]</li>

<li>A box contains \(20\) nuts. Some have a metric thread, some have a British Standard Fine (BSF) thread and some have a&nbsp; British Standard Whitworth (BSW) thread.<br>

\(A=\) first nut picked out of the box is BSF.<br>

\(B=\) second nut picked out of the box is metric.<br>
\(A\) and \(B\) are [[2]]</li></ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200630130842+X4wZGe</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200914233042+XbZZDp</text></tag>
    </tags>
    
</question>
</quiz>