<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26566 -->
  <question type="stack">
    <name>
				<text>7.1.2.3 Classify the following matrices</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p><p>Classify the following \(3\times 3\) matrices. If multiple options are valid then choose the most restrictive option.<br></p><p>\(A=\left[\begin{array}{lll}\hfill 1&amp;\hfill 1&amp;\hfill 2\\ \hfill 1&amp;\hfill 1&amp;\hfill 1\\ \hfill 0&amp;\hfill 1&amp;\hfill 3 \end{array}\right]\), \(B=\left[\begin{array}{lll}\hfill 1&amp;\hfill 0&amp;\hfill 0\\ \hfill 1&amp;\hfill 1&amp;\hfill 0\\ \hfill 0&amp;\hfill 1&amp;\hfill 3 \end{array}\right]\), \(C=\left[\begin{array}{lll}\hfill 1&amp;\hfill 2&amp;\hfill 1\\ \hfill 0&amp;\hfill 1&amp;\hfill 4\\ \hfill 0&amp;\hfill 0&amp;\hfill 1 \end{array}\right]\), \(D=\left[\begin{array}{lll}\hfill 1&amp;\hfill 0&amp;\hfill 0\\ \hfill 0&amp;\hfill 3&amp;\hfill 0\\ \hfill 0&amp;\hfill 0&amp;\hfill 2 \end{array}\right]\)</p><p></p><p>Matrix \(A\) is: [[1]]</p><p>Matrix \(B\) is: [[3]]</p><p>Matrix \(C\) is: [[4]]</p><p>Matrix \(D\) is: [[2]]<br></p><br><p></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200619104755+gFQAva</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200912104250+H943Xd</text></tag>
    </tags>
    
</question>
</quiz>