<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 33692 -->
  <question type="stack">
    <name>
				<text>19.2.5.8 Solving an Exact Equation Ex5</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<ol class="HELM_parts">
    <li>
        <p>Is the equation \(x^{{@aa@}}\dfrac{\mathrm{d}y}{\mathrm{d}x} + {@(aa+1)*x^(aa-1)@}y = {@bb@}x^3\) exact? [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
    </li>

    <li>
        <p>Multiply \(x^{{@aa@}}\dfrac{\mathrm{d}y}{\mathrm{d}x} + {@(aa+1)*x^(aa-1)@}y&nbsp; = {@bb@}x^3\) by \(x\) and show that the resulting equation is exact and obtain its solution.</p>

        <p>Exact form: \(\dfrac{\mathrm{d}}{\mathrm{d}x} \bigg(\) [[input:ans2]] [[validation:ans2]]&nbsp; \(\bigg) = {@bb@}x^4\).&nbsp;[[feedback:prt2]]</p>

        <p>\(y =\) [[input:ans3]] [[validation:ans3]] \( + \dfrac{C}{x^4}\). [[feedback:prt3]] </p>

    </li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<ol class="HELM_parts">
  <li>
    <p>We are given the equation \(x^{{@aa@}}\dfrac{\mathrm{d}y}{\mathrm{d}x} + {@(aa+1)*x^(aa-1)@}y&nbsp; = {@bb@}x^3\). It cannot be written in the form \[ f(x) \dfrac{\mathrm{d}y}{\mathrm{d}x} + y f^{\prime}(x) = g(x).\] Therefore it is not exact.</p> 
	</li>
  
  <li>
    <p>Multiplying the equation by \(x\) yields the equation \(x^{{@aa+1@}}\dfrac{\mathrm{d}y}{\mathrm{d}x} + {@aa+1@}x^{{@aa@}}y = {@bb@}x^4.\)</p>
    <p>The equation is of an exact form since it can be written as \[ f(x) \dfrac{\mathrm{d}y}{\mathrm{d}x} + y f^{\prime}(x) = g(x)\] with \[f(x) = x^{{@aa+1@}} \quad \text{and} \quad g(x)={@bb@}x^4.\] Hence the exact form becomes \[\dfrac{\mathrm{d}}{\mathrm{d}x}(yf(x))
    = g(x), \quad \text{i.e} \quad \dfrac{\mathrm{d}}{\mathrm{d}x}(yx^{{@aa+1@}}) = {@bb@}x^4.\]</p>
    
    <p>Integrating both sides gives \[ yx^{{@aa+1@}} = \int {@bb@}x^4 \ \mathrm{d}x\] which implies  \[yx^{{@aa+1@}} = \dfrac{{@bb@}}{5}x^5 + C.\] Rearranging to find \(y\) we get \[y = {@(bb/5)*x^(5-aa-1)@} + \dfrac{C}{x^{{@aa+1@}}}.\]</p>
  </li>
</ol>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>opts: [["Yes",false],["No", true]];
ta1: "No";

aa: 2+rand(2);
bb:2+rand(5);

ta2: y*x^(aa+1);
ta3: (bb/5)*x^(5-aa-1);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text> \(x^{{@aa+1@}}\dfrac{\mathrm{d}y}{\mathrm{d}x} + {@aa+1@}x^{{@aa@}}y = {@bb@}x^4, y = {@ta3@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>opts</tans>
      <boxsize>1</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210616114333+tJ0Hox</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210616120132+MzgHaP</text></tag>
    </tags>
    <deployedseed>1120893740</deployedseed>
    <deployedseed>1200539485</deployedseed>
    <deployedseed>2137595740</deployedseed>
    <deployedseed>1511181185</deployedseed>
    <deployedseed>1965998421</deployedseed>
    <deployedseed>1135901494</deployedseed>
    <deployedseed>1513231001</deployedseed>
    <deployedseed>1273347158</deployedseed>
    <deployedseed>1338569865</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(opts))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>"Yes"</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(2*ta2,simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ev(2*ta3,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>