<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26287 -->
  <question type="stack">
    <name>
				<text>22.3.2.7 Orthogonal vectors</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<div class="HELM_keypoint">
<h4>Key Point</h4>
<p>Two \(n\times 1\) column vectors \(X\) and \(Y\) are <strong>orthogonal</strong> if \(X^TY=\mathbf{0}\) or (equivalently) \(Y^TX=\mathbf{0}\).</p>
</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200624093700+VPXJmv</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200624095811+CuEi0C</text></tag>
    </tags>
    
</question>
</quiz>