<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 33869 -->
  <question type="stack">
    <name>
				<text>6.2.2.3 Key point</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<div class="HELM_keypoint">
<h4>Key Point: </h4>
<p>The fundamental identity relating hyperbolic functions is:</p>
<p style="text-align: center;">\(\rm cosh^2 x - \sinh^2 x \equiv 1\)</p>
<p>This is the hyperbolic function equivalent of the trigonometric identity: \( \cos^2x + \sin^2 x \equiv 1 \)</p>
</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210716110741+ZJpuqF</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210716111404+44hP4v</text></tag>
    </tags>
    
</question>
</quiz>