<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27272 -->
  <question type="stack">
    <name>
				<text>36.1.5.4 (for Edinburgh) Median and quartiles</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Median and quartiles</h4>
<p>Some statistics are derived from quantitative data which are placed in <b>rank order</b>. Ranking data simply means that the data are placed in order from the smallest to the largest. The smallest observation thus gets rank \(1\), the second smallest gets rank \(2\) and so on. Such statistics are called <i>ordered statistics</i>. The median is an order statistic which is used as a measure of location. It divides a sample of data into two equally sized groups, one containing the smaller observations and the other containing the larger observations. Two other order statistics which we will describe here together with the median are the <b>Lower Quartile</b> and the <b>Upper Quartile</b>. The lower quartile, median and upper quartile are also called the first, second and third quartiles respectively. They divide a sample of data into four portions according to the sizes of the observations, each portion containing the same number of observations. As you will see the <i>Ogive</i> or <i>Cumulative Frequency Curve</i> enables us to find these statistics for large data sets.</p>

<p>Definitions for the three sample quartiles are given below. A number of slightly different definitions are used for the lower and upper quartile. One version is given here. The different definitions will usually lead to only slightly different results.</p>

<p><b>The Median</b>; this is the central value of a sample where the numbers of larger observations and smaller observations are equal.</p>

<p><b>The Lower Quartile</b>; this is the value which divides the observations into a lower \(25\%\) and an upper \(75\%\).</p>

<p><b>The Upper Quartile</b>; this is the value which divides the observations into a lower \(75\%\) and an upper \(25\%\).</p><br>
<h5>Calculations</h5>
<p>We calculate the median and the lower and upper quartiles as follows.
Let the number of observations be \(n\). We sort the observations into increasing
order and allocate ranks. If there are ties then we just use integer ranks as usual.
For example, the data \(7.8, 4.6, 5.3, 3.1, 9.5, 5.3\) is sorted into 
\(3.1, 4.6, 5.3, 5.3, 7.8, 9.5\) and assign the corresponding ranks of \(1, 2, 3, 4, 5, 6\)
so that the observations with ranks \(3\) and \(4\) are both \(5.3\). 
We use the bracket notation \(x_{(i)}\) to indicate the observation with rank \(i\), so 
\(x_{(1)} = 3.1\) is the smallest value and \(x_{(n)} = x_{(6)} = 9.5\) is the largest.</p>
<ul>
<li>For the median, calculate the rank \(r_2 = \dfrac{n+1}{2}\). If \(n\) is odd then \(r_2\)
is an integer and the sample median is the observation \(x_{(r_2)}\). If \(n\) is even
then \(r_2\) is not an integer and we average the values on either side with ranks 
\(\dfrac{n}{2}\) and \(\dfrac{n+2}{2}\). For example, if \(n=15\) then \(r_2=7.5\) and we
average the values with ranks \(7\) and \(8\), namely \(\dfrac{1}{2} \left( x_{(7)} + x_{(8)} \right)\).</li>

<li>For the lower quartile, calculate \(r_1 = \dfrac{n+2}{4}\). If \(r_1\) is an integer, 
the sample quartile is the observation with rank \(r_1\). If \(r_1\) is not an integer
then we take the <i>weighted</i> average of the values on either side. For example,
if \(n=15\) then \(r_1 = 4.25\). Since this rank is closer to \(4\) than to \(5\) the lower 
quartile is calculated by giving more weight to the lower value: 
\(\dfrac{1}{4} \left( 3x_{(4)} + x_{(5)} \right)\).</li>

<li>For the upper quartile, calculate \(r_3 = \dfrac{3n+2}{4}\). If \(r_1\) is an integer, 
the sample upper quartile is the observation with rank \(r_3\). If \(r_3\) is not an 
integer then take the <i>weighted</i> average of the values on either side. For example,
if \(n=15\) then \(r_3 = 11.75\). Since this rank is closer to \(12\) than to \(11\) the upper 
quartile is calculated by giving less weight to the lower value: 
\(\dfrac{1}{4} \left( x_{(11)} + 3x_{(12)} \right) \).</li>
</ul>
<p>Note that \(r_2 = \dfrac{2n+2}{4}\) so the three rank \(r\)-values are:
\[
  r_1 = \dfrac{n+2}{4}, \quad r_2 = \dfrac{2n+2}{4}, \quad \mbox{and}\quad r_3 = \dfrac{3n+2}{4}.
\]</p>
<p>If the rank is an integer then the median/quantile is the observation with that rank.
If not, then the median/quantile is calculated by one of the following <i>weighted</i>
average formulas below depending on where the \(r\)-value falls in the interval 
between the two integer neighbouring ranks of \(m\) and \(m+1\)</p>
\[\begin{align*}
	x_{(m + 0.25)} &amp;= \dfrac{1}{4} \left( 3x_{(m)} +  x_{(m+1)} \right)\\
	x_{(m + 0.5 )} &amp;= \dfrac{1}{2} \left( x_{(m)} +  x_{(m+1)} \right)\\
	x_{(m + 0.75)} &amp;= \dfrac{1}{4} \left( x_{(m)} + 3x_{(m+1)} \right)
\end{align*}\]
<hr>
<h4 class="HELM_example">Example</h4>
<p>For the simple data set:</p>
\[1.2, 3.0, 2.5, 5.1, 3.5, 4.1, 3.1, 2.4\]
the process is illustrated by replacing the members of the data set in rank order
\[1.2, 2.4, 2.5, 3.0, 3.1, 3.5, 4.1, 5.1\]<p></p>
<p>Here \(n=8\) so
</p><ul>
<li>\(r_1 = 2.5\) and the lower quartile is \(\dfrac{2.4 + 2.5}{2} = 2.45\)</li>
<li>\(r_2 = 4.5\) and the median is \(\dfrac{3.0 + 3.1}{2} = 3.05\)</li>
<li>\(r_3 = 6.5\) and the upper quartile is \(\dfrac{3.5 + 4.1}{2} = 3.8\)</li>
</ul><p></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200907003324+YZLlUM</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+201104095213+3MjgKN</text></tag>
    </tags>
    
</question>
</quiz>