<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 35630 -->
  <question type="stack">
    <name>
				<text>1.5.2.1 Substitution</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise <br></p>
<p>The formula for the volume of a cylinder is \(V = \pi r^2h \).
Find \(V\) when \(r = {@ru1@}\), \(h = {@hu1@}\).&nbsp;<span style="font-size: calc(0.90375rem + 0.045vw);">Round your answer to 3 significant figures and include the appropriate units.</span></p>
<br>
<p>\(V =\) [[input:ans1]] [[validation:ans1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>We have \(r = {@ru1@}\) and \(h = {@hu1@}\). Substituting these values into the formula we find:\[\begin{aligned} V &amp;=\pi r^2h\\ &amp;= \pi \cdot ({@ru1@})^2 \cdot {@hu1@}\\ &amp;= \pi \cdot {@ru1^2@} \cdot {@hu1@}\\ &amp;= \pi \cdot {@ru1^2*hu1@}\\ &amp;= {@stackunits(float(%pi*(r1^2)*h1), cm^3)@}\\ &amp;= {@ta@} \text{ }_{\text{3.s.f}} \qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\end{aligned}\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021072200</text>
		</stackversion>
    <questionvariables>
				<text>stack_unit_si_declare(true);
eqs: V = %pi*(r^2)*h;
r1:rand([4,5,6]);
ru1:stackunits(r1, cm);
h1:rand([10,15,20]);
hu1:stackunits(h1, cm);
ta:stackunits(dispsf(float(%pi*(r1^2)*h1),3), cm^3);
ta1:stackunits(%pi*(r1^2)*h1, cm^3);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>\(V=\pi\cdot{@r1@}^2\cdot{@h1@} = {@ta@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>units</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>4</insertstars>
      <syntaxhint>3 s. f.</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>pi</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>minsf:3</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>Units</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>[3,-1]</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210805115441+UpOmV8</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220113102443+bFDvPi</text></tag>
    </tags>
    <deployedseed>1996361289</deployedseed>
    <deployedseed>1262711757</deployedseed>
    <deployedseed>41847470</deployedseed>
    <deployedseed>460002768</deployedseed>
    <deployedseed>1292102533</deployedseed>
    <deployedseed>909780961</deployedseed>
    <deployedseed>39396133</deployedseed>
    <deployedseed>766419246</deployedseed>
    <deployedseed>2107104156</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore></expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>stackunits(dispsf(float(%pi*(r1^2)*h1/100^3),3), m^3)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>