<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24667 -->
  <question type="stack">
    <name>
				<text>18.1.1.5 Find values of a function with three variables</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Given \(f(x,y,u) = {@ff@}\), find the values of <span class="math inline">\(f({@x1@},{@y1@},{@u1@}), \, f({@x2@},{@y2@},{@u2@}).\)</span></p>
<ol class="HELM_parts">

<li><p>\(f({@x1@},{@y1@},{@u1@})=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p></li>

<li><p>\(f({@x2@},{@y2@}, {@u2@})=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p></li>

</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p dir="ltr" style="text-align: left;"></p>
<p>For the function \(f(x,y,u) = {@ff@}\).</p>
<ol class="HELM_parts">
    <li>
        <p>\(f({@x1@},{@y1@},{@u1@} )= {@f1@} = {@ta1@}.\)</p>
    </li>
    <li>
        <p>\(f({@x2@},{@y2@},{@u2@})={@f2@} = {@ta2@}.\)&nbsp;</p>
    </li>
</ol>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021120900</text>
		</stackversion>
    <questionvariables>
				<text>ordergreat(x,y,u);
x1:rand(6)-3;
y1: rand_with_prohib(-4,4,[x1]);
u1: rand_with_prohib(-4,4,[x1,y1]);
x2: rand_with_prohib(-2,6,[x1]);
y2: rand_with_prohib(-2,6,[x1,x2]);
u2: rand_with_prohib(-2,6,[x1,x2,y1,y2]);

a:rand(2)+1;
b:rand(2)+1;
c: rand_with_prohib(1,5,[a,b]);

ff: a*x^2 + b*y*u +c;   /* use this in the question text*/
simp:false;
f(x,y,u):=a*x^2 + b*y*u +c;
f1:f(x1,y1,u1);
f2:f(x2,y2,u2);
ta1: ev(f1,simp);
ta2:ev(f2,simp);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(f(x,y) = {@ff@}, \quad f({@x1@},{@y1@},{@u1@})= {@ta1@}, \quad f({@x2@},{@y2@},{@u2@}) = {@ta2@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201215140335+LJLmgG</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220304103253+qcgZzR</text></tag>
    </tags>
    <deployedseed>1915784369</deployedseed>
    <deployedseed>987169990</deployedseed>
    <deployedseed>544422320</deployedseed>
    <deployedseed>1015713948</deployedseed>
    <deployedseed>1566325582</deployedseed>
    <deployedseed>800005623</deployedseed>
    <deployedseed>2096277416</deployedseed>
    <deployedseed>441830647</deployedseed>
    <deployedseed>1186544133</deployedseed>
    <deployedseed>318434446</deployedseed>
    <deployedseed>39890811</deployedseed>
    <deployedseed>46728647</deployedseed>
    <deployedseed>2114192350</deployedseed>
    <deployedseed>1103112512</deployedseed>
    <deployedseed>1137344271</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1-1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2-1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>