<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62145 -->
  <question type="stack">
    <name>
				<text>43.1.1.4 Regression lines by eye</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>Regression lines by eye</h3>

Note that at a very simple level, we may look at the data and, using an `educated guess', draw
a line of regression `by eye' through a set of points. However, finding a regression line by eye is
unsatisfactory as a general statistical method since it involves guess-work in drawing the line with
the associated errors in any results obtained. The guess-work can be removed by the method of least
squares in which the equation of a regression line is calculated using data. Essentially, we calculate
the equation of the regression line by minimising the sum of the squared vertical distances between
the data points and the line.]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200630194440+9SFkFL</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200715053426+2kz6wA</text></tag>
    </tags>
    
</question>
</quiz>