<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 35934 -->
  <question type="stack">
    <name>
				<text>1.5.2.15 Transposition</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise <br></p>
<p>Transpose the formula {@eqn@} for {@b1@}</p>
<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>Since we are trying to make {@b1@} the subject of the formula we first, to both sides of the equation, {@addsub@} \({@coef1*b1@}\) which gives: \[{@eqn1@}\] We then subtract 17 from both sides of the equation giving: \[{@eqn2@}\]. Finally we divide both sides of the equation by {@coef2@} giving: \[{@ta1@}\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021072200</text>
		</stackversion>
    <questionvariables>
				<text>powerdisp:true

a1:rand([7,8]);
a2:rand([4,5]);
eqn: a1*u+a2*v-3*w=17;

n1: rand(3)+1;
b1: [u,v,w][n1];
ta1: [u=-a2/a1*v+3/a1*w+17/a1,v=-a1/a2*u+3/a2*w+17/a2,w=a1/3*u+a2/3*v-17/3][n1];
ta1b: [u=1/a1*(-a2*v+3*w+17),v=1/a2*(-a1*u+3*w+17),w=1/3*(a1*u+a2*v-17)][n1];

addsub: ["subtract","subtract","add"][n1];
coef1: [a1,a2,3][n1];
coef2: [-a1,-a2,3][n1];
eqn1: [a2*v-3*w=-a1*u+17,a1*u-3*w=-a2*v+17,a1*u+a2*v=3*w+17][n1];
eqn2: [a2*v-3*w-17=-a1*u,a1*u-3*w-17=-a2*v,a1*u+a2*v-17=3*w][n1];</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>Transpose {@eqn@} for {@b1@}</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>EqualComAssRules</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>[oneMul]</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>EqualComAssRules</answertest>
        <sans>ans1</sans>
        <tans>ta1b</tans>
        <testoptions>[oneMul]</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210816170218+BRmDWh</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+211213113304+A8n3oi</text></tag>
    </tags>
    <deployedseed>11960216</deployedseed>
    <deployedseed>344708093</deployedseed>
    <deployedseed>1783611020</deployedseed>
    <deployedseed>1882603940</deployedseed>
    <deployedseed>340061373</deployedseed>
    <deployedseed>510334182</deployedseed>
    <deployedseed>417745701</deployedseed>
    <deployedseed>1914580239</deployedseed>
    <deployedseed>1462510675</deployedseed>
    <deployedseed>563880077</deployedseed>
    <deployedseed>745026925</deployedseed>
    <deployedseed>2027255467</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1b</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>(-17+8*u+4*v)/3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore></expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>