<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 37810 -->
  <question type="stack">
    <name>
				<text>1.1.6.3 Simplifying Multiplication</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Express as simply as possible</p>

<ol class="HELM_parts">
    <li>\(({@a1@}) \times x \times ({@a2@}) \times y =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</li>
    <li>\({@a3@} \times x \times z \times ({@a4@}) =\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<ol class="HELM_parts">
  <li>
    <p>Because of the commutativity of multiplication we have:</p>
    <p>\(\qquad ({@a1@}) \times x \times ({@a2@}) \times y = ({@a1@}) \times ({@a2@}) \times x \times y\)</p>
    <p>and then because of the associativity of multiplication we have:</p>
    <p>\(({@a1@}) \times ({@a2@}) \times x \times y = ({@a1@} \times {@a2@}) \times x \times y = ({@a1*a2@}) \times x \times y = {@a1*a2@}xy\)</p>
  </li>
  
  <li>
    <p>Because of the commutativity of multiplication we have:</p>
    <p>\(\qquad ({@a3@}) \times x \times z \times ({@a4@}) = ({@a3@}) \times ({@a4@}) \times x \times z\)</p>
    <p>and then because of the associativity of multiplication we have:</p>
    <p>\(({@a3@}) \times ({@a4@}) \times x \times z = ({@a3@} \times {@a4@}) \times x \times z = ({@a5@}) \times x \times z = {@a5@}xy\)</p>
  </li>
</ol>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021072200</text>
		</stackversion>
    <questionvariables>
				<text>a1: rand([-4,-3])
a2: rand([-2,-1])
a3: rand([8,9])
a4: rand([-5,-4])
a5: a3*a4

ta1: a1*a2*x*y
ta2: a3*a4*x*z</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text><ol class="HELM_parts_inline">
<li>\(({@a1@}) \times x \times ({@a2@}) \times y\)</li>
<li>\({@a3@} \times x \times z \times ({@a4@})\)</li>
</ol>
</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>5</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>5</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>EqualComAss</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210906172053+ytsTnu</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+211208155909+Y9BYAy</text></tag>
    </tags>
    <deployedseed>919213575</deployedseed>
    <deployedseed>402117183</deployedseed>
    <deployedseed>1342050492</deployedseed>
    <deployedseed>96803344</deployedseed>
    <deployedseed>300578462</deployedseed>
    <deployedseed>846537979</deployedseed>
    <deployedseed>1664176460</deployedseed>
    <deployedseed>1887597946</deployedseed>
    <deployedseed>1287301470</deployedseed>
    <deployedseed>488180774</deployedseed>
    <deployedseed>1472879930</deployedseed>
    <deployedseed>1442172767</deployedseed>
    <deployedseed>377184717</deployedseed>
    <deployedseed>1819974831</deployedseed>
    <deployedseed>1869491027</deployedseed>
    <deployedseed>367867111</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>x*y*ev(a1*a2,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>x*z*a5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>a1*a2*x*y</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>a3*a4*x*z</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>