<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24666 -->
  <question type="stack">
    <name>
				<text>18.1.1.4 Function of Three Variables</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>In a similar way we can define a function of three independent variables. Let these variables be \(x,y\) and \(u\) and the function \(f(x,y,u)\).</p>
<hr>
<h4 class="HELM_example">Example</h4>
<p>Given \(f(x,y,u)=x^2+yu+2\), find \(f(0,1,0),\ f(-1,-1,2)\).</p>

<h4 class="HELM_solution">Solution</h4>
\(f(0,1,0)=0^2+1\times 0+2=2\), &nbsp;\( f(-1,-1,2)=1-2+2=1\)]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201215135136+yFzXFw</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+201215135242+6dfq8u</text></tag>
    </tags>
    
</question>
</quiz>