<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24298 -->
  <question type="stack">
    <name>
				<text>5.1.1.10 Task</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Suppose that you travel often from Nottingham to Milton Keynes which is a distance of {@a@} miles almost all of which is along the M1 motorway. Usually it takes {@b@} hours. Suppose also that, on one occasion, you have agreed to pick someone up at the Leicester junction (21) of the M1. This is {@c@} miles from the start of your journey in Nottingham. If you start your journey at {@d@} a.m., what time should you advise for the pick-up? Note. Hours in first box and minutes in the second.</p>
[[input:ans1]] [[feedback:prt1]] 
[[input:ans2]] [[feedback:prt2]]
[[validation:ans1]][[validation:ans2]]]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>Let \(d\) miles be the distance travelled in time \(t\) hours. Then speed is equal to \( d/t = {@a@} / {@b@} t\).</p>
<p> Note: This is valid only for the duration of the journey \(0 \leq t \leq {@b@}\). </p>
<p>This equation can be used to find the time at which \(d = {@c@}\). </p>
<p>Now \({@c@} = {@h@} t\). So \(t = {@c@}/{@h@} = {@k@}\), this should then be converted into minutes {@ta2@}. </p>
<p>Due to the assumption of travelling at a constant speed for the entire journey and coming on and off the motorway at the services will take no time at all, it would be far safer to say I may arrive between \({@ta1@} {@l@}\) and \({@ta1@} {@m@}\).]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021120900</text>
		</stackversion>
    <questionvariables>
				<text>a: rand([85,86,87,88,89,90]);
b: rand([1.5,1.75,2,2.25]);
c: rand([23,24,25,26,27,28]);
d: rand([7,8,9]);
e: d+((b.c)/a);
ta1: floor(e);
g: (e-ta1)*60;
ta2: round(g);
h: a/b;
k: c/h;
l: ta2-10;
m: ta2+10;</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>{@ta1@}, {@ta2@}</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>ta2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200801093903+7Nw9aB</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220224103209+7AP7AN</text></tag>
    </tags>
    <deployedseed>1154166041</deployedseed>
    <deployedseed>1815543089</deployedseed>
    <deployedseed>1728537429</deployedseed>
    <deployedseed>149057603</deployedseed>
    <deployedseed>781424270</deployedseed>
    <deployedseed>1287377324</deployedseed>
    <deployedseed>1804357666</deployedseed>
    <deployedseed>968022755</deployedseed>
    <deployedseed>1949899233</deployedseed>
    <deployedseed>978976349</deployedseed>
    <deployedseed>2006091526</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>