<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24353 -->
  <question type="stack">
    <name>
				<text>16.1.1.2 Identify graph of sequence</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p><p>Which of the following graphs correctly shows the sequence \(\displaystyle 1,-1,1,-1,\ldots\) ?</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+201005121058+J3ziUP</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210610132045+G1WHgq</text></tag>
    </tags>
    
</question>
</quiz>