<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24621 -->
  <question type="stack">
    <name>
				<text>14.3.2.3 Find volume of parabola around y-axis</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>The curve \(y={@f@}\) for \({@low@}&lt;x&lt;{@upp@}\) is rotated about the \(y\)-axis. Find the volume of the solid formed.</p>
<p>\(\quad V =\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_solution">Solution</h4>
<p>Using the latest Key Point, the required integral is \[\int_{@ta2dn@}^{@ta2up@} {@ta2int@} \,\mathrm{d}y,\] since \(c={@ta2dn@}\) and \(d={@ta2up@}\). Since at \(x={@low@}\) we have \(y={@ta2dn@}\) and&nbsp;at \(x={@upp@}\) we have \(y={@ta2up@}\).</p>
<p>We have \(y={@f@}\), so the volume becomes \[\begin{align} \int_{@ta2dn@}^{@ta2up@} {@pi*x_sq@}\,\mathrm{d}y &amp;= \pi \left[ {@integrate(x_sq,y)@} \right]_{@ta2dn@}^{@ta2up@} \\ &amp;=\pi \left[ {@subst(y=ta2up,integrate(x_sq,y))@}-{@subst(y=ta2dn,integrate(x_sq,y))@}
    \right] \\ &amp;= {@ta3@} \end{align}\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>x_sq: y;
f: x^2;
low: rand(5)+1;
upp: low+1;

ta2up: upp^2;
ta2dn: low^2;
ta2int: pi*x^2;

ta3: simplify(integrate(pi*x_sq,y,ta2dn,ta2up));</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(y=\) {@f@}; \({@low@}\le x\le {@upp@}, \, V= {@ta3@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.</span>]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201202115615+4O2l3p</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210607115131+64VkgC</text></tag>
    </tags>
    <deployedseed>1774035727</deployedseed>
    <deployedseed>1576390224</deployedseed>
    <deployedseed>255718321</deployedseed>
    <deployedseed>702519725</deployedseed>
    <deployedseed>744715560</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans3</name>
        <value>ta3+1</value>
      </testinput>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>