<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27224 -->
  <question type="stack">
    <name>
				<text>35.3.3.2 Independent or Not?</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p></p><p>Decide which of the following pairs of events (\(A\) and \(B\)) arising from the experiments described are independent.<br></p><ol><li>One card is drawn from each of two packs.<br>

\(A=\) a red card is drawn from pack \(1\).<br>

\(B=\) a picture card is drawn from pack \(2\).<br>
\(A\) and \(B\) are [[1]]</li>
<li>The daily traffic accidents in Hull involving pedal cyclists and motorcyclists are counted.<br>

\(A=\) three motor cyclists are injured in separate collisions with cars.<br>

\(B=\) one pedal cyclist is injured when hit by a bus.<br>
\(A\) and \(B\) are [[1]]</li>

<li>Two boxes contains 20 nuts each, some have a metric thread, some
have a British Standard Fine (BSF) threads and some have a British
Standard Whitworth (BSW) thread. A nut is picked out of each box.<br>

\(A=\) a nut picked out of the first box is BSF.<br>

\(B=\) a nut picked out of the second box is metric.<br>
\(A\) and \(B\) are [[1]]</li>
<li>A box contains 20 nuts, some have a metric thread, some have a
British Standard Fine (BSF) threads and some have British Standard
Whitworth (BSW) thread. Two nuts are picked out of the box.<br>
\(A=\) first nut picked out of the box is BSF.<br>
\(B=\) second nut picked out of the box is metric.<br>
\(A\) and \(B\) are [[2]]</li></ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200701111655+wp4Qkr</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200914233418+rny4tf</text></tag>
    </tags>
    
</question>
</quiz>