<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27333 -->
  <question type="stack">
    <name>
				<text>38.1.1.3 Probability Density Functions</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
Which of the following are probability density functions?<p></p><p><br></p><p style="text-align: center;"><img src="@@PLUGINFILE@@/38.1%20figure%207%20%282%29.svg" alt="(i) is the line y=0 everywhere except at the line segment from (0,0) to (1,1). (ii) is the line y=0 everywhere except at the line segment from (0,1) to (2,0)." width="604" height="212" role="presentation" class="img-responsive atto_image_button_text-bottom"></p><p>\[\text{(iii)} \qquad f(x) = \left\{\begin{array}{ll} x^2 - 4x + \dfrac{10}{3}, &amp; \text{for all $0 \leq x \leq 3$}\\ 0, &amp; \text{elsewhere.} \end{array}\right. \]<br></p><p></p><p></p><p></p><p></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_solution">Solution</h4>
<p>(i) We can write this function as \[f(x) = \left\{\begin{array}{ll} x, &amp; \text{for all $0 \leq x \leq 1$} \\ 0, &amp; \text{elsewhere.} \end{array}\right. \] We have that \( f(x) \geq 0\) for all \(x\), but we have that \( \displaystyle{\int_{-\infty}^{\infty}} f(x) \; dx = \displaystyle{\int_0^1} x \; dx = \dfrac{1}{2}.\)&nbsp;Thus this function is not a valid probability density function because the integral’s value is not \(1\).</p><p>(ii) We can write this function as \[ f(x) = \left\{\begin{array}{ll} 1 -\dfrac{x}{2}, &amp; \text{for all $0 \leq x \leq 2$} \\ 0, &amp; \text{elsewhere.} \end{array}\right. \] We have that \( f(x) \geq 0\) for all \(x\), and we have that \( \displaystyle{\int_{-\infty}^{\infty}} f(x) \; dx = \int_0^2 1 -\dfrac{x}{2} \; dx = 1.\)&nbsp;Thus, this function is a valid probability density function.</p><p>(iii) We have that&nbsp;\( \displaystyle{\int_{-\infty}^{\infty}} f(x) \; dx = \int_0^3 x^2 - 4x + \dfrac{10}{3} \; dx = 1.\) But, we see that \(f(x)&lt;0 \) for \(\dfrac{3}{2} \leq x \leq \dfrac{5}{2}.\) Hence, because \(f(x)\) takes on negative values, it is not a valid probability density function.</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>1.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200722115639+WOqwW1</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+211013101108+mBGMwp</text></tag>
    </tags>
    
</question>
</quiz>