<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24834 -->
  <question type="stack">
    <name>
				<text>L7.2 Q2 Improper integral 1</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>Compute the improper integral. If it convergent evaluate it, otherwise enter the value <code>none</code> as your answer. \[ I = \int_{{@alaraja@}}^{{@ylaraja@}} \textrm{e}^{-x} {@b@} \, \textrm{d}x. \] \( I = \) [[input:SAns]][[validation:SAns]]
<p></p>
<p>Give your answer in exact form.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[We have the integral: \[ I = \int_{0}^{\infty} \textrm{e}^{-x} {{@b@}} \, \textrm{d}x.\] We can use integration by parts: \(\displaystyle\int_a^{b} u(x)v'(x) \, \textrm{d}x = [u(x)v(x)]_a^b -\int_{a}^{b} u'(x)v(x) \, \textrm{d}x\)<p></p>We have: \(u(x)=\textrm{e}^{-x} \Rightarrow u'(x)=-\textrm{e}^{-x}\) and \(v'(x)= {@b@} \Rightarrow v(x)= -\frac{1}{{@a@}} \cos({{@a@}}x)\)<p><br></p><p>So our integral becomse: \begin{eqnarray*} I = \int_{{@alaraja@}}^{\infty} \textrm{e}^{-x} {{@b@}} \, \textrm{d}x &amp; = &amp; \bigg/_{\!\!\!\!0}^{\,\infty} -\textrm{e}^{-x} \cdot \frac{1}{{@a@}} \cdot \cos({@a@}x)-\int_{0}^{\infty} \textrm{e}^{-x} \cdot \frac{1}{{@a@}}\cos({@a@}x) \, \textrm{d}x  \\
&amp; = &amp; \frac{1}{{@a@}}-\frac{1}{{@a@}} \int_{0}^{\infty} \textrm{e}^{-x} \cos({@a@}x) \, \textrm{d}x 
\end{eqnarray*} </p><p>Applying integration by parts again, we get:</p><p>\(u(x)=\textrm{e}^{-x} \Rightarrow u'(x)=-\textrm{e}^{-x} \) and </p><p>\(v'(x)= \cos({@a@}x) \Rightarrow v(x)=\frac{1}{{@a@}} \sin({@a@}x)\)</p><p>So our initial integral becomes:</p>

\begin{eqnarray*} I = \frac{1}{{@a@}}-\frac{1}{{@a@}} \int_{0}^{\infty} \textrm{e}^{-x} \cos({@a@}x) \, \textrm{d}x &amp; = &amp; \frac{1}{{@a@}}-\frac{1}{{@a@}} \bigg(\bigg/_{\!\!\!\!0}^{\,\infty} \frac{1}{{@a@}} \textrm{e}^{-x} \sin({@a@}x)-\int_{0}^{\infty} -\textrm{e}^{-x}\cdot \frac{1}{{@a@}} \sin({@a@}x) \, \textrm{d}x \bigg) \\
&amp; = &amp; \frac{1}{{@a@}}-\frac{1}{{@a@}} \bigg( 0+ \frac{1}{{@a@}} \int_{0}^{\infty} \textrm{e}^{-x} \sin({{@a@}}x) \, \textrm{d}x \bigg) \\
&amp; = &amp; \frac{1}{{@a@}}-\frac{1}{{@aa@}}\int_{0}^{\infty} \textrm{e}^{-x} \sin({@a@}x) \, \textrm{d}x.
\end{eqnarray*}

Whoa! The original integrand appeared also on the right hand side of the equation! Let us move it to the left side.
\begin{eqnarray*}
\int_{0}^{\infty} \textrm{e}^{-x}\sin({@a@}x) \, \textrm{d}x+\frac{1}{{@aa@}} \int_{0}^{\infty} \textrm{e}^{-x}\sin({@a@}x) \, \textrm{d}x &amp; = &amp; \frac{1}{{@a@}} \\
{@dd@} \int_{0}^{\infty} \textrm{e}^{-x} \sin({@a@}x) \, \textrm{d}x &amp; = &amp; \frac{1}{{@a@}} \\
\int_{0}^{\infty} \textrm{e}^{-x} \sin({@a@}x) \, \textrm{d}x &amp; = &amp; {@TAns@}.
\end{eqnarray*}<p></p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>1</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>a : rand([2,3,4,5,6])
b : sin(a*x)
alaraja : 0 
ylaraja : inf
f1 : e^(-x)*b
aa : a^2
cc : 1-1/aa
TAns : integrate(f1,x,alaraja,ylaraja)
dd : 1+1/aa</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>\(\int_{{@alaraja@}}^{{@ylaraja@}} e^{-x} {@b@} \, dx = {@TAns@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>SAns</name>
      <type>algebraic</type>
      <tans>TAns</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>[[basic-calculus]],[[basic-algebra]]</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>SAns</sans>
        <tans>TAns</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:maths-moodle.maths.ed.ac.uk+180222200854+XECG8a</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+210211004722+eap2Gu</text></tag>
    </tags>
    <deployedseed>1409120371</deployedseed>
    <deployedseed>1429096572</deployedseed>
    <deployedseed>1385118194</deployedseed>
    <deployedseed>2103560966</deployedseed>
    <deployedseed>951047825</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>SAns</name>
        <value>TAns</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>SAns</name>
        <value>TAns+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>