<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 33816 -->
  <question type="stack">
    <name>
				<text>13.5.3.4 Integrals and logarithms- continuing</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<hr>
<h4 class="HELM_example">Example</h4>
<p>Evaluate the definite integral:</p>
<p>\({\displaystyle\int \dfrac{x^2}{x^3 +1 }\ \mathrm{d}x}\)</p>
<h4 class="HELM_solution">Solution</h4>
<p>In this Example the derivative of the denominator is \(3x^2\) whereas the numerator is just \(x^2\). We adjust the numerator as follows:</p>
<p>\({\displaystyle\int \dfrac{x^2}{x^3 +1 }\ \mathrm{d}x} = {\displaystyle\dfrac{1}{3}\int \dfrac{3x^2}{x^3 +1 }\ \mathrm{d}x}\)</p>
<p>and integrate by the rule to get </p>
\(\dfrac{1}{3}\ln|x^3 + 1| + c\)
<p></p>
<p>Note that the sort of procedure in the last Example is only possible because we can move <strong>constant factors</strong> through the integral sign. It would be wrong to try to move terms involving the <strong>variable</strong> \(x\) in a similar way.</p>
<p></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210628144009+biCdhl</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210628144557+J9Dmnb</text></tag>
    </tags>
    
</question>
</quiz>