<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 84666 -->
  <question type="stack">
    <name>
				<text>18.2.1.4 First partial derivatives - functions of several var</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p></p>
<p><strong>Functions of several variables</strong></p>
<p>As we have seen, a function of two variables \(f(x,y)\) has two partial derivatives, \(\dfrac{\partial f}{\partial x}\) and \(\dfrac{\partial f}{\partial y}\). In an exactly analogous way a function of three variables \(f(x,y,u)\) has three partial derivatives
    \(\dfrac{\partial f}{\partial x}\), \(\dfrac{\partial f}{\partial y}\) and \(\dfrac{\partial f}{\partial u}\), and so on for functions of more than three variables. Each partial derivative is obtained in the same way as stated in the following Key
    Point:
</p>

<div class="HELM_keypoint">
    <h4>Key Point</h4>
    <p><strong>The Partial Derivatives of \(f(x,y,u,v,w,\dots)\)</strong></p>
    <p>For a function of several variables \(z=f(x, y, u, v, w, \dots)\) the partial derivative of \(f\) with respect to \(v\) (say) is denoted by \(\dfrac{\partial f}{\partial v}\) and is obtained by differentiating \(f(x, y, u, v, w, \dots)\) with respect
        to \(v\) in the usual way but treating all the other variables as if they were constants.</p>
    <p>Alternative notations for \(\dfrac{\partial f}{\partial v}\) when \(z = f(x, y, u, v, w, \dots )\) are \(f_v(x, y, u, v, w \dots )\) and \(f_v\) and \(\dfrac{\partial z}{\partial v}\).</p>
</div>

<hr>
<h4 class="HELM_example">Example</h4>
<p>Find \(\dfrac{\partial f}{\partial x}\) and \(\dfrac{\partial f}{\partial u}\) for \(f(x, y, u, v) = x^2+xy^2+y^2u^3-7uv^4\)</p>

<h4 class="HELM_solution">Solution</h4>
<p>\(\dfrac{\partial f}{\partial x}=2x+y^2+0+0=2x+y^2\),</p>
<p>\(\dfrac{\partial f}{\partial u}=0+0+y^2\cdot 3u^2-7v^4=3y^2u^2-7v^4\)</p>
<hr>

<h4 class="HELM_example">Example</h4>
<p>The pressure, \(P\), for one mole of an ideal gas is related to its absolute temperature, \(T\), and specific volume, \(v\), by the equation \(Pv=RT\) where \(R\) is the gas constant. Obtain simple expressions for</p>
<ol>
    <li> the coefficient of thermal expansion, \(\alpha\), defined by: \(\alpha = \dfrac{1}{v}\left(\dfrac{\partial v}{\partial T}\right)_P\)</li><br>
    <li> the isothermal compressibility, \(\kappa_T\), defined by: \(\kappa_T=-\dfrac{1}{v}\left(\dfrac{\partial v}{\partial P}\right)_T\)</li>
</ol>

<h4 class="HELM_solution">Solution</h4>
<p>
    \(v=\dfrac{RT}{P}\quad\Rightarrow\quad \left(\dfrac{\partial v}{\partial T}\right)_P=\dfrac{R}{P}\)
</p>
<p>so \(\alpha=\dfrac{1}{v}\left(\dfrac{\partial v}{\partial T}\right)_P=\dfrac{R}{Pv}=\dfrac{1}{T}\)
</p><br>
<p>
    \(v=\dfrac{RT}{P}\quad\Rightarrow\quad \left(\dfrac{\partial v}{\partial P}\right)_T=-\dfrac{RT}{P^2}\)
</p>
<p>so \(\kappa_T=-\dfrac{1}{v}\left(\dfrac{\partial v}{\partial P}\right)_T=\dfrac{RT}{vP^2}=\dfrac{1}{P}\)
</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+220304105218+wgCWAA</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220304113954+5cRbEt</text></tag>
    </tags>
    
</question>
</quiz>