<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27250 -->
  <question type="stack">
    <name>
				<text>36.1.4.2 Graphical representations/ Pie charts</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Pie charts</h4>
<p>Pie charts are often seen in magazines and newspapers. A pie chart is simply a circular diagram where the circle is divided into sectors and the angles of the sectors are proportional to the quantity represented. The quantities may be the frequencies of values of a categorical variable in a sample. Since the total area of the circle is fixed, pie charts are considered to be useful for representing proportions of a total.</p>
<p>The following data represent the time spent weekly on a variety of activities by the
full-time employees of an engineering company, based on their&nbsp;assigned female/male gender at birth (AFAB/AMAB).</p>
<table class="HELM_table" style="width:75%">
<thead>
<tr><th>Hours spent on:</th><th>AMAB</th><th>AFAB</th></tr>
</thead><tbody>
<tr><td>Travel to and from work</td><td>10.5</td><td>8.4</td></tr>
<tr><td>Paid activities in employment</td><td>47.0</td><td>37.0</td></tr>
<tr><td>Personal sport and leisure activities</td><td>8.2</td><td>3.6</td></tr>
<tr><td>Personal development</td><td>5.6</td><td>6.4</td></tr>
<tr><td>Family activities</td><td>8.4</td><td>18.2</td></tr>
<tr><td>Sleep</td><td>56.0</td><td>56.0</td></tr>
<tr><td>Other</td><td>32.3</td><td>38.4</td></tr>
</tbody>
</table><br><br>
<p>To construct a pie chart showing how the AMAB employees spend their time we proceed as follows. Note that the total number of hours spent is 168 (\(7\times 24\)).
</p><ol>
<li>Express the time spent on any given activity as a proportion of the total time spent;</li>
<li>Multiply the number obtained by 360 thus converting the proportion to an angle, in degrees;</li>
<li>Draw a chart consisting of (in this case) 6 sectors having the angles given in the table below  subtended at the centre of the circle.</li>
</ol>
<table class="HELM_table" style="text-align: center; width: 100%;">
<thead>
<tr style="border-bottom: 1px solid #000;"><th>Hours spent on:</th><th>AMAB</th><th>Proportion of Time</th><th>Sector Angle<br><span style="">(rounded to 1 d.p.)</span></th></tr></thead><tbody>
<tr><td>Travel to and from work</td><td>\(10.5\)</td><td>\(\dfrac{10.5}{168}\)</td><td>\(\dfrac{10.5}{168}\times 360^{\circ}=22.5^{\circ}\)</td></tr>
<tr><td>Paid activities in employment</td><td>\(47.0\)</td><td>\(\dfrac{47}{168}\)</td><td>\(\dfrac{47}{168}\times 360^{\circ}=100.7^{\circ}\)</td></tr>
<tr><td>Personal sport and leisure activities</td><td>\(8.2\)</td><td>\(\dfrac{8.2}{168}\)</td><td>\(\dfrac{8.2}{168}\times 360^{\circ}=17.6^{\circ}\)</td></tr>
<tr><td>Personal development</td><td>\(5.6\)</td><td>\(\dfrac{5.6}{168}\)</td><td>\(\dfrac{5.6}{168}\times 360^{\circ}=12.0^{\circ}\)</td></tr>
<tr><td>Family activities</td><td>\(8.4\)</td><td>\(\dfrac{8.4}{168}\)</td><td>\(\dfrac{8.4}{168}\times 360^{\circ}=18.0^{\circ}\)</td></tr>
<tr><td>Sleep</td><td>\(56.0\)</td><td>\(\dfrac{56}{168}\)</td><td>\(\dfrac{56}{168}\times 360^{\circ}=120.0^{\circ}\)</td></tr>
<tr><td>Other</td><td>\(32.3\)</td><td>\(\dfrac{32.3}{168}\)</td><td>\(\dfrac{32.3}{168}\times 360^{\circ}=69.2^{\circ}\)</td></tr>
</tbody>
</table><br>	
<p>The pie chart obtained is illustrated below.</p><br><p style="text-align: center;"><img src="@@PLUGINFILE@@/36-1-fig-2-pie_chart.svg" alt="" width="324" height="248" class="img-responsive atto_image_button_text-bottom"><br></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200708020508+eJhB0W</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+211007131147+H8uxza</text></tag>
    </tags>
    
</question>
</quiz>