<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24304 -->
  <question type="stack">
    <name>
				<text>5.3.1.2 Amplitude, frequency and period</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p><ol class="HELM_parts"><p>The figures below show plots of the functions</p><p>\((\text{a})\,\,y=\sin x,\quad (\text{b})\,\,y=2\sin x,\quad (\text{c})\,\,y=\cos x,\quad\text{and}\quad (\text{d})\,\,y=\cos\dfrac{x}{2}\)&nbsp;<br><br></p><p><img src="@@PLUGINFILE@@/5-3-sin_plots.svg" alt="" width="417" height="292" role="presentation" class="img-responsive atto_image_button_text-bottom"><br></p><p><img src="@@PLUGINFILE@@/5-3-cos_plots.svg" alt="" width="416" height="169" role="presentation" class="img-responsive atto_image_button_text-bottom"><br><br>State the correct amplitude, frequency, and period for the four functions:

<br><br></p><p></p><li>The amplitude is [[1]], the frequency is [[5]]\(\pi\), and the period if [[9]]\(\pi\).<p></p><p></p></li><li>The amplitude is [[1]], the frequency is [[5]]\(\pi\), and the period if [[9]]\(\pi\).</li><li>The amplitude is [[2]], the frequency is [[5]]\(\pi\), and the period if [[9]]\(\pi\).<br></li><li>The amplitude is [[1]], the frequency is [[6]]\(\pi\), and the period if [[10]]\(\pi\).<br></li></ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>In general, functions of the form</p><p>\(y=a\sin(bx),\quad\text{and}\quad y=a\cos(bx)\)</p><p>Have amplitude equal to \(a\), frequency equal to \(\dfrac{b}{2\pi}\), and period equal to \(\dfrac{2\pi}{b}\).</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200821075715+qXDu8t</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200821080336+wZlWXh</text></tag>
    </tags>
    
</question>
</quiz>