<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26726 -->
  <question type="stack">
    <name>
				<text>10.1.1.6 What is a Complex Number?</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<div class="HELM_keypoint">
<h4>Key Point </h4>
<p>Two complex numbers <span class="math inline">\(z=a+i b\)</span> and <span class="math inline">\(w=c+i d\), where \(a,b,c,d\) are real,</span> are said to be <span><strong>equal</strong></span> if and only if both their real parts are the same and both their imaginary parts are the same, that is <span class="math display">\[a=c\qquad \text{and}\qquad b=d.\]</span></p>
</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200625155607+zxBQSH</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210706164825+UAY9QS</text></tag>
    </tags>
    
</question>
</quiz>