<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27357 -->
  <question type="stack">
    <name>
				<text>39.1.1.1 Introduction and the CLT</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>The Normal Distribution</h3>
<p>The normal distribution is the most widely used model for the distribution of a random variable. There is a very good reason for this. Practical experiments involve measurements and measurements involve errors. However you go about measuring a quantity, inaccuracies of all sorts can make themselves felt. For example, if you are measuring a length using a device as crude as a ruler, you may find errors arising due to: </p>
<ul>
<li> the calibration of the ruler itself;</li>
<li>parallax errors due to the relative positions of the object being measured, the ruler and your eye;</li>
<li>rounding errors;</li>
<li>‘guesstimation’ errors if a measurement is between two marked lengths on the ruler;</li>
<li> mistakes.</li></ul>
<p> If you use a meter with a digital readout, you will avoid some of the above errors but others, often present in the design of the electronics controlling the meter, will be present. Errors are unavoidable and are usually the sum of several factors. The behaviour of variables which are the sum of several other variables is described by a very important and powerful result called <b>the Central Limit Theorem</b> which we will study later in this Workbook. For now we will quote the result so that the importance of the normal distribution will be appreciated. </p>

<h4>The Central Limit Theorem</h4>
<p>Let \(X\) be the sum of \(n\) independent random variables \(X_i, i=1,2,\dots n\) each having a distribution with mean \(\mu_i\) and variance \(\sigma^2_i\quad (\sigma_i^2&lt;\infty)\), respectively, then the distribution of \(X\) has expectation and variance given by the expressions </p>
<p> \(\displaystyle{\qquad\text{E}(X)=\sum_{i=1}^n\mu_i\qquad\text{and}\qquad \text{V}(X)=\sum_{i=1}^n\sigma^2_i}\) </p>
<p> and becomes <b>normal</b> as \(n\rightarrow\infty\). </p>
<p> Essentially we are saying that a quantity which represents the combined effect of a number of variables will be approximately normal no matter what the original distributions are provided that \(\sigma^2&lt;\infty\). This statement is true for the vast majority of distributions you are likely to meet in practice. This is why the normal distribution is crucially important to engineers. A quotation attributed to rof. G. Lippmann, (1845-1921, winner of the Nobel prize for hysics in 1908) ‘Everybody believes on the law of errors, experimenters because they think it is a mathematical theorem and mathematicians because they think it is an experimental fact.’ </p> 
<p>You may think that anything you measure follows an approximate normal distribution. Unfortunately this is not the case. While the heights of human beings follow a normal distribution, weights do not. Heights are the result of the interaction of many factors (outside one’s control) while weights principally depend on lifestyle (including how much and what you eat and drink!) In practice, it is found that weight is skewed to the right but that the square root of human weights is approximately normal. </p>


<h4>Defining a Normal Distribution</h4>
<p>The probability density function of a normal distribution with mean \(\mu\) and variance \(\sigma^2\) is given by the formula \[f(x)=\dfrac{1}{\sigma\sqrt{2\pi}}e^{-(x-\mu)^2/2\sigma^2}\] This curve is always bell-shaped with the centre of the bell located at the value of \(\mu\). See Figure below. The height of the bell is controlled by the value of \(\sigma\). As with all normal distribution curves it is symmetrical about the centre and decays as \(x\to \pm \infty\). As with any probability density function the area under the curve is equal to 1.</p> 
<p></p><center><img src="@@PLUGINFILE@@/39-1-fig-1-normaldist.svg" alt="The Normal Probability Distribution Function" class="img-responsive atto_image_button_text-bottom" width="332" height="75"><br></center><p></p> 
<p></p><center><strong>Figure</strong>: Normal Probability Density Function (PDF)</center><p></p> 
<p>A normal distribution is completely defined by specifying its mean (the value of \(\mu\)) and its variance (the value of \(\sigma^2\).) The normal distribution with mean \(\mu\) and variance \(\sigma^2\) is written \(N(\mu,\sigma^2)\). Hence the distribution \(N(20,25)\) has a mean of \(20\) and a standard deviation of \(5\); remember that the second parameter is the variance which is the square of the standard deviation. </p>
<div class="HELM_keypoint">
<h4>Key Point</h4>
<p> A normal distribution has mean \(\mu\) and variance \(\sigma^2\). A random variable \(X\) following this distribution is usually denoted by \(N(\mu, \sigma^2)\) and we often write
\[X ∼ N(\mu, \sigma^2)\] </p>
</div>
<p> </p>
<p>Clearly, since \(\mu\) and \(\sigma^2\) can both vary, there are infinitely many normal distributions and it is impossible to give tabulated information concerning them all. 
</p><p> For example, if we produce piston rings which have a target mean internal diameter of \(45\) mm then we may realistically expect the actual diameter to deviate from this value. Such deviations are well-modelled by the normal distribution. Suppose we decide that diameters in the range \(44.95\) mm to \(45.05\) mm are acceptable, we may then ask the question ‘What proportion of our manufactured output is satisfactory?’ </p> 
<p>Without tabulated data concerning the appropriate normal distribution we cannot easily answer this question (because the integral used to calculate areas under the normal curve is intractable.)</p>
<p>Since tabulated data allow us to apply the distribution to a wide variety of statistical situations, and we cannot tabulate all normal distributions, we tabulate only one - the standard normal distribution - and convert all problems involving the normal distribution into problems involving the standard normal distribution.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200727115245+YQ92pe</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200904105413+4fUPjH</text></tag>
    </tags>
    
</question>
</quiz>