<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24684 -->
  <question type="stack">
    <name>
				<text>NOT IN QUIZ 19.1.3.3 Classify the differential equation</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p dir="ltr">&nbsp;Classify the differential equation \[ \dfrac{d^2x}{ dt^2}+3 \dfrac{dx}{dt}=x.\]&nbsp;</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p dir="ltr" style="text-align: left;">The highest order of the derivatives is two and there is no non-linear term.</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:moodle2.maths.ed.ac.uk+210315125302+51dMez</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+210319153855+B4o9yL</text></tag>
    </tags>
    
</question>
</quiz>