<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24354 -->
  <question type="stack">
    <name>
				<text>16.1.1.3 Identify convergence</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p><p> Is the sequence \(\displaystyle 1,-1,1,-1,\ldots\) convergent or divergent?</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+201005133904+dlITVq</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+201130143846+S7uDW6</text></tag>
    </tags>
    
</question>
</quiz>