<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26602 -->
  <question type="stack">
    <name>
				<text>8.3.3.2 Engineering Example_Currents in three loops</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_example">Engineering Example: Currents in three loops</h4>

<p>In the circuit shown find the currents <span class="math inline">\((i_1, i_2, i_3)\)</span> in the loops.</p>

<figure class="HELM_figure"><img src="@@PLUGINFILE@@/8-3-fig-1-currents.svg" alt="" role="presentation" width="267" height="199"></figure>

<h4 class="HELM_solution">Solution</h4>

<p>Loop 1 gives <span class="math display">\[2(i_1)+3(i_1-i_2)=5\quad\rightarrow\quad 5i_1-3i_2=5\]</span> Loop 2 gives <span class="math display">\[6(i_2-i_3)+3(i_2-i_1)=4\quad\rightarrow\quad -3i_1+9i_2-6i_3=4\]</span> Loop 3 gives <span class="math display">\[6(i_3-i_2)+4(i_3)=6-5\quad\rightarrow\quad -6i_2+10i_3=1\]</span>    Note that in loop 3, the current generated by the <span class="math inline">\(6v\)</span> cell is positive and for the <span class="math inline">\(5v\)</span> cell negative in the direction of the arrow.</p>
<p>In matrix form <span class="math display">\[\left[\begin{matrix}\hfill 5&amp;\hfill -3&amp;\hfill 0\\
\hfill -3&amp;\hfill 9&amp;\hfill -6\\
\hfill 0&amp;\hfill -6&amp;\hfill 10\end{matrix}\right]
\left[\begin{matrix}\hfill i_1\\
\hfill i_2\\
\hfill i_3\end{matrix}\right]=
\left[\begin{matrix}\hfill 5\\
\hfill 4\\
\hfill 1\end{matrix}\right]\]</span> Solving gives <span class="math display">\[i_1= \frac{34}{15},\quad i_2= \frac{19}{9},\quad i_3= \frac{41}{30}\]</span></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200630144745+tvaSJZ</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210630142641+FUgFkn</text></tag>
    </tags>
    
</question>
</quiz>