<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 37777 -->
  <question type="stack">
    <name>
				<text>1.1.6.1 Combining Numbers Together Using +,−,×,÷</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span style="color: inherit; font-family: inherit;">Combining Numbers Together Using \(+,-,\times,\div\)</span><br></h3>
<h5><span style="color: inherit; font-family: inherit;">Addition (\(+\))</span><br></h5>
<p>If the letters \(x\) and \(y\) represent two numbers, then their <strong>sum</strong> is written as \(x+y\). Note that \(x+y\) is the same as \(y+x\) just as \(4+7\) is equal to \(7+4\).</p>
<h5><span style="color: inherit; font-family: inherit;">Subtraction (\(-\))</span><br></h5>
<p>Subtracting \(y\) from \(x\) yields \(x-y\). Note that \(x-y\) is not the same as \(y-x\) just as \(11-7\) is not the same as \(7-11\), however in both cases the difference is said to be 4.</p>
<h5><span style="color: inherit; font-family: inherit;">Multiplication (\(\times\))</span><br></h5>
<p>The instruction to multiply \(x\) and \(y\) together is written as \(x\times y\). Usually the multiplication sign is omitted and we write simply \(xy\). An alternative notation is to use a dot to represent multiplication and so we could write \(x.y\) The quantity \(xy\) is called the <strong>product</strong> of \(x\) and \(y\). As discussed earlier multiplication is both commutative and associative:</p>
<p>\(\qquad \text{i.e.}\qquad x\times y=y\times x\qquad \text{and}\qquad (x\times y)\times z = x\times (y\times z)\)</p>
<p>This last expression can thus be written \(x\times y\times z\) without ambiguity. When mixing numbers and symbols it is usual to write the numbers first. Thus \(3\times x\times y\times 4 = 3\times 4\times x\times y = 12xy\).</p>
<hr>
<h4 class="HELM_example">Example</h4>
<p>Simplify </p>
<ol class="HELM_parts_inline">
<li>\(9(2y)\)</li>
<li>\(-3(5z)\)</li>
<li>\(4(2a)\)</li>
<li>\(2x \times (2y)\)</li>
</ol>
<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
<li>Note that \(9(2y)\) means \(9\times(2\times y)\). Because of the associativity of multiplication \(9\times (2\times y)\) means the same as \((9 \times 2) \times y\), that is \(18y\).</li>
<li>\(−3(5z)\) means \(−3 \times (5 \times z)\). Because of associativity this is the same as \((−3 \times 5) \times z\), that is \(−15z\).</li>
<li>\(4(2a)\) means \(4 \times (2 \times a)\). We can write this as \((4 \times 2) \times a\), that is \(8a\).</li>
<li>Because of the associativity of multiplication, the brackets are not needed and we can write \(2x \times (2y) = 2x \times 2y\) which equals
<p>\(2 \times x \times 2 \times y = 2 \times 2 \times x \times y = 4xy\).</p></li></ol>
<hr>
<h4 class="HELM_example">Example</h4>
<p>What is the distinction between \(9(−2y)\) and \(9 − 2y\)?</p>
<h4 class="HELM_solution">Solution</h4>
<p>The expression \(9(−2y)\) means \(9 \times (−2y)\). Because of associativity of multiplication we can write this as \(9 \times (−2) \times y\) which equals \(-18y\).</p>
<p>On the other hand \(9 − 2y\) means subtract \(2y\) from \(9\). This cannot be simplified.</p>
<hr>
<h5><span style="color: inherit; font-family: inherit;">Division (\(\div\))</span><br></h5>
<p>The quantity \(x\div y\) means \(x\) divided by \(y\). This is also written as \(x/y\) or \(\dfrac{x}{y}\) and is known as the <strong>quotient</strong> of \(x\) and \(y\). In the expression \(\dfrac{x}{y}\) the symbol \(x\) is called the <strong>numerator</strong> and the symbol \(y\) is called the <strong>denominator</strong>. Note that \(x/y\) is not the same as \(y/x\). Division by 1 leaves a quantity unchanged so that \(\dfrac{x}{1}\) is simply \(x\).</p>
<h5><span style="color: inherit; font-family: inherit;">Algebraic Expressions</span><br></h5>
<p>A quantity made up of symbols and the operations \(+\), \(-\), \(\times\) and \(/\) is called an <strong>algebraic expression</strong>. One algebraic expression divided by another is called an algebraic fraction. Thus \(\dfrac{x+7}{x-3} \hbox{ and } \dfrac{3x-y}{2x+z}\) are algebraic fractions. The <strong>reciprocal</strong> of an algebraic fraction is found by inverting it. Thus the reciprocal of \(\dfrac{2}{x}\) is \(\dfrac{x}{2}\). The reciprocal of \(\dfrac{x+7}{x-3}\) is \(\dfrac{x-3}{x+7}\).</p>
<hr>
<h4 class="HELM_example">Example</h4>
<p>State the reciprocal of each of the following expressions:</p>
<ol class="HELM_parts_inline">
<li>\(\dfrac{y}{z}\)</li>
<li>\(\dfrac{x+z}{a-b}\)</li>
<li>\(3y\)</li>
<li>\(\dfrac{1}{ a+2b}\)</li>
<li> \(-\dfrac{1}{y}\)</li>
</ol>
<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
<li>\(\dfrac{z}{y}\)</li>
<li>\(\dfrac{a-b}{x+z}\)</li>
<li>\(3y\) is the same as \(\dfrac{3y}{1}\) so the reciprocal of \(3y\) is \(\dfrac{1}{3y}\).</li>
<li>The reciprocal of \(\dfrac{1}{a + 2b}\) is \(\dfrac{a + 2b}{1}\) or simply \(a + 2b\).</li>
<li>The reciprocal of \(-\dfrac{1}{y}\) is \(-\dfrac{y}{1}\) or simply \(−y\).</li>
</ol>
<hr>
<p>Finding the reciprocal of complicated expressions can cause confusion. Study the following Example carefully.</p>
<hr>
<h4 class="HELM_example">Example</h4>
<p>Obtain the reciprocal of:</p>
<ol class="HELM_parts_inline">
<li>\(p+q\)</li>
<li>\(\dfrac{1}{R_1} + \dfrac{1}{R_2}\)</li>
</ol>
<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
<li>Because \(p + q\) can be thought of as \(\dfrac{p + q}{1}\) its reciprocal is \(\dfrac{1}{p + q}\). Note in particular that the reciprocal of \(p + q\) is not \(\dfrac{1}{p} + \dfrac{1}{q}\). This distinction is important and a common cause of error. To avoid an error carefully identify the numerator and denominator in the original expression before inverting.</li>
<li>The reciprocal of \(\dfrac{1}{R_1} + \dfrac{1}{R_2}\) is \(\dfrac{1}{\dfrac{1}{R_1} + \dfrac{1}{R_2}}\). To simplify this further requires knowledge of the addition of algebraic fractions which is dealt with in later workbooks. It is important to note that the reciprocal of \(\dfrac{1}{R_1} + \dfrac{1}{R_2}\) is not \(R_1 + R_2\).</li>
</ol>
<hr>
<h5><span style="color: inherit; font-family: inherit;">The equals sign (\(=\))</span><br></h5>
<p>The equals sign, \(=\), is used in several different ways. Firstly, an equals sign is used in <strong>equations</strong>. The left-hand side and right-hand side of an equation are equal only when the variable involved takes specific values known as <strong>solutions</strong> of the equation. For example, in the equation \(x-8=0\), the variable is \(x\). The left-hand side and right-hand side are only equal when \(x\) has the value \(8\). If \(x\) has any other value the two sides are not equal.</p>
<p>Secondly, the equals sign is used in <strong>formulae</strong>. Physical quantities are often related through a formula. For example, the formula for the length, \(C\), of the circumference of a circle expresses the relationship between the circumference of the circle and its radius, \(r\). This formula states \(C=2\pi r\). When used in this way the equals sign expresses the fact that the quantity on the left is found by evaluating the expression on the right.</p>
<p>Thirdly, an equals sign is used in <strong>identities</strong>. An identity looks just like an equation, but it is true for <em>all</em> values of the variable. We shall see shortly that \((x-1)(x+1)=x^2-1\) for any value of \(x\) whatsoever. This mean that the quantity on the left means exactly the same as that on the right whatever the value of \(x\). To distinguish this usage from other uses of the equals symbol it is more correct to write \((x-1)(x+1)\equiv x^2-1\), where \(\equiv\) means ‘is identically equal to’. However, in practice, the equals sign is often used. We will only use \(\equiv\) where it is particularly important to do so.</p><p><br></p>
<h5><span style="color: inherit; font-family: inherit;">The ‘not equals’ sign (\(\ne\))</span><br></h5>
<p>The sign \(\ne\) means ‘is not equal to’. For example, \(5\ne 6\), \(7\ne -7\).</p><p><br></p>
<h5><span style="color: inherit; font-family: inherit;">The notation for the change in a variable (\(\delta\))</span><br></h5>
<p>The <strong>change</strong> in the value of a quantity is found by subtracting its initial value from its final value. For example, if the temperature of a mixture is initially \(13^\circ\)C and at a later time is found to be \(17^\circ\)C, the change in temperature is \(17-13=4^\circ\)C. The Greek letter \(\delta\) is often used to indicate such a change. If \(x\) is a variable we write \(\delta x\) to stand for a change in the value of \(x\). We sometimes refer to \(\delta x\) as an <strong>increment</strong> in \(x\). For example if the value of \(x\) changes from 3 to 3.01 we could write \(\delta x= 3.01-3 = 0.01\). It is important to note that this is <strong>not</strong> the product of \(\delta\) and \(x\), rather the whole symbol ‘\(\delta x\)’ means ‘the increment in \(x\)’.</p><p><br></p>
<h5><span style="color: inherit; font-family: inherit;">Sigma (or summation) notation (\(\sum\))</span><br></h5>
<p>This provides a concise and convenient way of writing long sums.</p>
<p>The sum</p>
<p>\(\qquad x_1+x_2+x_3+x_4 +\ldots +x_{11}+x_{12}\)</p>
<p>is written using the capital Greek letter sigma, \(\sum\), as</p>
<p>\(\qquad \sum_{k=1}^{12} x_k\)</p>
<p>The symbol \(\sum\) stands for the sum of all the values of \(x_k\) as \(k\) ranges from 1 to 12. Note that the lower-most and upper-most values of \(k\) are written at the bottom and top of the sigma sign respectively.</p> 
<hr>
<h4 class="HELM_example">Example</h4>
<p>Write out explicitly what is meant by \(\sum_{k=1}^{5}k^3\)</p>
<h4 class="HELM_solution">Solution</h4>
<p>We must let k range from 1 to 5:</p>
<p>\(\qquad \sum_{k=1}^{5}k^3 = 1^3+2^3+3^3+4^3+5^3\)</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+210903161420+EtjKEa</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+211208145312+GOFyic</text></tag>
    </tags>
    
</question>
</quiz>