<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26230 -->
  <question type="stack">
    <name>
				<text>7.2.6.5 Communication network example</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_example">Example</h4>
<h4>Communication network</h4>
<h5>Problem in words</h5>
<p>The figure below represents a communication network. Vertices \(a, b, f\) and \(g\) represent offices. Vertices \(c, d\) and \(e\) represent switching centres. The numbers marked along the edges represent the number of connections between any two vertices. Calculate the number of routes from \(a\) and \(b\) to \(f\) and \(g.\)
</p>
<figure class="HELM_figure">
<img src="@@PLUGINFILE@@/7-2-fig-3-communications-network-fixed.svg" alt="A network with 7 vertices, separated into 3 columns. Column 1: a, b. Column 2: c, d, e. Column 3: f, g. In each column, each vertex has edges joining it to each vertex of the next column. Each edge is weighted with an integer, representing the number of connection between the endpoint vertices. Here are the weights: edge a-c has weight 3; edge a-d has weight 4; edge a-e has weight 1; edge b-c has weight 2; edge b-d has weight 1; edge b-e has weight 3; edge c-f has weight 2; edge c-g has weight 1; edge d-f has weight 6; edge d-g has weight 3; edge e-f has weight 1; edge e-g has weight 2. This completes the description." width="180" height="168" class="img-responsive atto_image_button_text-bottom">
<figcaption>A communication network where \(a\), \(b\), \(f\) and \(g\) are offices and \(c\), \(d\) and \(e\) are switching centres</figcaption>
</figure>
<h5>Mathematical statement of the problem</h5>
<p>The number of routes from \(a\) to \(f\) can be calculated by taking the number via \(c\) plus the number via \(d\) plus the number via \(e\). In each case this is given by multiplying the number of connections along the edges connecting \(a\) to \(c\), \(c\) to \(f\) etc. This gives the result: \[\text{Number of routes from }a\text{ to }f = 3\times2 + 4\times 6 + 1\times 1=31.\] The nature of matrix multiplication means that the number of routes is obtained by multiplying the matrix representing the number of connections from \(ab\) to \(cde\) by the matrix representing the number of connections from
\(cde\) to \(fg\).</p>
<h5>Mathematical analysis</h5>
<p>The matrix representing the number of routes from \(ab\) to \(cde\) is:
\[ \begin{matrix} &amp; \begin{matrix} c &amp; d &amp; e \end{matrix}\\ \begin{matrix} a \\ b\end{matrix}\!\!\!\!\! &amp; \begin{bmatrix} 3 &amp; 4 &amp; 1 \\ 2 &amp; 1 &amp; 3 \end{bmatrix}\end{matrix}\] The matrix representing the number of routes from \(cde\) to \(fg\) is:
\[\begin{matrix} &amp; \begin{matrix} f &amp; g \end{matrix} \\ \begin{matrix} c \\ d \\ e \end{matrix}\!\!\!\!\! &amp; \begin{bmatrix} 2 &amp; 1 \\ 6 &amp; 3 \\ 1 &amp; 2 \end{bmatrix} \end{matrix}\] The product of these two matrices gives the total number of routes: \begin{align*} \begin{bmatrix} 3 &amp; 4 &amp; 1 \\ 2 &amp; 1 &amp; 3  \end{bmatrix}\begin{bmatrix} 2 &amp; 1 \\ 6 &amp; 3 \\ 1 &amp; 2 \end{bmatrix} &amp; = \begin{bmatrix} 3\times 2+4\times 6+1\times1 &amp; 3\times 1+4\times 3+1\times 2 \\ 2\times2+1\times 6+3\times 1 &amp; 2\times 1+1\times 3+3\times 2 \end{bmatrix} \\ &amp; = \begin{bmatrix} 31 &amp; 17 \\ 13 &amp; 11 \end{bmatrix} \end{align*}</p>
<h5>Interpretation</h5>
We can interpret the resulting (product) matrix by labelling the columns and rows: \[\begin{matrix} &amp; \begin{matrix} f &amp; g \end{matrix} \\ \begin{matrix} a \\ b \end{matrix}\!\!\!\!\! &amp; \begin{bmatrix} 31 &amp; 17 \\ 13 &amp; 11 \end{bmatrix} \end{matrix}\]
Hence there are \(31\) routes from \(a\) to \(f\), \(17\) from \(a\) to \(g\), \(13\) from \(b\) to \(f\) and \(11\) from \(b\) to \(g\).]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200626112322+pLRODT</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210701151504+NDKEzC</text></tag>
    </tags>
    
</question>
</quiz>