<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24281 -->
  <question type="stack">
    <name>
				<text>16.5.1.1 Maclaurin and Taylor series</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span>Maclaurin and Taylor series</span></h3><br>
<p>As we shall see, many functions can be represented by power series. In fact we have already seen in earlier Sections examples of such a representation: <span class="math display">\[\begin{aligned}
\dfrac{1}{1-x}&amp;=1+x+x^2+\cdots\qquad |x| &lt;1\\
\ln (1+x)&amp;=x-\dfrac{x^2}{2}+\dfrac{x^3}{3}-\cdots \qquad -1 &lt; x \leq 1\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\\
e^x&amp;=1+x+\dfrac{x^2}{2!}+\dfrac{x^3}{3!}+\cdots\qquad \text{all}\ x\end{aligned}\]</span> The first two examples show that, as long as we constrain <span class="math inline">\(x\)</span> to lie within the domain <span class="math inline">\(|x|&lt;1\)</span>    (or, equivalently, <span class="math inline">\(-1&lt;x&lt;1\)</span>), then in the first case <span class="math inline">\(\dfrac{1}{1-x}\)</span> has the <span><strong>same numerical value</strong></span> as <span class="math inline">\(1+x+x^2+\cdots\)</span>    and in the second case <span class="math inline">\(\ln
(1+x)\)</span> has the same numerical value as <span class="math inline">\(x-\dfrac{x^2}{2}+\dfrac{x^3}{3}-\cdots\)</span>.</p>
<p>In the third example we see that <span class="math inline">\(e^x\)</span> has the same numerical value as <span class="math inline">\(1+x+\dfrac{x^2}{2!}+\cdots\)</span> but in this case there is no restriction to be placed on the value of <span class="math inline">\(x\)</span>    since <span>this</span> power series converges for all values of <span class="math inline">\(x\)</span>. The figure below shows this situation geometrically. As more and more terms are used from the series <span class="math inline">\(1+x+\dfrac{x^2}{2!}+\dfrac{x^3}{3!}\cdots\)</span>    the curve representing <span class="math inline">\(e^x\)</span> is a better and better approximation. In (a) we show the linear approximation to <span class="math inline">\(e^x\)</span>. In (b) and (c) we show, respectively, the quadratic and cubic
    approximations.
</p>
<figure class="HELM_figure">
    <img src="@@PLUGINFILE@@/16-5-fig-1-taylor1.svg" alt="Linear, quadratic and cubic approximations to e^x" width="624" height="167" class="img-responsive atto_image_button_text-bottom">
    <figcaption>Linear, quadratic and cubic approximations to \(e^x\)</figcaption>
</figure>
        <p>These power series representations are extremely important, from many points of view. Numerically, we can simply replace the function <span class="math inline">\(\dfrac{1}{1-x}\)</span> by the quadratic expression <span class="math inline">\(1+x+x^2\)</span>            as long as <span class="math inline">\(x\)</span> is so small that powers of <span class="math inline">\(x\)</span> greater than or equal to 3 can be ignored in comparison to quadratic terms. This approach can be used to approximate more complicated functions in terms of simpler polynomials. Our aim now is to see how these power series expansions are obtained.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201117170700+eLrunu</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210610105348+dqfv0H</text></tag>
    </tags>
    
</question>
</quiz>