<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62138 -->
  <question type="stack">
    <name>
				<text>23.5.1.3 Half-range Fourier series</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[The point is that all three periodic functions \(f_1(t)\), \(f_2(t)\), \(f_3(t)\) will give rise to a <b>different</b> Fourier series but all will represent the function \(f(t) = t^2\) over \(0 &lt; t &lt; \pi\). Fourier series obtained by
extending functions in this sort of way are often referred to as <b>half-range</b> series.
<p></p>
Normally, in applications, we require either a Fourier Cosine series (so we would complete a definition
as in (a) above to obtain an <b>even</b> periodic function) or a Fourier Sine series (for which, as in (b)
above, we need an <b>odd</b> periodic function.)
<p></p>
The above considerations apply equally well for a function defined over any interval.
<hr>
<h4 class="HELM_example">Example</h4>
<p></p>
Obtain the half range Fourier Sine series to represent \(f(t) = t^2\), \(0 &lt; t &lt; 3\).
<h4 class="HELM_solution">Solution</h4>
We first extend \(f(t)\) as an odd periodic function \(F(t)\) of <b>period 6</b>: \(f(t) = -t^2\), \(-3 &lt; t &lt; 0\).
So 
\[
F(t) = 
\begin{cases}
-t^2 \hspace{1cm} -3 &lt; t &lt; 0 \\
t^2 \hspace{1.4cm} 0 &lt; t &lt; 3
\end{cases}
\]<p style="text-align: center;"><img src="@@PLUGINFILE@@/23-5-1-example-F%28t%29.svg" alt="example-F(t) odd" width="261" height="97" class="img-responsive atto_image_button_text-bottom"><br></p>
We now evaluate the Fourier series of \(F(t)\) by standard techniques but take advantage of the
symmetry and put \(a_n = 0\), \(n = 0, 1, 2, \dots \).
<p></p>
Using the results for the Fourier Sine coefficients for period \(T\) as
\[ b_n = \dfrac{2}{T} \displaystyle{ \int _{-\frac{T}{2}} ^{\frac{T}{2}} } F(t) \sin \left( \dfrac{2n\pi t}{T} \right) dt ,\]
we put \(T = 6\) and, since the integrand is even (a product of 2 odd functions), we can write
\[ b_n = \dfrac{2}{3} \displaystyle{ \int _{0} ^{3} } F(t) \sin \left( \dfrac{2n\pi t}{6} \right) dt =  \dfrac{2}{3} \displaystyle{ \int _{0} ^{3} } t^2 \sin \left( \dfrac{n \pi t}{3} \right) dt \]
(Note that we always integrate over the originally defined range, in this case \(0 &lt; t &lt; 3\).)
<p></p>
We now have to integrate by parts (twice!)
\[ 
\begin{align}
b_n &amp;= \dfrac{2}{3} \left\{ \left[ -\dfrac{3t^2}{n\pi} \cos \left( \dfrac{n\pi t}{3} \right) \right] _0 ^3 + 2\left( \dfrac{3}{n\pi} \right) \int _0 ^3 t \cos \left( \dfrac{n \pi t}{3} \right) dt \right\} \\
&amp;= \dfrac{2}{3} \left\{ -\dfrac{27}{n \pi} \cos n\pi + \dfrac{6}{n\pi} \left[ \dfrac{3}{n\pi} t \sin \left( \dfrac{n \pi t}{3} \right) \right] _0^3 - \left( \dfrac{6}{n\pi} \right) \left( \dfrac{3}{n\pi} \right) \int _0^3 \sin \left( \dfrac{n \pi t}{3} \right) dt \right\} \\
&amp;=  \dfrac{2}{3} \left\{ -\dfrac{27}{n \pi} \cos n\pi -\dfrac{18}{n^2\pi^2} \left[ -\dfrac{3}{n\pi}\cos \left( \dfrac{n\pi t}{3} \right) \right] _0^3  \right\} \\
&amp;= \dfrac{2}{3} \left\{ - \dfrac{27}{n\pi} \cos n\pi + \dfrac{54}{n^3\pi^3} \left( \cos n\pi -1 \right) \right\}
\end{align}
\]
so
<p></p>
\( b_n =
\begin{cases}
-\dfrac{18}{n\pi} \hspace{2.4cm} n = 2, 4, 6, \dots \\
\dfrac{18}{n\pi} - \dfrac{72}{n^3\pi^3} \hspace{1cm} n = 1, 3, 5 \dots
\end{cases}
\)
<p></p>
So the required Fourier series is 
\[ F(t) = 18 \left( \dfrac{1}{\pi} - \dfrac{4}{\pi^3} \right) \sin \left( \dfrac{\pi t}{3} \right) - \dfrac{18}{2\pi}\sin \left( \dfrac{2\pi t}{3} \right) + 18 \left( \dfrac{1}{3\pi} - \dfrac{4}{27\pi^3} \right) \sin (\pi t) - \dots \]]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[\[ 
\begin{align}
b_n &amp;= \dfrac{2}{3} \left\{ \left[ -\dfrac{3t^2}{n\pi} \cos \left( \dfrac{n\pi t}{3} \right) \right] _0 ^3 + 2\left( \dfrac{3}{n\pi} \right) \int _0 ^3 t \cos \left( \dfrac{n \pi t}{3} \right) dt \right\} \\
&amp;= \dfrac{2}{3} \left\{ -\dfrac{27}{n \pi} \cos n\pi + \dfrac{6}{n\pi} \left[ \dfrac{3}{n\pi} t \sin \left( \dfrac{n \pi t}{3} \right) \right] _0^3 - \left( \dfrac{6}{n\pi} \right) \left( \dfrac{3}{n\pi} \right) \int _0^3 \sin \left( \dfrac{n \pi t}{3} \right) dt \right\} \\
&amp;=  \dfrac{2}{3} \left\{ -\dfrac{27}{n \pi} \cos n\pi -\dfrac{18}{n^2\pi^2} \left[ -\dfrac{3}{n\pi}\cos \left( \dfrac{n\pi t}{3} \right) \right] _0^3  \right\} \\
&amp;= \dfrac{2}{3} \left\{ - \dfrac{27}{n\pi} \cos n\pi + \dfrac{54}{n^3\pi^3} \left( \cos n\pi -1 \right) \right\}
\end{align}
\]]]></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200906080452+9hOHvV</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220518105720+3UbcaP</text></tag>
    </tags>
    
</question>
</quiz>