<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27189 -->
  <question type="stack">
    <name>
				<text>35.1.1.3 Subsets exercise</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>A factory produces cars over a five day period; Monday to Friday. Consider the following sets,</p>
<p>\(A = \{\hbox{cars produced from Monday to Friday}\}\)</p>
<p>\(B = \{\hbox{cars produced from Monday to Thursday}\}\)</p>
<p>\(D = \{\hbox{cars produced on Wednesday}\}\)</p>
<p>\(E = \{\hbox{cars produced on Wednesday and Thursday}\}\)</p>
<p>Which of the following statements are correct?</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p dir="ltr" style="text-align: left;">Let \(A = \{\hbox{Mo},&nbsp;\hbox{Tu},&nbsp;\hbox{We},&nbsp;\hbox{Th},&nbsp;\hbox{Fi} \}\),&nbsp;\(B = \{\hbox{Mo},&nbsp;\hbox{Tu},&nbsp;\hbox{We},&nbsp;\hbox{Th} \}\),&nbsp;\(D = \{\hbox{We}\}\) and&nbsp;\(E = \{\hbox{We},&nbsp;\hbox{Th} \}\), then \(D \subset E \subset B \subset A\). Note that \(E \subset A\) because \(E \subset B\) and \(B \subset A\).</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200703191227+Xk5pzJ</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+211004094139+KAl2VL</text></tag>
    </tags>
    
</question>
</quiz>