<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24528 -->
  <question type="stack">
    <name>
				<text>13.4 Introduction</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span>Integration by parts</span><br></h3>

<p>Integration by Parts is a technique for integrating products of functions. In this Section you will learn to recognise when it is appropriate to use the technique and have the opportunity to practise using it for finding both definite and indefinite integrals.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200803123638+a1T5KW</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200803130711+0fkEzO</text></tag>
    </tags>
    
</question>
</quiz>