<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24570 -->
  <question type="stack">
    <name>
				<text>15.2.1.7 Distribution of particles in a plane</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Distribution of particles in a plane </h4>
<p>If the particles are distributed in a plane then the position of the centre of mass can be calculated in a similar way.</p>
<figure class="HELM_figure"><img src="@@PLUGINFILE@@/15-2-fig-6-Particles.svg?time=1623340599524" alt="15-2-fig-6.svg" width="600" height="259" class="img-responsive atto_image_button_text-bottom"><br>
    <figcaption>These masses are distributed throughout the \(xy\) plane</figcaption>
</figure>
<p>Now we must consider the moments of the individual masses taken about the \(x\)-axis <strong>and</strong> about the \(y\)-axis. For example, in the Figure above, mass \(m_i\) has a moment \(m_iy_i\) about the \(x\)-axis and a moment \(m_ix_i\) about the
    \(y\)-axis. Now we define the centre of mass at that point \((\bar{x}, \bar{y})\) such that the total mass, \(M = m_1 + m_2 + \dots + m_n\), placed at this point would have the same moment about each of the axes as the sum of the individual moments
    of the particles about these axes.</p>
<div class="HELM_keypoint">
    <h4>Key Point</h4>
    <p>The <strong>centre of mass</strong> of \(m_1, m_2, \dots, m_n\) located at \((x_1,y_1), (x_2,y_2), \dots, (x_n,y_n)\) has coordinates \((\bar{x},\bar{y})\) where \[ \bar{x}=\dfrac{\sum_{i=1}^n m_i x_i}{\sum_{i=1}^n m_i}, \quad \quad \bar{y}=\dfrac{\sum_{i=1}^n
        m_i y_i}{\sum_{i=1}^n m_i}. \] </p>
</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200817132656+QDmOqo</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210610155709+CO3jlA</text></tag>
    </tags>
    
</question>
</quiz>