<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25974 -->
  <question type="stack">
    <name>
				<text>19.2.4.2 Solve an exact equation</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Solve the differential equation \[\dfrac{\mathrm{d}}{\mathrm{d}x}(y \cos({@b@}x)) = {@a@}\cos({@b@}x)\]</p>
<p>\(y=\) [[input:ans1]] [[validation:ans1]]\(+ C \sec({@b@}x)\)</p><p>where \(C\) is an arbitrary constant.</p><p>[[feedback:prt1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p dir="ltr" style="text-align: left;">We are given&nbsp;the differential equation \[\dfrac{\mathrm{d}}{\mathrm{d}x}(y \cos({@b@}x)) = {@a@}\cos({@b@}x)\]</p><p dir="ltr" style="text-align: left;">Integrating both sides with respect to \(\mathrm{d}x\) gives \[\int d(y \cos({@b@}x)) \, \mathrm{d}x = \int {@a@} \cos({@b@}x) \, \mathrm{d}x.\]</p><p dir="ltr" style="text-align: left;">Then we have \[ y \cos({@b@}x)&nbsp; = \dfrac{{@a@}\sin({@b@}x)}{{@b@}} + C.\]&nbsp;</p><p dir="ltr" style="text-align: left;">You can now make \(y\) the subject to get \[y = {@ta1@} + C \sec({@b@}x).\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>a: rand(4)+2;
b:rand(6)+2;
ta1: a*tan(b*x)/b</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(\dfrac{d}{dx}(y \cos({@b@}x)) = {@a@}\cos({@b@}x)\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:moodle2.maths.ed.ac.uk+210315173405+DpYqZL</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210614120209+7BZOai</text></tag>
    </tags>
    <deployedseed>1408164508</deployedseed>
    <deployedseed>920706736</deployedseed>
    <deployedseed>1452380673</deployedseed>
    <deployedseed>1561119397</deployedseed>
    <deployedseed>1005586191</deployedseed>
    <deployedseed>1398515759</deployedseed>
    <deployedseed>152633459</deployedseed>
    <deployedseed>344244274</deployedseed>
    <deployedseed>72970589</deployedseed>
    <deployedseed>985096308</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1-1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>