<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27196 -->
  <question type="stack">
    <name>
				<text>35.2.1.2 Sample space of a single coin toss</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p> If the sample space can be written in the form of a list (possibly infinite) then it is called a <strong>discrete sample space</strong> (e.g. number of tosses of a fair coin before Heads occurs). If this is not possible then it is called a <strong>continuous sample space</strong> (e.g. positions where shells land in a tank battle). The following example, therefore, concerns a discrete sample space.</p>
<hr>
<h4 class="HELM_example">Example</h4>
<p>Obtain the sample space of the experiment <em>throwing a single coin</em>.</p>
<h4 class="HELM_solution">Solution</h4>
<p> Consider the experiment of throwing a coin which can land Heads up \((H)\) or Tails up \((T)\). We list the outcomes as a set \(\{H, T\}\) − the order being unimportant. This set, \(\{H, T\}\), is the sample space. On any particular throw of a coin, Heads or Tails are equally likely to occur. We say that, for a <strong>fair</strong> coin, \(H\) and \(T\) are <strong>equally likely outcomes</strong>. </p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200629165320+Gq2pS1</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200629165930+PA9f9W</text></tag>
    </tags>
    
</question>
</quiz>