<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24426 -->
  <question type="stack">
    <name>
				<text>11.2.3.3 Example</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<hr>
<h4 class="HELM_example">Example</h4>
<p>Find the value of the derivative of \(y = x^3\) where \(x = 2\). Interpret your result.</p>

<h4 class="HELM_solution">Solution</h4>
<p> We have \(y = x^3\) and so \( \dfrac{dy}{dx} = 3x^2\).
</p><p></p>
When \(x = 2\), \(\dfrac{dy}{dx} = 3(2)^2 = 12\), that is, \( \dfrac{dy}{dx}(2) = 12\)  (Equivalently, \(y'(2) = 12\)).
<p></p>
The derivative is positive when \(x = 2\) and so \(y\) is increasing at this point. When \(x = 2\), \(y\) is increasing at a rate of \(12\) vertical units per horizontal unit.
<hr>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200725100535+xsFRL8</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210901194632+qS5MI5</text></tag>
    </tags>
    
</question>
</quiz>