<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24612 -->
  <question type="stack">
    <name>
				<text>14.3.1.1 Volumes generated by rotating curves about the x-axis</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>Volumes generated by rotating curves about the x-axis</h3>

<p>The figure below shows a graph of the function \(y=2x\) for \(x\) between 0 and 3.</p><p style="text-align: center;"><img src="@@PLUGINFILE@@/14-3-fig-8-linegraph.svg" alt="A graph of the function y = 2x, for x between 0 and 3." width="145" height="128" class="img-responsive atto_image_button_text-bottom"><br></p>
<p>Imagine rotating the line \(y=2x\) by one complete revolution (\(360^{\circ}\) or \(2\pi\) radians) around the \(x\)-axis. The surface so formed is the surface of a cone as shown in the figure below. Such a three-dimensional shape is known as a <strong>solid of revolution</strong>.
    We now discuss how to obtain the volumes of such solids of revolution.</p><p style="text-align: center;"><img src="@@PLUGINFILE@@/14-3-fig-9-conerevolution.svg" alt="When line y = 2x is rotated around the axis a solid is generated. It's a cone, vertex at origin, axis = x-axis, base at x=3." width="199" height="261" class="img-responsive atto_image_button_text-bottom"><br></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201201142215+4EgRjz</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210607110534+Zv0Zi9</text></tag>
    </tags>
    
</question>
</quiz>