<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25733 -->
  <question type="stack">
    <name>
				<text>2.4.5.5 Task: Complete sketch (odd) *****</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>Extend the graph below in order to produce a graph of an odd function.</p><p><img src="@@PLUGINFILE@@/block6-fig5-1.svg" alt="block6-fig5" width="188" height="101" class="img-responsive atto_image_button_text-bottom"><br></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p><img src="@@PLUGINFILE@@/block6-fig5%28new%29-1.svg" alt="block6-fig5(new)" width="188" height="101" class="img-responsive atto_image_button_text-bottom"><br></p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200618184041+MRONRt</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200618192007+NlR8qu</text></tag>
    </tags>
    
</question>
</quiz>