<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26581 -->
  <question type="stack">
    <name>
				<text>8.2.1.6 Engineering Example: Currents in two loops</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_example">Engineering Example: Currents in two loops</h4>

<p> In the circuit shown below find the currents \((i_1,\ i_2)\) in the loops: </p>

<div style="text-align: center;"><img src="@@PLUGINFILE@@/8-2-fig-1-currents.svg" alt="Figure displaying currents in two loops" class="img-responsive atto_image_button_text-bottom" width="257" height="115"></div>
<p style="text-align: center;"><strong>Figure 1</strong></p>
<p style="text-align: left;"><strong><br></strong></p>

<h4 class="HELM_solution">Solution</h4>

<p>We note that the current across the <span class="math inline">\(3\ \Omega\)</span> resistor (travelling top to bottom in the diagram) is given by <span class="math inline">\((i_1-i_2)\)</span>. With this proviso we can apply Kirchhoff’s law: <span class="math display">\[\begin{matrix}
    \text{In the left-hand loop} &amp; 3(i_1-i_2)+5i_1=3  &amp; \rightarrow &amp; 8i_1-3i_2=3 \\
    \text{In the right-hand loop} &amp; 3(i_2-i_1)+6i_2=4 &amp; \rightarrow &amp; -3i_1+9i_2=4 \\
\end{matrix}\]</span></p>
<p>In matrix form: <span class="math display">\[\left[
\begin{matrix} 
    8 &amp; -3\\
    -3 &amp;  9 \\
\end{matrix} 
\right]
\left[
\begin{matrix}
    i_1 \\
    i_2 \\
\end{matrix}
\right]=
\left[\begin{matrix}
    3 \\
    4 \\
\end{matrix}
\right].\]</span></p>

<p>The inverse of <span class="math display">\[\left[\begin{matrix}\hfill 8&amp; -3 \\ -3&amp; 9\end{matrix}\right]\]</span> is <span class="math display">\[\dfrac{1}{63}\left[\begin{matrix}9&amp;3\\3&amp;8\end{matrix}\right]\]</span> so solving gives
    <span class="math display">\[\left[\begin{matrix}i_1\\ i_2\end{matrix}\right]=\dfrac{1}{63}\left[\begin{matrix}9&amp;3\\3&amp;8\end{matrix}\right] \left[\begin{matrix}3\\ 4\end{matrix}\right] =\dfrac{1}{63}\left[\begin{matrix}39\\ 41\end{matrix}\right]\]
        </span> so <span class="math inline">\(\displaystyle{i_1= \frac{39}{63}}\)</span>, <span class="math inline">\(\displaystyle{i_2= \frac{41}{63}}\)</span>.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200622131101+pjnbg2</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210630131235+wx1noM</text></tag>
    </tags>
    
</question>
</quiz>