<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62096 -->
  <question type="stack">
    <name>
				<text>23.2.2.1 Functions of period 2*pi</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Functions of period \(2\pi\)</h4>
We now discuss how to represent periodic non-sinusoidal functions \(f(t)\) of period \(2\pi\) in terms of
sinusoids, i.e. how to obtain Fourier series representations. As already discussed we expect such
Fourier series to contain harmonics of frequency \( \dfrac{n}{2\pi}\) (\(n = 1, 2, 3, \dots \) ) and, if the periodic function has a non-zero average value, a constant term.
<p></p>
Thus we seek a Fourier series representation of the general form
\[ f(t) = \dfrac{a0}{2} + a_1 \cos t + a_2 \cos 2t + \dots + b_1 \sin t + b_2 \sin 2t + \dots \]
The reason for labelling the constant term as \( \dfrac{a0}{2}\) will be discussed later. The amplitudes \(a_1, a_2, \dots , b_1, b_2, \dots \) of the sinusoids are called <b>Fourier coefficients</b>
<p></p>
Obtaining the Fourier coefficients for a given periodic function \(f(t)\) is our main task and is referred
to as Fourier Analysis. Before embarking on such an analysis it is instructive to establish, at least
qualitatively, the plausibility of approximating a function by a few terms of its Fourier series.

<hr>
<h4 class="HELM_example">Example</h4>
<p></p>
Consider the square wave of period \(2\pi\) one period of which is shown in the Figure below.<p><br></p><p style="text-align: center;"><img src="@@PLUGINFILE@@/23-2-2-square-pulse.svg" alt="square-pulse" width="330" height="113" class="img-responsive atto_image_button_text-bottom"><br></p>
<ol class="HELM_parts_inline">
<li>Write down the analytic description of this function,</li>
<p></p>
<li>State whether you expect the Fourier series of this function to contain a constant term,</li>
<p></p>
<li>List any other possible features of the Fourier series that you might
expect from the graph of the square-wave function.</li>
</ol>
<h4 class="HELM_solution">Solution</h4>
<ol class="HELM_parts">
<li>We have
\[ 
\begin{equation}
f(t) = \begin{cases} 
4 \hspace{1cm} -\dfrac{\pi}{2} &lt; t &lt; \dfrac{\pi}{2} \\
0&nbsp;&nbsp;\hspace{1cm} -\pi &lt; t &lt; -\dfrac{\pi}{2}, \dfrac{\pi}{2} &lt; t &lt; \pi
\end{cases}
\end{equation}
\]
\[f(t + 2\pi) = f(t)\]
</li><li>The Fourier series will contain a constant term since the square wave here is non-negative and
cannot therefore have a zero average value. This constant term is often referred to as the d.c.
(direct current) term by engineers.</li>
<li>Since the square wave is an even function (i.e. the graph has symmetry about the \(y\) axis) then
its Fourier series will contain cosine terms but not sine terms because only the cosines are even
functions. (Well done if you spotted this at this early stage!)</li></ol>
<p>&nbsp;</p><hr>

It is possible to show, and we will do so later, that the Fourier series representation of this square
wave is
\[ 2 + \dfrac{8}{\pi} \left\{ \cos t - \dfrac{1}{3}\cos 3t + \dfrac{1}{5}\cos 5t - \dfrac{1}{7}\cos 7t + \dots \right\} \]
i.e the Fourier coefficients are 
\[ \dfrac{a_0}{2} = 2, \hspace{0.5cm} a_1 = \dfrac{8}{\pi}, \hspace{0.5cm} a_2 = 0, \hspace{0.5cm} a_3 = -\dfrac{8}{3\pi}, \hspace{0.5cm} a_4 = 0, \hspace{0.5cm} a_5 = \dfrac{8}{5\pi}, \dots \]

Note, as well as the presence of the constant term and of the cosine (but not sine) terms, that only
odd harmonics are present i.e. sinusoids of period \(2\pi\), \( \dfrac{2\pi}{3}\), \( \dfrac{2\pi}{5}\), \( \dfrac{2\pi}{7}\), \(\dots\), or of frequency \(1,3,5,7, \dots\), times the fundamental frequency \(\dfrac{1}{2\pi}\).
We now show in the Figure below graphs of
<p></p>
\( (i)\) the square wave <p></p>
\( (ii)\)&nbsp; the first two terms of the Fourier series representing the square wave<p></p>
\( (iii)\) the first three terms of the Fourier series representing the square wave <p></p>
\( (iv)\) the first four terms of the Fourier series representing the square wave <p></p>
\( (v)\)&nbsp; the first five terms of the Fourier series representing the square wave <p></p>
Note: We show the graphs for \(0 &lt; t &lt; \pi\) only since the square wave and its Fourier series are even.<p style="text-align: center;"><img src="@@PLUGINFILE@@/23-2-2-square-wave-formation.svg" alt="square-wave-formation" width="606" height="408" class="img-responsive atto_image_button_text-bottom"><br></p>
We can clearly see from the Figure above that as the number of terms is increased the graph of the Fourier series gradually approaches that of the original square wave - the ripples increase in number but
decrease in amplitude. (The behaviour near the <b>discontinuity</b>, at \(t = \dfrac{\pi}{2}\), is slightly more complicated and it is possible to show that however many terms are taken in the Fourier series, some "overshoot" will always occur. This effect, which we do not discuss further, is known as the Gibbs Phenomenon.)&nbsp;<p></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200815124327+NK87Iw</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220518092406+rj3a2L</text></tag>
    </tags>
    
</question>
</quiz>