<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24471 -->
  <question type="stack">
    <name>
				<text>11.7.0 Introduction implicit differentiation</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Introduction </h4>

This Section introduces implicit differentiation which is used to differentiate functions expressed in implicit form (where the variables are found together). Examples are \(x^3 + xy + y^2 = 1\), and \( \dfrac{x^2}{a^2} + \dfrac{y^2}{b^2} = 1 \) which represents an ellipse.]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200809091135+puBdgj</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200809091559+tUJEJp</text></tag>
    </tags>
    
</question>
</quiz>