<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27199 -->
  <question type="stack">
    <name>
				<text>35.2.1.5 Tree diagram and number of possible outcomes</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4 class="HELM_example"> Example </h4>
<p> A car manufacturer offers certain options on its family cars. Customers may order: </p>
<ul>
<li> either automatic gearboxes or manual gearboxes </li>
<li> either sunroof or air-conditioning </li>
<li> either steel wheels or alloy wheels </li>
<li> either solid colour paint or metallic paint </li>
</ul>
<p> Find the number of outcomes in the sample space of options that it is possible to
order and represent them using a suitable diagram. </p>
<h4 class="HELM_solution"> Solution </h4>
<p> A suitable diagram is shown in the figure below. The diagram makes it easy to find the number of outcomes simply by counting. It also points the way to a formula for calculating the number of outcomes.</p><p><img src="@@PLUGINFILE@@/3-5-2-tree.svg" alt="Tree diagram bifurcating where choices in car equipment differ" role="presentation"><br></p>

<p><strong>Figure:</strong> Tree diagram. </p>
<p> In this case there is a total of 16 outcomes in the sample space of options. Note that in each case the customer makes two choices. This implies that there are \[2 \times 2 \times 2 \times 2 = 16\]options in total.
</p>
<p> Diagrams such as the one above are called <strong>tree diagrams</strong>. They are only suitable in simple situations. </p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200630122243+ucrMCj</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200803120136+pP9Azj</text></tag>
    </tags>
    
</question>
</quiz>