<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62192 -->
  <question type="stack">
    <name>
				<text>23.6 Introduction</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Introduction</h4>

In this Section we show how a Fourier series can be expressed more concisely if we introduce the
complex number \(i\) where \(i^2 = -1\). By utilising the Euler relation:
\[ e^{i\theta} \equiv \cos \theta  + i \sin \theta \]
we can replace the trigonometric functions by complex exponential functions. By also combining the
Fourier coefficients \(a_n\) and \(b_n\) into a complex coefficient \(c_n\) through
\[ c_n = \dfrac{1}{2} \left( a_n - i b_n \right) \]
we find that, for a given periodic signal, both sets of constants can be found in one operation.
<p></p>
We also obtain Parseval's theorem which has important applications in electrical engineering.
<p></p>
The complex formulation of a Fourier series is an important precursor of the Fourier transform which
attempts to Fourier analyse non-periodic functions.]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200913063839+3pWSDB</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200913063839+UGP5rH</text></tag>
    </tags>
    
</question>
</quiz>