<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26733 -->
  <question type="stack">
    <name>
				<text>10.2.2.1 The polar form of a complex number</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>The polar form of a complex number </h3>
<p>We have seen that the complex number <span class="math inline">\(z=a+i b\)</span> can be represented by a line pointing out from the origin and ending at a point with Cartesian coordinates <span class="math inline">\((a,b)\)</span></p><br>
<figure class="HELM_figure"><img src="@@PLUGINFILE@@/10.2-fig3-polar_form.svg" alt="" width="287" height="201" class="img-responsive atto_image_button_text-bottom"><br></figure>
<br>
<p>To locate the point P we introduce polar coordinates <span class="math inline">\((r,\theta)\)</span> where <span class="math inline">\(r\)</span> is the positive distance from <span class="math inline">\(0\)</span> and <span class="math inline">\(\theta\)</span>    is the angle measured from the positive <span class="math inline">\(x\)</span>-axis, as shown in the figure above. From the properties of the right-angled triangle there is an obvious relation between <span class="math inline">\((a,b)\)</span> and
    <span class="math inline">\((r,\theta )\)</span>: <span class="math inline">\[
a=r\cos\theta\qquad b=r\sin\theta\]</span> or equivalently, <span class="math inline">\[
r=\sqrt{a^2+b^2}\qquad \tan\theta = \frac{b}{a}.\]</span> This leads to an alternative way of writing a complex number: <span class="math display">\[\begin{aligned}
z=a+i b&amp;=r\cos\theta +i r\sin\theta\\
&amp;=r(\cos\theta
+i\sin\theta)\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\qquad\end{aligned}\]</span></p>
<p>The angle <span class="math inline">\(\theta\)</span> is called the <strong> argument </strong>of <span class="math inline">\(z\)</span> and written, for short, arg<span class="math inline">\((z)\)</span>. The non-negative real number <span class="math inline">\(r\)</span>    is the modulus of <span class="math inline">\(z\)</span>. We normally consider <span class="math inline">\(\theta\)</span> measured in radians to lie in the interval <span class="math inline">\(-\pi&lt;\theta\le \pi\)</span> although any value
    <span class="math inline">\(\theta +2k\pi\)</span> for integer <span class="math inline">\(k\)</span> will be equivalent to <span class="math inline">\(\theta\)</span>. The angle <span class="math inline">\(\theta\)</span> may be expressed in radians or degrees, it is standard
        to use radians and we should be explicit that we use degrees is this is the case.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200629120734+dhDaSv</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210701151150+LuSBzy</text></tag>
    </tags>
    
</question>
</quiz>