<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62194 -->
  <question type="stack">
    <name>
				<text>23.6.2 Revision of the exponential form of a complex number</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Revision of the exponential form of a complex number</h4>

Recall that a complex number in <b>Cartesian</b> form which is written as
\[z = a + ib \]
where \(a\) and \(b\) are real numbers and \(i^2 = -1\), can be written in <b>polar</b> form as
\[ z = r \left( \cos \theta + i \sin \theta \right) \]
where \(r = \lvert z \rvert =\sqrt{a^2 + b^2} \) and \(\theta\), the <b>argument</b> or <b>phase</b> of \(z\) is such that
\[ a = r \cos \theta \hspace{1cm} b = r \sin \theta \]
A more concise version of the polar form of \(z\) can be obtained by defining a <b>complex exponential</b>
quantity \(e^{i\theta}\) by Euler's relation
\[ e^{i\theta} \equiv \cos \theta + i \sin \theta \]
The polar angle \(\theta\) is normally expressed in <b>radians</b>. Replacing i by \( - i\) we obtain the alternative form
\[ e^{-i\theta} \equiv \cos \theta - i \sin \theta \]]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200913162024+jxlAKM</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200914182636+uYUnXN</text></tag>
    </tags>
    
</question>
</quiz>