<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25998 -->
  <question type="stack">
    <name>
				<text>18.2.3.1 Engineering example</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3><span>Engineering example</span><br></h3>

<p><strong>The ideal gas law and Redlich-Kwong equation</strong></p>

<p><strong>Introduction</strong></p>

<p>In Chemical Engineering it is often necessary to be able to equate the pressure, volume and temperature of a gas. One relevant equation is the ideal gas law \(PV = n R T\) &nbsp; (1) where \(P\) is pressure, \(V\) is volume, \(n\) is the number of moles of gas, \(T\) is temperature and \(R\) is the ideal gas constant (= 8.314 J mol\(^{-1}\) K\(^{-1}\), when all quantities are in S.I. units).</p>
<p>The ideal gas law has been in use since 1834, although its special cases at constant temperature (Boyle’s Law, 1662) and constant pressure (Charles’ Law, 1787) had been in use many decades previously. While the ideal gas law is adequate in many circumstances, it has been superseded by many other laws where, in general, simplicity is weighed against accuracy.</p>
<p>One such law is the Redlich-Kwong equation \(P = \dfrac{R T}{V-b}  -  \dfrac{a}{\sqrt{T}  V (V+b)}\) &nbsp; (2) where, in addition to the variables in the ideal gas law, the extra parameters \(a\) and \(b\) are dependent upon the particular gas under consideration. Clearly, in both equations the temperature, pressure and volume will be positive. Additionally, the Redlich-Kwong equation is only valid for values of volume greater than the parameter \(b\) - in practice however, this is not a limitation, since the gas would condense to a liquid before this point was reached.</p>

<p><strong>Problem in words</strong> Show that for both Equations (1) and (2)</p> 
<ol> 
<li> for constant temperature, the pressure decreases as the volume increases (Note : in the Redlich-Kwong equation, assume that \(T\) is large.)</li>
<li> for constant volume, the pressure increases as the temperature increases. </li>
</ol> 
<p><strong>Mathematical statement of problem</strong> 
</p><p>For both Equations (1) and (2), and for the allowed ranges of the variables, show that</p> 
<ol> 
<li>\(\dfrac{\partial P}{\partial V} &lt; 0\) for \(T\) = constant</li> 
<li>\(\dfrac{\partial P}{\partial T} &gt; 0\) for \(V\) = constant</li> 
</ol> 
<p>Assume that \(T\) is sufficiently large so that terms in \(T^{-1/2}\) may be neglected when compared to terms in \(T\).</p> 
<p><strong>Mathematical analysis</strong></p> 
<ol> 
<li><strong>Ideal gas law</strong><br> 
<p>This can be rearranged as \(P = \dfrac{n R  T}{V}\) so that 
</p><ol>
<li><p> at constant temperature \(\dfrac{\partial P}{\partial V} = \dfrac{-n R \, T}{V^2} &lt; 0\) as all quantities are positive </p></li> 
<li> <p> for constant volume \(\dfrac{\partial P}{\partial T} = \dfrac{n R}{V} &gt; 0\) as all quantities are positive</p></li> 
</ol>
</li><li><strong>Redlich-Kwong equation</strong> 
<p>\(P = \dfrac{RT}{V - b}  -  \dfrac{a}{\sqrt{T}{V(V+b)}} \\ \phantom{P}= RT(V-b)^{-1} - a  T^{-1/2}  (V^2+V b)^{-1}\)</p> <p>so that </p><ol> 
<li><p> at constant temperature \(\dfrac{\partial P}{\partial V} = -R T(V-b)^{-2} + a  T^{-1/2} (V^2+V b)^{-2}(2V+b)\) which, for large \(T\), can be approximated by \( \dfrac{\partial P}{\partial V} \approx \dfrac{-R \, T}{(V-b)^2}  &lt; 0\) as all quantities are positive</p></li> 
<li> for constant volume \(\dfrac{\partial P}{\partial T} = R(V-b)^{-1} + \frac{1}{2} \, a \, T^{-3/2}(V^2+Vb)^{-1}  &gt; 0\) as all quantities are positive<p></p></li>
</ol></li> </ol> 
<p>In practice, the restriction on \(T\) is not severe, and regions in which \({\dfrac{\partial P}{\partial V} &lt; 0}\) does not apply are those in which the gas is close to liquefying and, therefore, the entire Redlich-Kwong equation no longer applies.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200819132536+8v9hdG</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200819142725+wSAafR</text></tag>
    </tags>
    
</question>
</quiz>