<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26800 -->
  <question type="stack">
    <name>
				<text>22.2.2.4 Writing system of 1st order lin. eqs. in matrix form</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Write in matrix form the pair of <span><strong>coupled</strong></span> differential equations <br><br><span class="math inline">\(\qquad \left.\begin{array}{l}\dot x={@eq1@}\\ \dot y={@eq2@}\end{array}\right\}\ \ \)</span></p><br>
\(\left[\begin{array}{c}\dot x\\ \dot y\end{array}\right]=\) [[input:ans1]][[validation:ans1]]\(\left[\begin{array}{c}x\\ y\end{array}\right]\) </div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>\(\qquad\left[\begin{array}{c}\dot x\\ \dot y\end{array}\right]={@mat@}\left[\begin{array}{c}x\\ y\end{array}\right]\)<br>
<br>\(\qquad \ \ \dot{\underline{x}}\ \  \ \ =\ \ \ A\ \underline{x} \qquad\qquad\ \ \ \!\!\!\!\)</span></p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020070100</text>
		</stackversion>
    <questionvariables>
				<text>ordergreat(x,y)
mat:rand(matrix([4,4],[4,4]))+matrix([1,1],[1,1])
eq1:mat[1,1]*x+mat[1,2]*y
eq2:mat[2,1]*x+mat[2,2]*y</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>{@mat@}</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>mat</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>matrix,.</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>2</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mat</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>x_correct</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>x_wrong</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200618145431+bm4Tvx</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200827141447+v1xBXh</text></tag>
    </tags>
    <deployedseed>1850080931</deployedseed>
    <deployedseed>2131604502</deployedseed>
    <deployedseed>601423159</deployedseed>
    <deployedseed>2007913165</deployedseed>
    <deployedseed>2140383476</deployedseed>
    <deployedseed>309240784</deployedseed>
    <deployedseed>381104214</deployedseed>
    <deployedseed>1887476697</deployedseed>
    <deployedseed>1351291173</deployedseed>
    <deployedseed>1697182726</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>mat</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>x_correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(2*mat,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>x_wrong</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>