<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26518 -->
  <question type="stack">
    <name>
				<text>6.2.2.7 Common hyperbolic identities</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<div class="HELM_keypoint">
<h4>Key Point</h4>
<p>The fundamental identity relating hyperbolic functions is:\[\cosh^2x - \sinh^2x \equiv 1.\] This is the hyperbolic function equivalence of the trigonometric identity: \(\cos^2x + \sin^2x \equiv 1\).</p></div><br><br>
<div class="HELM_keypoint">

<h4>Key Point: More hyperbolic identities</h4>
\[\begin{align*} \cosh(x+y) &amp;\equiv \cosh x \cosh y+\sinh x \sinh y \\ \sinh(x+y) &amp;\equiv \sinh x \cosh y+\cosh x\sinh y \\ \sinh2x &amp;\equiv 2\sinh x \cosh x \\ \cosh2x &amp;\equiv \cosh^2x+\sinh^2x \equiv 2\cosh^2x-1 \equiv 1+2\sinh^2x \end{align*} \]
</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200622144250+d5x7Om</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200723144548+hIb68B</text></tag>
    </tags>
    
</question>
</quiz>