<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24292 -->
  <question type="stack">
    <name>
				<text>5.1.1.6</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>So, in the context of modelling motion under gravity, the <strong>initial velocity</strong> determines the vertical displacement of the line, its <strong>intercept</strong> on the \(v\)-axis, and the <strong>acceleration</strong> determines the <strong>slope</strong>. Again, given the modelling context, both of these influence the range of validity of the model since they alter the time taken for the rock to reach the ground and this fixes the upper limit on time.</p> <p>Like velocity, acceleration has direction as well as magnitude. As long as position is being measured downwards, and only gravity is considered to act, falling objects do not provide any examples of negative accelerations - but rocket motion does. Where downwards accelerations are represented as positive, an upwards acceleration will be negative. So a model of the motion of a rocket accelerating away from the Earth could include a constant negative acceleration. Horizontal acceleration, say of a road vehicle, in the same direction as position as being measured, is represented as positive. Deceleration, for example when this vehicle is being braked, implies that velocity is decreasing with time, and is represented as negative. In mathematical modelling, it is usual to refer to <em>acceleration</em>, whether it represents positive acceleration or deceleration. Suppose that we are describing the motion of a rocket taking off vertically during its initial booster stage of 10 s. We might model the acceleration as a constant \(-20\ \text{m/s}^{-2}\). The negative sign arises because downwards is being taken as the positive direction but the acceleration is upwards. Since the rocket is starting from rest, an appropriate function is </p><p>\(\qquad v = -20t \) \((0\le t\le 10)\) </p><p>This should describe the variation of its velocity with time until the end of the initial booster stage of its flight. Figure 4 shows the corresponding graph of velocity against time. Note the way in which the graph slopes downwards to the right. This function describes an increasingly <em>negative</em> velocity as time passes, consistent with an increasing <em>upwards</em> velocity. The corresponding graph for a positive acceleration of the same magnitude would slope upwards towards the right.</p>
<p><img width="248" height="153" class="img-responsive atto_image_button_text-bottom" alt="velocity time graph y= -20t" src="@@PLUGINFILE@@/5.1%20linear%20functions%20figure%202.svg"></p> <p><strong>Figure </strong>: Variation of velocity of rocket during the initial booster stage.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200724080330+vK9NZr</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220215174056+AxSOBZ</text></tag>
    </tags>
    
</question>
</quiz>