<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24361 -->
  <question type="stack">
    <name>
				<text>16.1.2.5 Sequences may be neither geometric nor arithmetic</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>The examples of arithmetic and geometric sequences are particularly nice and relatively easy to work with. It is important to note that a general sequence will not fall into either category. One easily defined example is the harmonic sequence&nbsp;</p><p>\[1,\frac12,\frac13,\frac14\,\dots,\]</p><p>whose \(n\)th term is \(\dfrac{1}{n}\), is neither arithmetic nor geometric.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+201005164345+VwbulE</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+201116183734+pDr7Y9</text></tag>
    </tags>
    
</question>
</quiz>