<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24610 -->
  <question type="stack">
    <name>
				<text>14.2.2.5 Finding more r.m.s. of more functions</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Calculate the r.m.s. value of the following to 4 decimal places:</p><br>

<ol class="HELM_parts">
    <li>
        <p>\(f(t)={@f1@}\) across \([{@a1@},{@b1@}]\)</p>
        <p>r.m.s. value \(=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
    </li>
    <li>
        <p>\(f(t)={@f2@}\) across \([{@a2@},{@b2@}]\)</p>
        <p>r.m.s. value \(=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>
    </li>
    <li>
        <p>\(f(t)={@f3@}\) across \([{@a3@},{@b3@}]\)</p>
        <p>r.m.s. value \(=\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>
    </li>
</ol>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>Remember the Key Point, r.m.s. value \(\displaystyle = \sqrt{\frac{1}{b-a} \int_a^b [f(t)]^2 \ \mathrm{d}t}\).</p><br>

<ol class="HELM_parts">
    <li>Using the Key Point with \(a={@a1@}\), \(b={@b1@}\) and \(f(t)={@f1@}\) we have: \[\begin{aligned}\text{r.m.s. value }&amp;= \sqrt{\frac{1}{{@b1@}-{@a1@}} \int_{@a1@}^{@b1@} \left[{@f1@}\right]^2 \ \mathrm{d}t}\\ &amp;= \sqrt{\frac{1}{{@b1-a1@}}\left[{@integrate(f1^2,t)@}\right]_{@a1@}^{@b1@}}\\
        &amp;= \sqrt{\frac{1}{{@b1-a1@}}\left[{@subst(t=b1,integrate(f1^2,t))@}-{@subst(t=a1,integrate(f1^2,t))@}\right]}\\ &amp;\approx {@ta1@} \end{aligned}\]</li>

    <li>Using the Key Point with \(a={@a2@}\), \(b={@b2@}\) and \(f(t)={@f2@}\) we have: \[\begin{aligned}\text{r.m.s. value }&amp;= \sqrt{\frac{1}{{@b2@}-({@a2@})} \int_{@a2@}^{@b2@} \left[{@f2@}\right]^2 \ \mathrm{d}t}\\ &amp;= \sqrt{\frac{1}{{@b2-a2@}}\left[{@integrate(f2^2,t)@}\right]_{@a2@}^{@b2@}}\\
        &amp;= \sqrt{\frac{1}{{@b2-a2@}}\left[{@subst(t=b2,integrate(f2^2,t))@}-{@subst(t=a2,integrate(f2^2,t))@}\right]}\\ &amp;\approx {@ta2@} \end{aligned}\]</li>

    <li>Using the Key Point with \(a={@a3@}\), \(b={@b3@}\) and \(f(t)={@f3@}\) we have: \[\begin{aligned}\text{r.m.s. value }&amp;= \sqrt{\frac{1}{{@b3@}-({@a3@})} \int_{@a3@}^{@b3@} \left[{@f3@}\right]^2 \ \mathrm{d}t}\\ &amp;= \sqrt{\frac{1}{{@b3-a3@}}\left[{@integrate(f3^2,t)@}\right]_{@a3@}^{@b3@}}\\
        &amp;= \sqrt{\frac{1}{{@b3-a3@}}\left[{@subst(t=b3,integrate(f3^2,t))@}-\left({@subst(t=a3,integrate(f3^2,t))@}\right)\right]}\\ &amp;\approx {@ta3@} \end{aligned}\]</li>
</ol>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>const1: rand([1,2,3]);
const2: rand([1,2,3]);

f1: sqrt(t+const1);
a1: 0;
b1: 3;
ta1: sqrt(1/(b1-a1)*integrate(f1^2,t,a1,b1));
ta1: decimalplaces(float(ta1),4);

f2: e^(const2*t);
a2: -1;
b2: 1;
ta2: sqrt(1/(b2-a2)*integrate(f2^2,t,a2,b2));
ta2: decimalplaces(float(ta2),4);

f3: const1+e^(const2*t);
a3: -1;
b3: 1;
ta3: sqrt(1/(b3-a3)*integrate(f3^2,t,a3,b3));
ta3: decimalplaces(float(ta3),4);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\( ({@f1@} , [{@a1@},{@b1@}], \, m={@ta1@}), \quad ({@f2@} , [{@a2@},{@b2@}], \, m={@ta2@}), \quad ({@f3@} , [{@a3@},{@b3@}], \, m={@ta3@}) \)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.</span>]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4,maxdp:4</options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4,maxdp:4</options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>mindp:4,maxdp:4</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>NumDecPlaces</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack-back1.maths.ed.ac.uk+201130141243+8SLHDT</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210608083726+Ch3tM7</text></tag>
    </tags>
    <deployedseed>1247152815</deployedseed>
    <deployedseed>1646395292</deployedseed>
    <deployedseed>1397273226</deployedseed>
    <deployedseed>1113152399</deployedseed>
    <deployedseed>2029369252</deployedseed>
    <deployedseed>826156328</deployedseed>
    <deployedseed>202647328</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>dispdp(ta1,4)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dispdp(ta2,4)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>dispdp(ta3,4)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>dispdp(ev(ta1+1,simp),4)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dispdp(ev(ta2+1,simp),4)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>dispdp(ev(ta3+1,simp),4)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>-ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>-ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>-ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>