<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62108 -->
  <question type="stack">
    <name>
				<text>23.2.5.1 Fourier series for functions of general period</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Fourier series for functions of general period</h4>
This is a straightforward extension of the period \(2\pi\) case that we have already discussed.
Using \(x\) (instead of \(t\)) temporarily as the variable. We have seen that a \(2\pi\) periodic function \(f(x)\)
has a Fourier series
\[ f(x) = \dfrac{a_0}{2} + \sum _{n = 1}^\infty ( a_n \cos nx + b_n \sin nx) \] 
with
\[ a_n = \dfrac{1}{\pi} \displaystyle{ \int _{-\pi} ^\pi } f(x) \cos nx \hspace{0.15cm} dt \hspace{2cm} n = 0, 1, 2, \dots\]
\[ b_n = \dfrac{1}{\pi} \displaystyle{ \int _{-\pi} ^\pi } f(x) \sin nx \hspace{0.15cm} dt \hspace{2cm} n = 1, 2, \dots\]
Suppose we now change the variable to \(t\) where \(x = \dfrac{2\pi}{T}t\).
<p></p>
Thus \( x = \pi \) corresponds to \( t = T/2 \) and \( x = - \pi \) corresponds  to \( t = -T/2\).
<p></p>
Hence regarded as a function of \(t\), we have a function with period \(T\).
<p></p>
Making the substitution \(x = \dfrac{2\pi}{T}t\), and hence \(dx = \dfrac{2\pi}{T} dt\), in the expressions for \(a_n\) and \(b_n\) we obtain
\[
\begin{align}
a_n &amp;= \dfrac{2}{T} \int _{-\frac{T}{2}} ^{\frac{T}{2}} f(t) \cos \left( \dfrac{2n\pi t}{T} \right) dt \hspace{2cm} n = 0, 1, 2, \dots \\
b_n &amp;= \dfrac{2}{T} \int _{-\frac{T}{2}} ^{\frac{T}{2}} f(t) \sin \left( \dfrac{2n\pi t}{T} \right) dt \hspace{2cm} n = 1, 2, \dots
\end{align}
\]
These integrals give the Fourier coefficients for a function of period \(T\) whose Fourier series is
\[ f(t) = \dfrac{a_0}{2} + \sum _{n = 1} ^\infty \left\{ a_n\cos \left( \dfrac{2n\pi t}{T}\right) + b_n \sin\left( \dfrac{2n\pi t}{T}\right) \right\} \]
Various other notations are commonly used in this case e.g. it is sometimes convenient to write
the period \(T = 2l\). (This is particularly useful when Fourier series arise in the solution of partial
differential equations.) Another alternative is to use the angular frequency \(\omega\) and put \(T = 2\pi/\omega\).]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200816154045+muFYAs</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220518094650+GgoUfy</text></tag>
    </tags>
    
</question>
</quiz>