<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26055 -->
  <question type="stack">
    <name>
				<text>4.5.4.3 Exercise 3 Max current (NOT IN QUIZ)</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>The current in an a.c. circuit is</p>
 \[ i(t)={@ta9@} \sin ({@na32@}t)\ \mathrm{amp},\]
<p>where \(t\) is measured in seconds.</p>
<p>What is the maximum current and at what times does it occur?
</p><p>Please write "pi" for the value \(\pi\).</p>
<p>The maximum current is [[input:ans1]] \(\mathrm{amp}\), [[validation:ans1]][[feedback:prt1]]</p>
<p>which occurs at times \( t = \) [[input:ans2]] \(\mathrm{s}\), for \(n=1,2,3 \dots \) [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p dir="ltr" style="text-align: left;"></p><p>The current in an a.c. circuit is</p>\[ i(t)={@ta9@} \sin ({@na32@}t)\ \mathrm{amp}.\]<br><p></p><p dir="ltr" style="text-align: left;">The maximum value of the sine function is 1, therefore the maximum value is \(i(t) = {@ta9@} \, \mathrm{amp}\).</p><p dir="ltr" style="text-align: left;">Then to find the value of \(t\) for which this maximum occurs, find \(\sin ({@na32@}t) = 1\). Firstly \({@na32@}t = {@2*n*%pi + %pi/2@}\) for \( n \in&nbsp;\mathbb{Z}\). Hence, \[t =&nbsp;{@(1/na32)*(2*n*%pi + %pi/2)@} = {@(4*n+1)*(pi/(2*na32))@} = {@ta10@}.\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021072200</text>
		</stackversion>
    <questionvariables>
				<text>ta9: rand(30)+2;
na32: pi*(rand(10)+2); 
ta10: ev(pi/(2*na32)+ (2*n*pi)/(na32),simp);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\[ i(t)={@ta9@} \sin ({@na32@})t  \text{ amp } \]
\[ I={@ta9@}, \quad t={@ta10@}\]</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta9</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta10</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta9</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta10</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200713133540+c2ZHds</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210909155837+Fmt3kN</text></tag>
    </tags>
    <deployedseed>265448111</deployedseed>
    <deployedseed>1358361576</deployedseed>
    <deployedseed>1400667838</deployedseed>
    <deployedseed>1487983630</deployedseed>
    <deployedseed>1658819146</deployedseed>
    <deployedseed>2067666000</deployedseed>
    <deployedseed>1428494969</deployedseed>
    <deployedseed>91052921</deployedseed>
    <deployedseed>976300055</deployedseed>
    <deployedseed>2097357742</deployedseed>
    <deployedseed>762156464</deployedseed>
    <deployedseed>981743553</deployedseed>
    <deployedseed>276763732</deployedseed>
    <deployedseed>1459586073</deployedseed>
    <deployedseed>1381950444</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta9</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta10</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(ta9+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta10+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>