<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24297 -->
  <question type="stack">
    <name>
				<text>5.1.1.13 DELETE</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>During one winter, the roads in a rural area were completely free from snow when it started snowing at midnight. It snowed heavily all night and day. By {@a@} a.m. it was {@b@} cm deep. To save money the local authorities wait until the snow is {@c@} cm deep before ploughing the snow away from the roads. Forecast how many hours and minutes after midnight the ploughing should start [[input:ans1]] [[feedback:prt1]] [[input:ans2]] [[feedback:prt2]][[validation:ans1]] [[validation:ans2]], enter the number of assumptions you have made [[input:ans3]][[validation:ans3]][[feedback:prt3]] to check your assumptions.
 <ol>
</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>If the depth of snow is represented by <em>D</em> cm and the time by <em>t</em> hours after midnight then a linear model for the depth as a function of time may be written:</p>
<p><em>D</em> = <em>at</em> + <em>b</em></p>
<p>where <em>a</em> and <em>b</em> are constants to be found from the other information in the problem or from assumptions. As there was no snow at midnight </p>
<p> 0 = 0 + <em>b</em>    or     <em>b</em>= 0</p>
<p>It is stated that <em>D</em> = {@b@} when <em>t</em> = {@a@}, i.e.</p>
{@b@} = {@a@}<em>a</em>     or <em>a</em> = {@b@}/{@a@}</p>
<p>So the equation for estimating the depth of snow at different times is:</p>
<p><em>D</em> = {@b@}<em>t</em>/{@a@}</p>
<p>The model will predict a level of {@c@}cm at a time <em>t</em> given by</p>
<p>{@c@} ={@b@} <em>t</em> / {@a@}     or    <em>t</em> = {@c@}*{@a@}/{@b@} = {@d@}</p>
<p>i.e.  The snow plough will be needed {@ta1@} hours and {@ta2@} after midnight.</p>
<p>Assumptions:</p>
<p>Assume snow build up is steady</p>
<p>No drifts</p>
<p> Assume that the snow is constant</p>
<p> Assume change in the time of the day does not affect the situation</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021120900</text>
		</stackversion>
    <questionvariables>
				<text>a: rand([8,9,10,11]);
b: rand([15,16,17,18,19,20,21,22]);
c: rand([30,31,32,33,34,35,36,37]);
d: a+(a.(c-b)/b);
ta1: floor(d)
g: (d-ta1)*60
ta2: round(g)
ta3: (2 or 3 or 4 or 5)
</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>a</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta3</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>0.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200726192646+ryBc8t</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220224121751+xkT1Nl</text></tag>
    </tags>
    
</question>
</quiz>