<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26052 -->
  <question type="stack">
    <name>
				<text>4.5.4.9 Exercise Weight on spring (NOT IN QUIZ)</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>The motion of a weight on a 
spring is given by \[\displaystyle{y= {@na4a@}\cos({@na4c@}t)- 
{@na4b@}\sin({@na4c@}t)}\]</p>
<p>Obtain \(C\) and \(\alpha\) such that \( \displaystyle{y=C\sin ({@na4c@}t+\alpha)}\).</p>
<p> \( y =\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>Using the formulae gives:</p>
\[ C=\sqrt(({@na4a@})^2+ ({@na4b@})^2)={@ta4c@} \]
\[\cos\alpha=-\dfrac{{@na4b@}}{{@ta4c@}} \qquad \sin 
\alpha = -\dfrac{{@na4a@}}{{@ta4c@}}\]
<p> We see that \(\alpha\) is in the second 
quadrant and \(\alpha={@ta4a@}\).</p>
<p>Therefore,</p>
\[y={@ta4f@} \]]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021071900</text>
		</stackversion>
    <questionvariables>
				<text>na4a: 2/(rand(4)+1);
na4b: 2/(rand(5)+2);
na4c: rand(7)+1;
ta4c: sqrt(na4a^2+ na4b^2);
ta4a: decimalplaces(float(atan(na4a/na4b)),4);
ta4f: ta4c*sin(na4c*t+ta4a);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\[y= {@ta4f@}\]</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta4f</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta4f</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200713115736+M3RVA8</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210723180019+Tl2jXE</text></tag>
    </tags>
    <deployedseed>530070191</deployedseed>
    <deployedseed>95658768</deployedseed>
    <deployedseed>738232789</deployedseed>
    <deployedseed>1550032515</deployedseed>
    <deployedseed>1997197952</deployedseed>
    <deployedseed>2047153847</deployedseed>
    <deployedseed>979133114</deployedseed>
    <deployedseed>773680377</deployedseed>
    <deployedseed>299169984</deployedseed>
    <deployedseed>167414871</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta4f</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>sin(7*t+%pi/4)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty></expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>