<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26727 -->
  <question type="stack">
    <name>
				<text>10.1.1.8 What is a Complex Number?</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<div class="HELM_keypoint">
<h4>Key Point </h4>
<p>The <span><strong>modulus</strong></span> of a complex number <span class="math inline">\(z=a+i b\)</span> is denoted by <span class="math inline">\(|z|\)</span> and is defined by</p>
<p><span class="math inline">\(\qquad |z|=\sqrt{a^2+b^2}\)</span></p>
<p>so that the modulus is always a non-negative real number.</p>
</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200625155700+Ut8Nzc</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210706164612+fj15Zi</text></tag>
    </tags>
    
</question>
</quiz>