<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26741 -->
  <question type="stack">
    <name>
				<text>10.2.2.5 Argument</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Assume \(k&gt;0\) is a real number.&nbsp; For each complex number specify the argument.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200629215914+G7Stzl</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200914161445+LzAWtj</text></tag>
    </tags>
    
</question>
</quiz>