<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26522 -->
  <question type="stack">
    <name>
				<text>6.2.4.5 (not in quiz) Express exp in terms of sinh, cosh (exercise 2(c))</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise&nbsp;</p>
<p>Express the following in terms of \(\sinh({@aa@}{@xx@})\) and \(\cosh(-{@aa@}{@xx@})\).</p>
<p>\(&nbsp; {@uu@} e^{-{@aa@}{@xx@}} - {@vv@} e^{{@aa@}{@xx@}} = \) [[input:ans1]] [[validation:ans1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>Using \(e^{{@aa@}{@xx@}} = \cosh {@aa@}{@xx@} + \sinh {@aa@}{@xx@} \) and \(e^{-{@aa@}{@xx@}} = \cosh {@aa@}{@xx@} - \sinh {@aa@}{@xx@}\), we get</p>
<p>\[ \begin{align*} &amp; {@uu@} e^{-{@aa@}{@xx@}} - {@vv@} e^{{@aa@}{@xx@}} \\
&amp; = {@uu@}( \cosh {@aa@}{@xx@} - \sinh {@aa@}{@xx@} ) - {@vv@} (\cosh {@aa@}{@xx@} + \sinh {@aa@}{@xx@})  \ \\
&amp; = {@ta@}
\end{align*}\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020070100</text>
		</stackversion>
    <questionvariables>
				<text>aa: rand(6)+2;
uu: rand(6)+2;
vv: rand(6)+2;
xx: rand([x,y,z,t]);
exp: cosh(aa*xx) + sinh(aa*xx);
expn: cosh(aa*xx) - sinh(aa*xx);
sum: -vv*exp+uu*expn;
ta: expand(sum);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text><![CDATA[[[feedback:prt1]]]]></text>
		</specificfeedback>
    <questionnote>
				<text>\( {@uu@} e^{-{@aa@}{@xx@}} - {@vv@} e^{{@aa@}{@xx@}}  = {@ta@} \)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>25</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>e,[[BASIC-ALGEBRA]]</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200625191052+rb1UkK</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200716111557+jsfGK4</text></tag>
    </tags>
    <deployedseed>1060509446</deployedseed>
    <deployedseed>122038673</deployedseed>
    <deployedseed>1407754912</deployedseed>
    <deployedseed>1626150642</deployedseed>
    <deployedseed>19620855</deployedseed>
    <deployedseed>1363038024</deployedseed>
    <deployedseed>1748525147</deployedseed>
    <deployedseed>560113779</deployedseed>
    <deployedseed>621447793</deployedseed>
    <deployedseed>1671015681</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(ta+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>