<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 27179 -->
  <question type="stack">
    <name>
				<text>35.1.3.1 The intersection of sets</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h3>The intersection and union of sets</h3><h4>
Intersection</h4><p>Given two sets, \(A\) and \(B\), the <b>intersection</b> of \(A\) and \(B\) is a set which contains elements that are common both to \(A\) and \(B\). We write \(A \cap B\) to denote the intersection of 
\(A\) and \(B\).  Mathematically we write this as:<br></p>
<div class="HELM_keypoint">
<h4>Key Point</h4>
<p>\(A \cap B = \{x:x \in A \hbox{ and }x \in B\}\).</p>
</div><br>
<p>This says that the intersection contains all the elements \(x\) such that \(x\) belongs to \(A\) and also \(x\) belongs to \(B\).  Note that \(A \cap B\) and \(B \cap A\) are identical. The intersection of two sets can be represented by a Venn diagram as shown in the following figure.</p><p><img src="@@PLUGINFILE@@/35-1-fig-3-intersection.svg" alt="" width="235" height="173" class="img-responsive atto_image_button_text-bottom"><br></p>
Figure: The overlapping area represents \(A \cap B\).]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200701002701+JtFWWt</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200727224140+ddB4GX</text></tag>
    </tags>
    
</question>
</quiz>