<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26737 -->
  <question type="stack">
    <name>
				<text>10.2.3.3 Key Point-Multiplication and division</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>We see that in calculating the quotient that the moduli <span class="math inline">\(r\)</span> and <span class="math inline">\(t\)</span> <span><strong>divide</strong></span> whilst the arguments \(\text{arg}(z)=\theta\) and \(\text{arg}(w)=\phi\)<span class="math inline"></span> <span><strong>subtract</strong></span>.</p>
<div class="HELM_keypoint">
<h4>Key Point</h4>
<p>If <span class="math inline">\(z=r(\cos\theta +i\sin\theta )\)</span> and <span class="math inline">\(w=t(\cos\phi
+i\sin\phi )\)</span> then <span class="math display">\[zw=rt(\cos (\theta +\phi)+i\sin (\theta + \phi ))\qquad 
{z\over w}={r\over t}(\cos (\theta -\phi) +i\sin
(\theta -\phi ))\]</span></p>
</div>

<p>We conclude that addition and subtraction are most easily carried out in Cartesian form whereas multiplication and division are most easily carried out in polar form.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200629125022+OmNJC5</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200727124312+UgHkLm</text></tag>
    </tags>
    
</question>
</quiz>