<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24580 -->
  <question type="stack">
    <name>
				<text>16.3.1.3 binomial coefficients</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p><p></p><p>As a real number \(x\) approaches zero we note hiigher powers of&nbsp; \(x\) become even closer to zero and may even fall within a margin of error we can tolerate.&nbsp;</p><p>Assume, for this exercise, that \(x^3\) and higher powers may be ignored in comparison to lower order terms. <br><br>Find a quadratic polynomial approximating \( (1−x)^{{@pp@}} \). Hint:&nbsp;First expand \( (1+x)^{{@pp@}} \) using the binomial series with \(p=\displaystyle{@pp@} \)and with \( x \) replaced by \( (−x) \):<br></p><p>\( (1-x)^{@pp@} \approx \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p><p><strong></strong></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p dir="ltr" style="text-align: left;"></p>
<p style="text-align: left;"><span style="font-size: 0.9375rem;">The binomial series gives \[ (x-1)^{@pp@} \approx {@ta@}.\]<br></span></p><p style="text-align: left;"><span style="font-size: 0.9375rem;">What we have done in this exercise is to approximate the function \((1-x)^{{@pp@}}\) by the simpler (polynomial) function \(\displaystyle {@ta@}\) which is reasonable provided \(x\) is very small.&nbsp;</span><br></p>
<p></p>

<p dir="ltr" style="text-align: left;"></p>
<p style="text-align: left;">In the case of \((1-x)^{\frac{1}{2}}\)this approximation is well illustrated geometrically by drawing the curves \(y=(1-x)^{\frac{1}{2}}\) and \(y=1-\dfrac{1}{2}x-\dfrac{1}{8}x^2\). The two curves coincide when \(x\) is ‘small’.</p>
<p style="text-align: center;"><img src="@@PLUGINFILE@@/quadratic%20polynomial%20graph.svg" alt="plot of x minus 1 over 8 x squared with dotted line 1 minus x all to the power of a half" class="img-responsive atto_image_button_text-bottom" width="457" height="195"></p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2020120600</text>
		</stackversion>
    <questionvariables>
				<text>pp: 1/(rand(8)+2);
ta: 1-(pp)*x+(pp*(pp-1))/2*x^2</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\[ (1-x)^{{@pp@}}= {@ta@}\]</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
      <node>
        <name>1</name>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>(1-x)^(1/pp)</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
					<text><p dir="ltr" style="text-align: left;">This is just the function and the required polynomial approximation.</p></text>
				</truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+201014124005+uLVzjO</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210610102711+75gShL</text></tag>
    </tags>
    <deployedseed>1041493062</deployedseed>
    <deployedseed>911925760</deployedseed>
    <deployedseed>128511278</deployedseed>
    <deployedseed>1807004749</deployedseed>
    <deployedseed>1393782279</deployedseed>
    <deployedseed>1626669995</deployedseed>
    <deployedseed>1268105003</deployedseed>
    <deployedseed>391115733</deployedseed>
    <deployedseed>541931792</deployedseed>
    <deployedseed>89702201</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>ev(ta+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>(1-x)^pp</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>