<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26742 -->
  <question type="stack">
    <name>
				<text>10.2.3.7 Modulus</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Which of the following are true?</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>Let \(z = r(\cos \theta + i\sin \theta)\). Then \(|z| = r \implies |z|^2 = r^2\). Also we have \(z^2 = r^2 (\cos \theta + i\sin \theta)^2\) which implies that \(|z^2| = r^2\). Therefore \( |z^2| = |z|^2 \).</p><p><br></p>

<p>Let \(z = a + ib\). Then \(zz^* = (a+ib)(a-ib) = a^2 + b^2\). Note that \(|z| = \sqrt{a^2 + b^2}\) which implies that \(|z|^2 = a^2+b^2\). Therefore \( z z^* = |z|^2 \).</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200629220021+Ljb8SE</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210702115912+Mvkqye</text></tag>
    </tags>
    
</question>
</quiz>