<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 24679 -->
  <question type="stack">
    <name>
				<text>19.1.1.3 Initial Condition</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p>In the above question we call <span class="math inline">\(t\)</span> the independent variable and <span class="math inline">\(\theta\)</span> the dependent variable. Since the condition is given at <span class="math inline">\(t=0\)</span> we refer to it as
    an initial condition. For future reference, the solution of the above differential equation which satisfies the initial condition is <span class="math inline">\(\theta=\theta_\text{s}+(\theta_0-\theta_\text{s})\text{e}^{-kt}\)</span></p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:moodle2.maths.ed.ac.uk+210305174232+RlVHDB</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+210315130205+I47XHE</text></tag>
    </tags>
    
</question>
</quiz>