<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26734 -->
  <question type="stack">
    <name>
				<text>10.2.2.2 Key Point -Polar form</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<div class="HELM_keypoint">
    <h4>Key Point</h4>
    <p>If <span class="math inline">\(z=a+i b\), \(a\neq 0\)</span> then <span class="math display">\[z=r(\cos\theta+i\sin\theta)\]</span> in which <span class="math display">\[r=|z|=\sqrt{a^2+b^2}\quad \text{and}\quad \tan(\theta) =\dfrac{b}{a}.\] Note to find \(\theta=\text{arg}(z)\) we need to know which quadrant \(z\) lies by examining the signs of \(a\) and \(b\). The inverse tan function will give an angle in the first or fourth quadrant so we need to add or subtract \(\pi\) if the complex number lies in the second or third quadrant.<br></span></p>
</div>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200629121527+U6l02p</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+220824090819+c3JoCK</text></tag>
    </tags>
    
</question>
</quiz>