<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 62116 -->
  <question type="stack">
    <name>
				<text>23.3 Introduction</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<h4>Introduction</h4>
In this Section we examine how to obtain Fourier series of periodic functions which are either even
or odd. We show that the Fourier series for such functions is considerably easier to obtain as, if the
signal is even only cosines are involved whereas if the signal is odd then only sines are involved. We
also show that if a signal reverses after half a period then the Fourier series will only contain odd
harmonics.]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200825173711+l7tX4n</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200825173834+0vEdGf</text></tag>
    </tags>
    
</question>
</quiz>