<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25720 -->
  <question type="stack">
    <name>
				<text>2.4.2.1 Continuous vs discontinuous discussion *****</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>

<p>Explain the distinction between a continuous and a discontinuous function. Draw a graph showing an example of each type of function.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200618160306+dmCCVO</text></tag>
      <tag><text>version:stack2.maths.ed.ac.uk+200622151923+7QQs6v</text></tag>
    </tags>
    
</question>
</quiz>