<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 26249 -->
  <question type="stack">
    <name>
				<text>7.3.2.6 Compute a 3x3 determinant</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>By expanding along the first row, compute the determinant of the matrix \(A = {@matrixa@}\).</p><p>The determinant \(\Delta\) is: [[input:ans1]][[validation:ans1]][[feedback:prt1]]</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text><![CDATA[<p>For \(A={@matrixa@}\), we have: \begin{align*} a_{11} &amp; = {@matrixa[1,1]@}, &amp; a_{12} &amp; = {@matrixa[1,2]@}, &amp; a_{13} &amp; = {@matrixa[1,3]@}, \\ A_{11} &amp; = \det{{@AA11@}} = {@determinant(AA11)@}, &amp; A_{12} &amp; = \det{{@AA12@}} = {@determinant(AA12)@}, &amp; A_{13} &amp; = \det{{@AA13@}} = {@determinant(AA13)@}. \end{align*}</p>
<p>Therefore \[\Delta = a_{11}A_{11}+a_{12}A_{12}+a_{13}A_{13} = {@ta1@}.\]</p>]]></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
				<text>2021072200</text>
		</stackversion>
    <questionvariables>
				<text>matrixa: matrix([rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)],[rand_with_step(-5,5,1),rand_with_step(-5,5,1),rand_with_step(-5,5,1)]);
matrixa: ev(matrixa,simp);

AA11: matrix([matrixa[2,2],matrixa[2,3]],[matrixa[3,2],matrixa[3,3]]);
AA11: ev(AA11,simp);
AA12: matrix([matrixa[2,1],matrixa[2,3]],[matrixa[3,1],matrixa[3,3]]);
AA12: ev(AA12,simp);
AA13: matrix([matrixa[2,1],matrixa[2,2]],[matrixa[3,1],matrixa[3,2]]);
AA13: ev(AA13,simp);

ta1: ev(determinant(matrixa),simp);</text>
		</questionvariables>
    <specificfeedback format="html">
				<text></text>
		</specificfeedback>
    <questionnote>
				<text>\(A = {@matrixa@}, D = {@ta1@}\)</text>
		</questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
		</prtcorrect>
    <prtpartiallycorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
		</prtpartiallycorrect>
    <prtincorrect format="html">
				<text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
		</prtincorrect>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed>MinorAndCofactor</variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>determinant, minor</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
				<text></text>
			</feedbackvariables>
      <node>
        <name>0</name>
        <answertest>CasEqual</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
					<text></text>
				</truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
					<text></text>
				</falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200725160655+SG3NLc</text></tag>
      <tag><text>version:stack-back1.maths.ed.ac.uk+210910161800+Yfo5du</text></tag>
    </tags>
    <deployedseed>401428587</deployedseed>
    <deployedseed>256470495</deployedseed>
    <deployedseed>1580423024</deployedseed>
    <deployedseed>311223496</deployedseed>
    <deployedseed>717670310</deployedseed>
    <deployedseed>478689502</deployedseed>
    <deployedseed>1948238459</deployedseed>
    <deployedseed>975517033</deployedseed>
    <deployedseed>491816791</deployedseed>
    <deployedseed>67200360</deployedseed>
    <deployedseed>1294055332</deployedseed>
    <deployedseed>1284459475</deployedseed>
    <deployedseed>1969144389</deployedseed>
    <deployedseed>2131917626</deployedseed>
    <deployedseed>691664144</deployedseed>
    <deployedseed>138140317</deployedseed>
    <deployedseed>1729697927</deployedseed>
    <deployedseed>1066017578</deployedseed>
    
    <qtest>
      <testcase>1</testcase>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <testinput>
        <name>ans1</name>
        <value>-ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <testinput>
        <name>ans1</name>
        <value>matrixa[1,1]*determinant(AA11)+matrixa[1,2]*determinant(AA12)+matrixa[1,3]*determinant(AA13)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
</question>
</quiz>