<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 25721 -->
  <question type="stack">
    <name>
				<text>2.4.2.2 Quadratic continuity</text>
		</name>
    <questiontext format="moodle_auto_format">
				<text><![CDATA[<p class="HELM_exercise">Exercise</p>
<p>Study the graphs of the functions \(y=x^2\) and \(y=-x^2\), \(x \in \mathbb{R}\) and fill in the blanks:</p><p>\(y=x^2\) is [[1]] and</p><p>\(y=-x^2\) is [[1]].<br></p><p><br></p><p>Note: after checking the answer you will need to drag the items out of the boxes before allowing you to drag back in for the next attempt.</p>]]></text>
		</questiontext>
    <generalfeedback format="moodle_auto_format">
				<text></text>
		</generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <tags>
      <tag><text>stamp:stack2.maths.ed.ac.uk+200618160812+hj63MU</text></tag>
      <tag><text>version:moodle2.maths.ed.ac.uk+200929092935+TvHLGB</text></tag>
    </tags>
    
</question>
</quiz>